Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.pair
% (98136)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 218
% Passive clauses: 4375
% Generated clauses: 24563
% Final active clauses: 163
% Final passive clauses: 3769
% Input formulas: 68
% Initial clauses: 104
% 
% Duplicate literals: 761
% Trivial inequalities: 295
% Fw subsumption resolutions: 1923
% Fw demodulations: 4300
% Bw demodulations: 51
% 
% Simple tautologies: 302
% Equational tautologies: 52
% Forward subsumptions: 11907
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 499
% Factoring: 1
% Forward superposition: 8752
% Backward superposition: 7493
% Self superposition: 137
% Equality factoring: 84
% Equality resolution: 28
% 
% Split clauses: 286
% Split components: 577
% SAT solver clauses: 866
% SAT solver unit clauses: 21
% SAT solver binary clauses: 185
% 
% Sat splits: 286
% Sat splitting refutations: 21
% 
% TWLsolver clauses: 860
% TWLsolver calls for satisfiability: 156
% 
% Memory used [KB]: 15351
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pair
% (98155)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 233
% Passive clauses: 2013
% Generated clauses: 17113
% Final active clauses: 111
% Final passive clauses: 1029
% Input formulas: 68
% Initial clauses: 104
% 
% Duplicate literals: 64
% Trivial inequalities: 192
% Fw subsumption resolutions: 1031
% Fw demodulations: 2742
% Bw demodulations: 796
% Condensations: 317
% Global subsumptions: 1
% 
% Simple tautologies: 37
% Equational tautologies: 10
% Forward subsumptions: 8191
% Bw demodulations to eq. taut.: 29
% 
% Binary resolution: 762
% Factoring: 2
% Forward superposition: 4664
% Backward superposition: 4104
% Self superposition: 101
% Equality factoring: 6
% Equality resolution: 8
% 
% Split clauses: 125
% Split components: 250
% SAT solver clauses: 2019
% SAT solver unit clauses: 134
% SAT solver binary clauses: 64
% SAT solver learnt clauses: 8
% SAT solver learnt literals: 18
% 
% Sat splits: 125
% Sat splitting refutations: 24
% 
% TWLsolver clauses: 2009
% TWLsolver calls for satisfiability: 553
% 
% Memory used [KB]: 17270
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 137 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pair
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pair
% SZS output start Proof for thm.pair
fof(f9227,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f260,f368,f367,f259,f370,f366,f611,f616,f1329,f617,f1337,f1338,f1340,f1342,f1346,f1348,f1350,f1352,f1356,f1358,f1360,f1362,f1370,f1372,f1433,f1434,f1435,f1436,f1437,f1438,f1439,f1440,f1441,f1442,f1443,f1444,f2709,f2711,f2710,f3522,f3538,f3536,f4241,f4245,f4243,f4247,f4251,f4249,f4495,f4499,f4501,f4503,f4497,f4518,f4520,f8106,f8367,f8369,f8370,f8371,f8372,f8373,f8374,f8459,f9205])).
fof(f9205,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f9201,f263])).
fof(f263,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pair.ELIM_PEXISTS_EVAL',aHL_FALSITY)).
fof(f9201,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f9183,f8497])).
fof(f8497,plain,(
  p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK6('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | $spl0),
  inference(forward_demodulation,[],[f8478,f335])).
fof(f335,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X11,X7,X6] : ! [X12,X13] : s(X11,'const.pair.UNCURRY_2'(s(fun(X7,fun(X6,X11)),X12),s('type.pair.prod'(X7,X6),X13))) = s(X11,happ(s(fun('type.pair.prod'(X7,X6),X11),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,X11)),X12))),s('type.pair.prod'(X7,X6),X13)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','ARITY_1_2_const.pair.UNCURRY')).
fof(f8478,plain,(
  p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),sK6('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f366,f288])).
fof(f288,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f215])).
fof(f215,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f214])).
fof(f214,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X7] : ! [X15] : (p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X15)))) <=> ? [X9] : p(s(bool,happ(s(fun(X7,bool),X15),s(X7,X9)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.bool.EXISTS_THM')).
fof(f9183,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f381,f9139,f468])).
fof(f468,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f459,f263])).
fof(f459,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f296,f264])).
fof(f264,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pair.ELIM_PEXISTS_EVAL',aHL_BOOL_CASES)).
fof(f296,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X22,X23] : ((p(s(bool,X22)) => p(s(bool,X23))) => ((p(s(bool,X23)) => p(s(bool,X22))) => s(bool,X22) = s(bool,X23)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.bool.IMP_ANTISYM_AX')).
fof(f9139,plain,(
  ( ! [X7] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),X7))))) ) | $spl2),
  inference(superposition,[],[f8499,f712])).
fof(f712,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK28(X4,X2,X1)))),s(X2,sK29(X4,X2,X1))))) )),
  inference(backward_demodulation,[],[f702,f709])).
fof(f709,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK28(X4,X2,X1)))),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X1,X2),X4)))))) )),
  inference(backward_demodulation,[],[f701,f337])).
fof(f337,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X2),X4))))),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X1,X2),X4)))))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X2),X4))))),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X1,X2),X4)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X2),X4))))),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X1,X2),X4)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11,X7,X6] : ! [X15,X5] : s(X11,'const.pair.UNCURRY_2'(s(fun(X7,fun(X6,X11)),X15),s('type.pair.prod'(X7,X6),X5))) = s(X11,happ(s(fun(X6,X11),happ(s(fun(X7,fun(X6,X11)),X15),s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X6),X5))))),s(X6,'const.pair.SND_1'(s('type.pair.prod'(X7,X6),X5)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.ELIM_UNCURRY')).
fof(f701,plain,(
  ( ! [X2,X0,X1] : (s(X0,sK28(X2,X1,X0)) = s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2)))) )),
  inference(superposition,[],[f323,f321])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK28(X2,X1,X0)),s(X1,sK29(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK28(X2,X1,X0)),s(X1,sK29(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK28,sK29])],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ! [X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X7,X6] : ! [X9] : ? [X25,X26] : s('type.pair.prod'(X7,X6),X9) = s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X25),s(X6,X26)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.ABS_PAIR_THM')).
fof(f323,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6,X7] : ! [X9,X14] : s(X7,X9) = s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X9),s(X6,X14)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.FST')).
fof(f702,plain,(
  ( ! [X4,X5,X3] : (s(X4,sK29(X5,X4,X3)) = s(X4,'const.pair.SND_1'(s('type.pair.prod'(X3,X4),X5)))) )),
  inference(superposition,[],[f322,f321])).
fof(f322,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X7,X6] : ! [X9,X14] : s(X6,X14) = s(X6,'const.pair.SND_1'(s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X9),s(X6,X14)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.SND')).
fof(f8499,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0))),s(sK0,X1))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f368,f289])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f216])).
fof(f381,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f263,f265])).
fof(f265,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.bool.EQ_CLAUSES|split|1')).
fof(f8459,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f8440,f8375])).
fof(f8375,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0))),s(sK0,X1))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f367,f319])).
fof(f319,plain,(
  ( ! [X6,X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ((! [X5,X6] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))))) & (~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK26(X2,X1,X0)))),s(X1,sK27(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26,sK27])],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : ((! [X5,X6] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))))) & (~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))) | ? [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))))),
  inference(rectify,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : ((! [X3,X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))))) & (~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))) | ? [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))))),
  inference(nnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : (? [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) <=> p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : ! [X2] : (? [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) <=> p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7,X6] : ! [X8] : (? [X9,X14] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,X9))),s(X6,X14)))) <=> p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X7,X6),bool),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,bool)),X8)))))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.PEXISTS_THM')).
fof(f8440,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3))),s(sK0,sK6(happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)),sK0))))) | $spl4),
  inference(unit_resulting_resolution,[],[f370,f288])).
fof(f8374,plain,(
  $false | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f367,f4657])).
fof(f4657,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl28),
  inference(subsumption_resolution,[],[f4625,f261])).
fof(f261,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f23])).
fof(f23,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.bool.TRUTH')).
fof(f4625,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl28),
  inference(superposition,[],[f283,f4547])).
fof(f4547,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | $spl28),
  inference(unit_resulting_resolution,[],[f261,f4497,f299])).
fof(f299,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f195])).
fof(f283,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X4) = s(X0,X5)))) & (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X1,X0))))) & s(X0,sK4(X1,X0)) != s(X0,sK5(X1,X0))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5])],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X4) = s(X0,X5)))) & (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s(X0,X2) = s(X0,X3))))),
  inference(ennf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X7] : ! [X5] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X7,bool),X5)))) <=> (p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X5)))) & ! [X9,X14] : ((p(s(bool,happ(s(fun(X7,bool),X5),s(X7,X9)))) & p(s(bool,happ(s(fun(X7,bool),X5),s(X7,X14))))) => s(X7,X9) = s(X7,X14))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.bool.EXISTS_UNIQUE_DEF')).
fof(f8373,plain,(
  $false | ($spl1 | $spl28)),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f266,f265,f270,f269,f279,f280,f287,f286,f285,f284,f283,f289,f288,f299,f296,f301,f300,f303,f302,f307,f308,f310,f309,f312,f311,f316,f315,f318,f317,f321,f322,f323,f325,f332,f333,f334,f335,f336,f337,f340,f339,f342,f341,f343,f344,f345,f346,f347,f348,f380,f350,f351,f352,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f382,f381,f391,f397,f442,f468,f509,f546,f575,f576,f595,f611,f655,f661,f699,f700,f701,f709,f710,f702,f712,f718,f719,f720,f825,f826,f827,f828,f829,f830,f831,f803,f871,f872,f912,f913,f914,f915,f916,f891,f918,f920,f934,f935,f936,f937,f938,f930,f993,f998,f999,f1003,f1051,f1030,f1056,f1196,f1225,f1270,f1272,f1273,f1274,f1275,f1276,f1277,f1233,f1278,f1279,f1298,f1299,f1300,f1326,f1329,f326,f379,f378,f377,f376,f375,f1091,f939,f306,f305,f304,f973,f1127,f1295,f841,f1209,f331,f330,f329,f328,f327,f320,f371,f313,f372,f314,f319,f373,f374,f659,f660,f658,f968,f1122,f1265,f570,f823,f571,f569,f292,f1479,f1480,f1481,f1482,f1483,f1484,f1453,f1486,f1488,f1489,f1516,f1517,f1518,f1519,f1520,f1521,f1512,f1531,f1553,f1561,f1563,f1605,f1607,f1669,f1737,f1738,f1920,f1828,f1837,f1839,f1848,f1857,f1859,f1942,f1945,f2278,f2041,f2049,f2050,f2051,f2052,f2053,f2283,f2062,f2073,f2075,f2308,f2325,f2326,f2327,f2328,f2329,f2344,f2409,f2410,f2416,f2417,f2357,f2423,f2424,f2441,f2442,f2384,f2448,f2449,f2403,f2479,f2481,f2464,f2463,f2503,f2505,f2488,f2487,f2601,f2604,f2607,f2615,f2618,f2621,f2648,f2654,f2663,f2664,f2549,f2670,f2671,f2688,f2689,f2576,f2695,f2696,f2595,f2709,f2713,f2733,f2734,f2736,f2956,f3146,f3158,f3160,f3271,f3260,f3339,f3313,f3341,f3410,f3352,f3418,f3471,f3460,f3523,f3524,f3525,f3526,f3527,f3528,f3529,f3530,f3522,f3638,f3639,f3640,f3641,f3642,f3643,f3633,f4014,f4063,f4065,f4141,f4080,f4149,f4230,f4232,f4234,f4236,f4238,f4240,f3517,f2575,f2548,f2383,f2356,f1944,f1244,f854,f849,f811,f4416,f4418,f4420,f4422,f4424,f4426,f367,f4497,f4552,f4551,f4554,f4547,f4619,f4620,f4722,f4723,f4753,f4838,f4839,f4851,f5026,f5027,f5028,f5029,f5030,f5031,f5128,f5129,f5130,f5131,f5132,f5133,f5249,f5327,f5328,f5329,f5330,f5257,f5258,f5337,f5277,f5278,f5280,f5281,f5351,f5353,f5354,f5302,f5357,f5358,f5360,f5361,f5362,f5363,f5470,f5471,f5472,f5473,f5493,f5494,f5495,f5496,f5497,f5498,f5513,f5514,f5515,f5516,f5571,f5573,f5574,f5523,f5577,f5578,f5580,f5581,f5582,f5583,f5636,f5637,f5638,f5639,f5640,f5641,f5642,f5643,f5632,f5633,f5728,f5729,f5730,f5731,f5732,f5733,f5734,f5735,f5724,f5725,f5840,f5841,f5842,f5843,f5844,f5845,f5846,f5847,f5835,f5848,f5838,f6055,f6172,f6173,f6174,f6202,f6205,f6274,f6273,f6275,f6276,f6277,f6278,f6279,f6280,f6281,f6502,f6335,f6336,f6347,f6350,f6351,f6352,f6353,f6354,f6359,f6360,f6361,f6362,f6363,f6364,f6365,f6366,f6367,f6368,f6370,f6371,f6372,f6373,f6374,f6376,f6377,f6380,f6381,f6382,f6383,f6384,f6385,f6387,f6388,f6389,f6390,f6391,f6392,f6393,f6394,f6395,f6396,f6397,f6398,f6399,f6400,f6401,f6402,f6403,f6404,f6411,f6412,f6432,f6435,f6436,f6437,f6438,f6439,f6444,f6445,f6447,f6448,f6449,f6450,f6452,f6453,f6454,f6455,f6456,f6458,f6459,f6462,f6463,f6464,f6465,f6467,f6468,f6469,f6470,f6471,f6472,f6473,f6474,f6475,f6476,f6477,f6478,f6479,f6480,f6481,f6482,f6483,f6484,f6491,f6492,f6570,f6572,f6573,f6574,f6575,f6562,f6563,f6576,f6577,f6578,f6579,f6739,f6740,f6741,f6742,f6743,f6660,f6661,f6744,f6745,f6746,f6747,f6697,f6698,f6699,f6700,f6701,f6733,f6736,f6737,f6984,f6987,f6990,f7008,f7010,f7012,f7027,f7029,f7031,f7062,f7064,f7066,f7092,f6903,f6904,f6909,f6929,f6930,f7235,f7234,f7236,f7237,f7238,f7239,f7240,f7241,f7242,f7327,f7329,f7340,f7343,f7344,f7345,f7346,f7347,f7348,f7356,f7357,f7359,f7360,f7361,f7362,f7363,f7364,f7365,f7366,f7367,f7368,f7369,f7370,f7371,f7372,f7373,f7374,f7375,f7376,f7377,f7378,f7387,f7388,f7398,f7401,f7402,f7403,f7404,f7405,f7406,f7429,f7430,f7431,f7432,f7433,f7434,f7435,f7436,f7437,f7438,f7439,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7448,f7457,f7458,f7563,f7562,f7877,f7565,f7566,f7567,f7621,f7622,f7623,f7624,f7625,f7626,f7627,f7666,f7667,f7668,f7669,f7671,f7673,f7684,f7687,f7688,f7689,f7690,f7691,f7692,f7702,f7703,f7705,f7706,f7707,f7708,f7710,f7711,f7712,f7713,f7714,f7716,f7717,f7720,f7721,f7722,f7723,f7724,f7725,f7727,f7728,f7729,f7730,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7739,f7740,f7741,f7742,f7743,f7744,f7745,f7746,f7756,f7757,f7771,f7782,f7785,f7786,f7787,f7788,f7789,f7790,f7800,f7801,f7802,f7803,f7804,f7805,f7806,f7807,f7809,f7810,f7811,f7812,f7813,f7814,f7818,f7819,f7820,f7821,f7822,f7824,f7825,f7828,f7829,f7830,f7831,f7832,f7833,f7835,f7836,f7837,f7838,f7839,f7840,f7841,f7842,f7843,f7844,f7845,f7846,f7847,f7848,f7849,f7850,f7851,f7852,f7853,f7854,f7864,f7865,f7981,f7982,f7983,f7984,f7985,f7986,f7987,f7988,f7989,f7992,f7995,f8025,f8026,f8027,f8028,f8029,f8031,f8032,f8033,f8034,f8035,f8052,f8053,f8054,f8055,f8129,f8131,f8132,f8063,f8135,f8136,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8147,f8148,f4657,f5953,f2345,f2347])).
fof(f2347,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))))),
  inference(unit_resulting_resolution,[],[f381,f330])).
fof(f2345,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f381,f330])).
fof(f5953,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),X1)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),X0)))) | s('type.pair.prod'(sK1,sK0),X0) = s('type.pair.prod'(sK1,sK0),X1)) ) | $spl28),
  inference(forward_demodulation,[],[f5952,f335])).
fof(f5952,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),X0)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),X1)))) | s('type.pair.prod'(sK1,sK0),X0) = s('type.pair.prod'(sK1,sK0),X1)) ) | $spl28),
  inference(forward_demodulation,[],[f5951,f335])).
fof(f5951,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),X0)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),X1)))) | s('type.pair.prod'(sK1,sK0),X0) = s('type.pair.prod'(sK1,sK0),X1)) ) | $spl28),
  inference(subsumption_resolution,[],[f5929,f261])).
fof(f5929,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),X0)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),X1)))) | s('type.pair.prod'(sK1,sK0),X0) = s('type.pair.prod'(sK1,sK0),X1)) ) | $spl28),
  inference(superposition,[],[f284,f4547])).
fof(f8148,plain,(
  ( ! [X584,X586,X582,X585,X581,X583] : (s(X581,sK28('const.pair.UNCURRY_2'(s(fun(X584,fun(X583,'type.pair.prod'(X581,X582))),X585),s('type.pair.prod'(X584,X583),X586)),X582,X581)) = s(X581,sK28(happ(s(fun(X583,'type.pair.prod'(X581,X582)),happ(s(fun(X584,fun(X583,'type.pair.prod'(X581,X582))),X585),s(X584,sK28(X586,X583,X584)))),s(X583,sK29(X586,X583,X584))),X582,X581))) )),
  inference(forward_demodulation,[],[f8088,f701])).
fof(f8088,plain,(
  ( ! [X584,X586,X582,X585,X581,X583] : (s(X581,'const.pair.FST_1'(s('type.pair.prod'(X581,X582),'const.pair.UNCURRY_2'(s(fun(X584,fun(X583,'type.pair.prod'(X581,X582))),X585),s('type.pair.prod'(X584,X583),X586))))) = s(X581,sK28(happ(s(fun(X583,'type.pair.prod'(X581,X582)),happ(s(fun(X584,fun(X583,'type.pair.prod'(X581,X582))),X585),s(X584,sK28(X586,X583,X584)))),s(X583,sK29(X586,X583,X584))),X582,X581))) )),
  inference(superposition,[],[f701,f712])).
fof(f8147,plain,(
  ( ! [X575,X580,X576,X578,X577,X579] : (s(X576,sK29('const.pair.UNCURRY_2'(s(fun(X578,fun(X577,'type.pair.prod'(X575,X576))),X579),s('type.pair.prod'(X578,X577),X580)),X576,X575)) = s(X576,sK29(happ(s(fun(X577,'type.pair.prod'(X575,X576)),happ(s(fun(X578,fun(X577,'type.pair.prod'(X575,X576))),X579),s(X578,sK28(X580,X577,X578)))),s(X577,sK29(X580,X577,X578))),X576,X575))) )),
  inference(forward_demodulation,[],[f8087,f702])).
fof(f8087,plain,(
  ( ! [X575,X580,X576,X578,X577,X579] : (s(X576,'const.pair.SND_1'(s('type.pair.prod'(X575,X576),'const.pair.UNCURRY_2'(s(fun(X578,fun(X577,'type.pair.prod'(X575,X576))),X579),s('type.pair.prod'(X578,X577),X580))))) = s(X576,sK29(happ(s(fun(X577,'type.pair.prod'(X575,X576)),happ(s(fun(X578,fun(X577,'type.pair.prod'(X575,X576))),X579),s(X578,sK28(X580,X577,X578)))),s(X577,sK29(X580,X577,X578))),X576,X575))) )),
  inference(superposition,[],[f702,f712])).
fof(f8144,plain,(
  ( ! [X476,X472,X474,X469,X471,X475,X473,X470] : (p(s(bool,happ(s(fun(X470,bool),happ(s(fun(X469,fun(X470,bool)),'const.pair.UNCURRY_2'(s(fun(X472,fun(X471,fun(X469,fun(X470,bool)))),X473),s('type.pair.prod'(X472,X471),X474))),s(X469,X475))),s(X470,X476)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X469,fun(X470,bool)),'const.pair.UNCURRY_2'(s(fun(X472,fun(X471,fun(X469,fun(X470,bool)))),X473),s('type.pair.prod'(X472,X471),X474))),s('type.pair.prod'(X469,X470),sK15(happ(s(fun(X471,fun(X469,fun(X470,bool))),happ(s(fun(X472,fun(X471,fun(X469,fun(X470,bool)))),X473),s(X472,sK28(X474,X471,X472)))),s(X471,sK29(X474,X471,X472))),X470,X469)))))) )),
  inference(forward_demodulation,[],[f8074,f712])).
fof(f8074,plain,(
  ( ! [X476,X472,X474,X469,X471,X475,X473,X470] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X469,fun(X470,bool)),'const.pair.UNCURRY_2'(s(fun(X472,fun(X471,fun(X469,fun(X470,bool)))),X473),s('type.pair.prod'(X472,X471),X474))),s('type.pair.prod'(X469,X470),sK15(happ(s(fun(X471,fun(X469,fun(X470,bool))),happ(s(fun(X472,fun(X471,fun(X469,fun(X470,bool)))),X473),s(X472,sK28(X474,X471,X472)))),s(X471,sK29(X474,X471,X472))),X470,X469))))) | p(s(bool,happ(s(fun(X470,bool),happ(s(fun(X469,fun(X470,bool)),happ(s(fun(X471,fun(X469,fun(X470,bool))),happ(s(fun(X472,fun(X471,fun(X469,fun(X470,bool)))),X473),s(X472,sK28(X474,X471,X472)))),s(X471,sK29(X474,X471,X472)))),s(X469,X475))),s(X470,X476))))) )),
  inference(superposition,[],[f373,f712])).
fof(f8143,plain,(
  ( ! [X466,X461,X463,X465,X467,X462,X468,X464] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X461,fun(X462,bool)),'const.pair.UNCURRY_2'(s(fun(X464,fun(X463,fun(X461,fun(X462,bool)))),X465),s('type.pair.prod'(X464,X463),X466))),s('type.pair.prod'(X461,X462),sK25(happ(s(fun(X463,fun(X461,fun(X462,bool))),happ(s(fun(X464,fun(X463,fun(X461,fun(X462,bool)))),X465),s(X464,sK28(X466,X463,X464)))),s(X463,sK29(X466,X463,X464))),X462,X461))))) | ~p(s(bool,happ(s(fun(X462,bool),happ(s(fun(X461,fun(X462,bool)),'const.pair.UNCURRY_2'(s(fun(X464,fun(X463,fun(X461,fun(X462,bool)))),X465),s('type.pair.prod'(X464,X463),X466))),s(X461,X467))),s(X462,X468))))) )),
  inference(forward_demodulation,[],[f8073,f712])).
fof(f8073,plain,(
  ( ! [X466,X461,X463,X465,X467,X462,X468,X464] : (~p(s(bool,happ(s(fun(X462,bool),happ(s(fun(X461,fun(X462,bool)),'const.pair.UNCURRY_2'(s(fun(X464,fun(X463,fun(X461,fun(X462,bool)))),X465),s('type.pair.prod'(X464,X463),X466))),s(X461,X467))),s(X462,X468)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X461,fun(X462,bool)),happ(s(fun(X463,fun(X461,fun(X462,bool))),happ(s(fun(X464,fun(X463,fun(X461,fun(X462,bool)))),X465),s(X464,sK28(X466,X463,X464)))),s(X463,sK29(X466,X463,X464)))),s('type.pair.prod'(X461,X462),sK25(happ(s(fun(X463,fun(X461,fun(X462,bool))),happ(s(fun(X464,fun(X463,fun(X461,fun(X462,bool)))),X465),s(X464,sK28(X466,X463,X464)))),s(X463,sK29(X466,X463,X464))),X462,X461)))))) )),
  inference(superposition,[],[f374,f712])).
fof(f8142,plain,(
  ( ! [X453,X455,X457,X459,X454,X460,X456,X458] : (p(s(bool,happ(s(fun('type.pair.prod'(X453,X454),bool),'const.pair.UNCURRY_2'(s(fun(X456,fun(X455,fun('type.pair.prod'(X453,X454),bool))),X457),s('type.pair.prod'(X456,X455),X458))),s('type.pair.prod'(X453,X454),sK22(happ(s(fun(X455,fun('type.pair.prod'(X453,X454),bool)),happ(s(fun(X456,fun(X455,fun('type.pair.prod'(X453,X454),bool))),X457),s(X456,sK28(X458,X455,X456)))),s(X455,sK29(X458,X455,X456))),X454,X453))))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X453,X454),bool),'const.pair.UNCURRY_2'(s(fun(X456,fun(X455,fun('type.pair.prod'(X453,X454),bool))),X457),s('type.pair.prod'(X456,X455),X458))),s(X453,X459),s(X454,X460))))) )),
  inference(forward_demodulation,[],[f8072,f712])).
fof(f8072,plain,(
  ( ! [X453,X455,X457,X459,X454,X460,X456,X458] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X453,X454),bool),'const.pair.UNCURRY_2'(s(fun(X456,fun(X455,fun('type.pair.prod'(X453,X454),bool))),X457),s('type.pair.prod'(X456,X455),X458))),s(X453,X459),s(X454,X460)))) | p(s(bool,happ(s(fun('type.pair.prod'(X453,X454),bool),happ(s(fun(X455,fun('type.pair.prod'(X453,X454),bool)),happ(s(fun(X456,fun(X455,fun('type.pair.prod'(X453,X454),bool))),X457),s(X456,sK28(X458,X455,X456)))),s(X455,sK29(X458,X455,X456)))),s('type.pair.prod'(X453,X454),sK22(happ(s(fun(X455,fun('type.pair.prod'(X453,X454),bool)),happ(s(fun(X456,fun(X455,fun('type.pair.prod'(X453,X454),bool))),X457),s(X456,sK28(X458,X455,X456)))),s(X455,sK29(X458,X455,X456))),X454,X453)))))) )),
  inference(superposition,[],[f375,f712])).
fof(f8141,plain,(
  ( ! [X445,X447,X449,X451,X446,X452,X448,X450] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X445,X446),bool),'const.pair.UNCURRY_2'(s(fun(X448,fun(X447,fun('type.pair.prod'(X445,X446),bool))),X449),s('type.pair.prod'(X448,X447),X450))),s(X445,X451),s(X446,X452)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X445,X446),bool),'const.pair.UNCURRY_2'(s(fun(X448,fun(X447,fun('type.pair.prod'(X445,X446),bool))),X449),s('type.pair.prod'(X448,X447),X450))),s('type.pair.prod'(X445,X446),sK12(happ(s(fun(X447,fun('type.pair.prod'(X445,X446),bool)),happ(s(fun(X448,fun(X447,fun('type.pair.prod'(X445,X446),bool))),X449),s(X448,sK28(X450,X447,X448)))),s(X447,sK29(X450,X447,X448))),X446,X445)))))) )),
  inference(forward_demodulation,[],[f8071,f712])).
fof(f8071,plain,(
  ( ! [X445,X447,X449,X451,X446,X452,X448,X450] : (~p(s(bool,happ(s(fun('type.pair.prod'(X445,X446),bool),'const.pair.UNCURRY_2'(s(fun(X448,fun(X447,fun('type.pair.prod'(X445,X446),bool))),X449),s('type.pair.prod'(X448,X447),X450))),s('type.pair.prod'(X445,X446),sK12(happ(s(fun(X447,fun('type.pair.prod'(X445,X446),bool)),happ(s(fun(X448,fun(X447,fun('type.pair.prod'(X445,X446),bool))),X449),s(X448,sK28(X450,X447,X448)))),s(X447,sK29(X450,X447,X448))),X446,X445))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X445,X446),bool),happ(s(fun(X447,fun('type.pair.prod'(X445,X446),bool)),happ(s(fun(X448,fun(X447,fun('type.pair.prod'(X445,X446),bool))),X449),s(X448,sK28(X450,X447,X448)))),s(X447,sK29(X450,X447,X448)))),s(X445,X451),s(X446,X452))))) )),
  inference(superposition,[],[f376,f712])).
fof(f8140,plain,(
  ( ! [X441,X443,X438,X444,X440,X442,X439] : (p(s(bool,happ(s(fun('type.pair.prod'(X438,X439),bool),'const.pair.UNCURRY_2'(s(fun(X441,fun(X440,fun('type.pair.prod'(X438,X439),bool))),X442),s('type.pair.prod'(X441,X440),X443))),s('type.pair.prod'(X438,X439),X444)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X438,X439),bool),'const.pair.UNCURRY_2'(s(fun(X441,fun(X440,fun('type.pair.prod'(X438,X439),bool))),X442),s('type.pair.prod'(X441,X440),X443))),s(X438,sK10(happ(s(fun(X440,fun('type.pair.prod'(X438,X439),bool)),happ(s(fun(X441,fun(X440,fun('type.pair.prod'(X438,X439),bool))),X442),s(X441,sK28(X443,X440,X441)))),s(X440,sK29(X443,X440,X441))),X439,X438)),s(X439,sK11(happ(s(fun(X440,fun('type.pair.prod'(X438,X439),bool)),happ(s(fun(X441,fun(X440,fun('type.pair.prod'(X438,X439),bool))),X442),s(X441,sK28(X443,X440,X441)))),s(X440,sK29(X443,X440,X441))),X439,X438)))))) )),
  inference(forward_demodulation,[],[f8070,f712])).
fof(f8070,plain,(
  ( ! [X441,X443,X438,X444,X440,X442,X439] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X438,X439),bool),'const.pair.UNCURRY_2'(s(fun(X441,fun(X440,fun('type.pair.prod'(X438,X439),bool))),X442),s('type.pair.prod'(X441,X440),X443))),s(X438,sK10(happ(s(fun(X440,fun('type.pair.prod'(X438,X439),bool)),happ(s(fun(X441,fun(X440,fun('type.pair.prod'(X438,X439),bool))),X442),s(X441,sK28(X443,X440,X441)))),s(X440,sK29(X443,X440,X441))),X439,X438)),s(X439,sK11(happ(s(fun(X440,fun('type.pair.prod'(X438,X439),bool)),happ(s(fun(X441,fun(X440,fun('type.pair.prod'(X438,X439),bool))),X442),s(X441,sK28(X443,X440,X441)))),s(X440,sK29(X443,X440,X441))),X439,X438))))) | p(s(bool,happ(s(fun('type.pair.prod'(X438,X439),bool),happ(s(fun(X440,fun('type.pair.prod'(X438,X439),bool)),happ(s(fun(X441,fun(X440,fun('type.pair.prod'(X438,X439),bool))),X442),s(X441,sK28(X443,X440,X441)))),s(X440,sK29(X443,X440,X441)))),s('type.pair.prod'(X438,X439),X444))))) )),
  inference(superposition,[],[f377,f712])).
fof(f8139,plain,(
  ( ! [X436,X432,X434,X431,X437,X433,X435] : (p(s(bool,happ(s(fun('type.pair.prod'(X431,X432),bool),'const.pair.UNCURRY_2'(s(fun(X434,fun(X433,fun('type.pair.prod'(X431,X432),bool))),X435),s('type.pair.prod'(X434,X433),X436))),s('type.pair.prod'(X431,X432),X437)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X431,X432),bool),'const.pair.UNCURRY_2'(s(fun(X434,fun(X433,fun('type.pair.prod'(X431,X432),bool))),X435),s('type.pair.prod'(X434,X433),X436))),s(X431,sK13(happ(s(fun(X433,fun('type.pair.prod'(X431,X432),bool)),happ(s(fun(X434,fun(X433,fun('type.pair.prod'(X431,X432),bool))),X435),s(X434,sK28(X436,X433,X434)))),s(X433,sK29(X436,X433,X434))),X432,X431)),s(X432,sK14(happ(s(fun(X433,fun('type.pair.prod'(X431,X432),bool)),happ(s(fun(X434,fun(X433,fun('type.pair.prod'(X431,X432),bool))),X435),s(X434,sK28(X436,X433,X434)))),s(X433,sK29(X436,X433,X434))),X432,X431)))))) )),
  inference(forward_demodulation,[],[f8069,f712])).
fof(f8069,plain,(
  ( ! [X436,X432,X434,X431,X437,X433,X435] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X431,X432),bool),'const.pair.UNCURRY_2'(s(fun(X434,fun(X433,fun('type.pair.prod'(X431,X432),bool))),X435),s('type.pair.prod'(X434,X433),X436))),s(X431,sK13(happ(s(fun(X433,fun('type.pair.prod'(X431,X432),bool)),happ(s(fun(X434,fun(X433,fun('type.pair.prod'(X431,X432),bool))),X435),s(X434,sK28(X436,X433,X434)))),s(X433,sK29(X436,X433,X434))),X432,X431)),s(X432,sK14(happ(s(fun(X433,fun('type.pair.prod'(X431,X432),bool)),happ(s(fun(X434,fun(X433,fun('type.pair.prod'(X431,X432),bool))),X435),s(X434,sK28(X436,X433,X434)))),s(X433,sK29(X436,X433,X434))),X432,X431))))) | p(s(bool,happ(s(fun('type.pair.prod'(X431,X432),bool),happ(s(fun(X433,fun('type.pair.prod'(X431,X432),bool)),happ(s(fun(X434,fun(X433,fun('type.pair.prod'(X431,X432),bool))),X435),s(X434,sK28(X436,X433,X434)))),s(X433,sK29(X436,X433,X434)))),s('type.pair.prod'(X431,X432),X437))))) )),
  inference(superposition,[],[f378,f712])).
fof(f8138,plain,(
  ( ! [X428,X430,X424,X426,X429,X425,X427] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X424,X425),bool),'const.pair.UNCURRY_2'(s(fun(X427,fun(X426,fun('type.pair.prod'(X424,X425),bool))),X428),s('type.pair.prod'(X427,X426),X429))),s(X424,sK20(happ(s(fun(X426,fun('type.pair.prod'(X424,X425),bool)),happ(s(fun(X427,fun(X426,fun('type.pair.prod'(X424,X425),bool))),X428),s(X427,sK28(X429,X426,X427)))),s(X426,sK29(X429,X426,X427))),X425,X424)),s(X425,sK21(happ(s(fun(X426,fun('type.pair.prod'(X424,X425),bool)),happ(s(fun(X427,fun(X426,fun('type.pair.prod'(X424,X425),bool))),X428),s(X427,sK28(X429,X426,X427)))),s(X426,sK29(X429,X426,X427))),X425,X424))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X424,X425),bool),'const.pair.UNCURRY_2'(s(fun(X427,fun(X426,fun('type.pair.prod'(X424,X425),bool))),X428),s('type.pair.prod'(X427,X426),X429))),s('type.pair.prod'(X424,X425),X430))))) )),
  inference(forward_demodulation,[],[f8068,f712])).
fof(f8068,plain,(
  ( ! [X428,X430,X424,X426,X429,X425,X427] : (~p(s(bool,happ(s(fun('type.pair.prod'(X424,X425),bool),'const.pair.UNCURRY_2'(s(fun(X427,fun(X426,fun('type.pair.prod'(X424,X425),bool))),X428),s('type.pair.prod'(X427,X426),X429))),s('type.pair.prod'(X424,X425),X430)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X424,X425),bool),happ(s(fun(X426,fun('type.pair.prod'(X424,X425),bool)),happ(s(fun(X427,fun(X426,fun('type.pair.prod'(X424,X425),bool))),X428),s(X427,sK28(X429,X426,X427)))),s(X426,sK29(X429,X426,X427)))),s(X424,sK20(happ(s(fun(X426,fun('type.pair.prod'(X424,X425),bool)),happ(s(fun(X427,fun(X426,fun('type.pair.prod'(X424,X425),bool))),X428),s(X427,sK28(X429,X426,X427)))),s(X426,sK29(X429,X426,X427))),X425,X424)),s(X425,sK21(happ(s(fun(X426,fun('type.pair.prod'(X424,X425),bool)),happ(s(fun(X427,fun(X426,fun('type.pair.prod'(X424,X425),bool))),X428),s(X427,sK28(X429,X426,X427)))),s(X426,sK29(X429,X426,X427))),X425,X424)))))) )),
  inference(superposition,[],[f379,f712])).
fof(f8136,plain,(
  ( ! [X414,X416,X413,X415,X417] : (p(s(bool,happ(s(fun(X413,bool),'const.pair.UNCURRY_2'(s(fun(X415,fun(X414,fun(X413,bool))),X416),s('type.pair.prod'(X415,X414),X417))),s(X413,sK9(happ(s(fun(X414,fun(X413,bool)),happ(s(fun(X415,fun(X414,fun(X413,bool))),X416),s(X415,sK28(X417,X414,X415)))),s(X414,sK29(X417,X414,X415))),X413))))) | ~p(s(bool,happ(s(fun(fun(X413,bool),bool),'?_1'),s(fun(X413,bool),'const.pair.UNCURRY_2'(s(fun(X415,fun(X414,fun(X413,bool))),X416),s('type.pair.prod'(X415,X414),X417))))))) )),
  inference(forward_demodulation,[],[f8066,f712])).
fof(f8066,plain,(
  ( ! [X414,X416,X413,X415,X417] : (~p(s(bool,happ(s(fun(fun(X413,bool),bool),'?_1'),s(fun(X413,bool),'const.pair.UNCURRY_2'(s(fun(X415,fun(X414,fun(X413,bool))),X416),s('type.pair.prod'(X415,X414),X417)))))) | p(s(bool,happ(s(fun(X413,bool),happ(s(fun(X414,fun(X413,bool)),happ(s(fun(X415,fun(X414,fun(X413,bool))),X416),s(X415,sK28(X417,X414,X415)))),s(X414,sK29(X417,X414,X415)))),s(X413,sK9(happ(s(fun(X414,fun(X413,bool)),happ(s(fun(X415,fun(X414,fun(X413,bool))),X416),s(X415,sK28(X417,X414,X415)))),s(X414,sK29(X417,X414,X415))),X413)))))) )),
  inference(superposition,[],[f302,f712])).
fof(f8135,plain,(
  ( ! [X412,X408,X410,X409,X411] : (p(s(bool,happ(s(fun(fun(X408,bool),bool),'!_1'),s(fun(X408,bool),'const.pair.UNCURRY_2'(s(fun(X410,fun(X409,fun(X408,bool))),X411),s('type.pair.prod'(X410,X409),X412)))))) | ~p(s(bool,happ(s(fun(X408,bool),'const.pair.UNCURRY_2'(s(fun(X410,fun(X409,fun(X408,bool))),X411),s('type.pair.prod'(X410,X409),X412))),s(X408,sK8(happ(s(fun(X409,fun(X408,bool)),happ(s(fun(X410,fun(X409,fun(X408,bool))),X411),s(X410,sK28(X412,X409,X410)))),s(X409,sK29(X412,X409,X410))),X408)))))) )),
  inference(forward_demodulation,[],[f8065,f712])).
fof(f8065,plain,(
  ( ! [X412,X408,X410,X409,X411] : (~p(s(bool,happ(s(fun(X408,bool),'const.pair.UNCURRY_2'(s(fun(X410,fun(X409,fun(X408,bool))),X411),s('type.pair.prod'(X410,X409),X412))),s(X408,sK8(happ(s(fun(X409,fun(X408,bool)),happ(s(fun(X410,fun(X409,fun(X408,bool))),X411),s(X410,sK28(X412,X409,X410)))),s(X409,sK29(X412,X409,X410))),X408))))) | p(s(bool,happ(s(fun(fun(X408,bool),bool),'!_1'),s(fun(X408,bool),happ(s(fun(X409,fun(X408,bool)),happ(s(fun(X410,fun(X409,fun(X408,bool))),X411),s(X410,sK28(X412,X409,X410)))),s(X409,sK29(X412,X409,X410)))))))) )),
  inference(superposition,[],[f301,f712])).
fof(f8063,plain,(
  ( ! [X397,X399,X401,X396,X398,X400] : (p(s(bool,'|dollar|exists_1'(s(fun(X396,bool),'const.pair.UNCURRY_2'(s(fun(X398,fun(X397,fun(X396,bool))),X399),s('type.pair.prod'(X398,X397),X400)))))) | s(X396,X401) = s(X396,sK7(X401,happ(s(fun(X397,fun(X396,bool)),happ(s(fun(X398,fun(X397,fun(X396,bool))),X399),s(X398,sK28(X400,X397,X398)))),s(X397,sK29(X400,X397,X398))),X396))) )),
  inference(superposition,[],[f292,f712])).
fof(f8132,plain,(
  ( ! [X389,X385,X387,X388,X386] : (~p(s(bool,'|dollar|exists_1'(s(fun(X385,bool),'const.pair.UNCURRY_2'(s(fun(X387,fun(X386,fun(X385,bool))),X388),s('type.pair.prod'(X387,X386),X389)))))) | p(s(bool,happ(s(fun(X385,bool),'const.pair.UNCURRY_2'(s(fun(X387,fun(X386,fun(X385,bool))),X388),s('type.pair.prod'(X387,X386),X389))),s(X385,sK6(happ(s(fun(X386,fun(X385,bool)),happ(s(fun(X387,fun(X386,fun(X385,bool))),X388),s(X387,sK28(X389,X386,X387)))),s(X386,sK29(X389,X386,X387))),X385)))))) )),
  inference(forward_demodulation,[],[f8061,f712])).
fof(f8061,plain,(
  ( ! [X389,X385,X387,X388,X386] : (p(s(bool,happ(s(fun(X385,bool),'const.pair.UNCURRY_2'(s(fun(X387,fun(X386,fun(X385,bool))),X388),s('type.pair.prod'(X387,X386),X389))),s(X385,sK6(happ(s(fun(X386,fun(X385,bool)),happ(s(fun(X387,fun(X386,fun(X385,bool))),X388),s(X387,sK28(X389,X386,X387)))),s(X386,sK29(X389,X386,X387))),X385))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X385,bool),happ(s(fun(X386,fun(X385,bool)),happ(s(fun(X387,fun(X386,fun(X385,bool))),X388),s(X387,sK28(X389,X386,X387)))),s(X386,sK29(X389,X386,X387)))))))) )),
  inference(superposition,[],[f288,f712])).
fof(f8131,plain,(
  ( ! [X383,X380,X382,X384,X381] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X380,bool),'const.pair.UNCURRY_2'(s(fun(X382,fun(X381,fun(X380,bool))),X383),s('type.pair.prod'(X382,X381),X384)))))) | p(s(bool,happ(s(fun(X380,bool),'const.pair.UNCURRY_2'(s(fun(X382,fun(X381,fun(X380,bool))),X383),s('type.pair.prod'(X382,X381),X384))),s(X380,sK5(happ(s(fun(X381,fun(X380,bool)),happ(s(fun(X382,fun(X381,fun(X380,bool))),X383),s(X382,sK28(X384,X381,X382)))),s(X381,sK29(X384,X381,X382))),X380))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X380,bool),'const.pair.UNCURRY_2'(s(fun(X382,fun(X381,fun(X380,bool))),X383),s('type.pair.prod'(X382,X381),X384))))))) )),
  inference(forward_demodulation,[],[f8130,f712])).
fof(f8130,plain,(
  ( ! [X383,X380,X382,X384,X381] : (p(s(bool,happ(s(fun(X380,bool),'const.pair.UNCURRY_2'(s(fun(X382,fun(X381,fun(X380,bool))),X383),s('type.pair.prod'(X382,X381),X384))),s(X380,sK5(happ(s(fun(X381,fun(X380,bool)),happ(s(fun(X382,fun(X381,fun(X380,bool))),X383),s(X382,sK28(X384,X381,X382)))),s(X381,sK29(X384,X381,X382))),X380))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X380,bool),'const.pair.UNCURRY_2'(s(fun(X382,fun(X381,fun(X380,bool))),X383),s('type.pair.prod'(X382,X381),X384)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X380,bool),happ(s(fun(X381,fun(X380,bool)),happ(s(fun(X382,fun(X381,fun(X380,bool))),X383),s(X382,sK28(X384,X381,X382)))),s(X381,sK29(X384,X381,X382)))))))) )),
  inference(forward_demodulation,[],[f8060,f712])).
fof(f8060,plain,(
  ( ! [X383,X380,X382,X384,X381] : (~p(s(bool,'|dollar|exists_1'(s(fun(X380,bool),'const.pair.UNCURRY_2'(s(fun(X382,fun(X381,fun(X380,bool))),X383),s('type.pair.prod'(X382,X381),X384)))))) | p(s(bool,happ(s(fun(X380,bool),happ(s(fun(X381,fun(X380,bool)),happ(s(fun(X382,fun(X381,fun(X380,bool))),X383),s(X382,sK28(X384,X381,X382)))),s(X381,sK29(X384,X381,X382)))),s(X380,sK5(happ(s(fun(X381,fun(X380,bool)),happ(s(fun(X382,fun(X381,fun(X380,bool))),X383),s(X382,sK28(X384,X381,X382)))),s(X381,sK29(X384,X381,X382))),X380))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X380,bool),happ(s(fun(X381,fun(X380,bool)),happ(s(fun(X382,fun(X381,fun(X380,bool))),X383),s(X382,sK28(X384,X381,X382)))),s(X381,sK29(X384,X381,X382)))))))) )),
  inference(superposition,[],[f286,f712])).
fof(f8129,plain,(
  ( ! [X377,X379,X376,X378,X375] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X375,bool),'const.pair.UNCURRY_2'(s(fun(X377,fun(X376,fun(X375,bool))),X378),s('type.pair.prod'(X377,X376),X379)))))) | p(s(bool,happ(s(fun(X375,bool),'const.pair.UNCURRY_2'(s(fun(X377,fun(X376,fun(X375,bool))),X378),s('type.pair.prod'(X377,X376),X379))),s(X375,sK4(happ(s(fun(X376,fun(X375,bool)),happ(s(fun(X377,fun(X376,fun(X375,bool))),X378),s(X377,sK28(X379,X376,X377)))),s(X376,sK29(X379,X376,X377))),X375))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X375,bool),'const.pair.UNCURRY_2'(s(fun(X377,fun(X376,fun(X375,bool))),X378),s('type.pair.prod'(X377,X376),X379))))))) )),
  inference(forward_demodulation,[],[f8128,f712])).
fof(f8128,plain,(
  ( ! [X377,X379,X376,X378,X375] : (p(s(bool,happ(s(fun(X375,bool),'const.pair.UNCURRY_2'(s(fun(X377,fun(X376,fun(X375,bool))),X378),s('type.pair.prod'(X377,X376),X379))),s(X375,sK4(happ(s(fun(X376,fun(X375,bool)),happ(s(fun(X377,fun(X376,fun(X375,bool))),X378),s(X377,sK28(X379,X376,X377)))),s(X376,sK29(X379,X376,X377))),X375))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X375,bool),'const.pair.UNCURRY_2'(s(fun(X377,fun(X376,fun(X375,bool))),X378),s('type.pair.prod'(X377,X376),X379)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X375,bool),happ(s(fun(X376,fun(X375,bool)),happ(s(fun(X377,fun(X376,fun(X375,bool))),X378),s(X377,sK28(X379,X376,X377)))),s(X376,sK29(X379,X376,X377)))))))) )),
  inference(forward_demodulation,[],[f8059,f712])).
fof(f8059,plain,(
  ( ! [X377,X379,X376,X378,X375] : (~p(s(bool,'|dollar|exists_1'(s(fun(X375,bool),'const.pair.UNCURRY_2'(s(fun(X377,fun(X376,fun(X375,bool))),X378),s('type.pair.prod'(X377,X376),X379)))))) | p(s(bool,happ(s(fun(X375,bool),happ(s(fun(X376,fun(X375,bool)),happ(s(fun(X377,fun(X376,fun(X375,bool))),X378),s(X377,sK28(X379,X376,X377)))),s(X376,sK29(X379,X376,X377)))),s(X375,sK4(happ(s(fun(X376,fun(X375,bool)),happ(s(fun(X377,fun(X376,fun(X375,bool))),X378),s(X377,sK28(X379,X376,X377)))),s(X376,sK29(X379,X376,X377))),X375))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X375,bool),happ(s(fun(X376,fun(X375,bool)),happ(s(fun(X377,fun(X376,fun(X375,bool))),X378),s(X377,sK28(X379,X376,X377)))),s(X376,sK29(X379,X376,X377)))))))) )),
  inference(superposition,[],[f285,f712])).
fof(f8055,plain,(
  ( ! [X356,X352,X354,X351,X357,X353,X355] : (s(X352,happ(s(fun(X351,X352),'const.pair.UNCURRY_2'(s(fun(X354,fun(X353,fun(X351,X352))),X355),s('type.pair.prod'(X354,X353),X356))),s(X351,X357))) = s(X352,'const.pair.UNCURRY_2'(s(fun(X353,fun(X351,X352)),happ(s(fun(X354,fun(X353,fun(X351,X352))),X355),s(X354,sK28(X356,X353,X354)))),s('type.pair.prod'(X353,X351),'const.pair.,_2'(s(X353,sK29(X356,X353,X354)),s(X351,X357)))))) )),
  inference(superposition,[],[f345,f712])).
fof(f8054,plain,(
  ( ! [X350,X344,X346,X349,X345,X347,X348] : (s(X345,happ(s(fun(X344,X345),'const.pair.UNCURRY_2'(s(fun(X347,fun(X346,fun(X344,X345))),X348),s('type.pair.prod'(X347,X346),X349))),s(X344,X350))) = s(X345,happ(s(fun(X346,X345),'_dst_x1x0_2'(s(fun(X346,fun(X344,X345)),happ(s(fun(X347,fun(X346,fun(X344,X345))),X348),s(X347,sK28(X349,X346,X347)))),s(X344,X350))),s(X346,sK29(X349,X346,X347))))) )),
  inference(superposition,[],[f346,f712])).
fof(f8053,plain,(
  ( ! [X341,X343,X339,X340,X342,X338] : (~p(s(bool,happ(s(fun(X338,bool),'const.pair.UNCURRY_2'(s(fun(X340,fun(X339,fun(X338,bool))),X341),s('type.pair.prod'(X340,X339),X342))),s(X338,X343)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X339,X338),bool),'const.pair.UNCURRY_1'(s(fun(X339,fun(X338,bool)),happ(s(fun(X340,fun(X339,fun(X338,bool))),X341),s(X340,sK28(X342,X339,X340)))))))))) )),
  inference(superposition,[],[f319,f712])).
fof(f8052,plain,(
  ( ! [X333,X335,X337,X332,X334,X336] : (~p(s(bool,happ(s(fun(X332,bool),'const.pair.UNCURRY_2'(s(fun(X334,fun(X333,fun(X332,bool))),X335),s('type.pair.prod'(X334,X333),X336))),s(X332,X337)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X333,fun(X332,bool)),happ(s(fun(X334,fun(X333,fun(X332,bool))),X335),s(X334,sK28(X336,X333,X334)))),s('type.pair.prod'(X333,X332),sK25(happ(s(fun(X334,fun(X333,fun(X332,bool))),X335),s(X334,sK28(X336,X333,X334))),X332,X333)))))) )),
  inference(superposition,[],[f374,f712])).
fof(f8035,plain,(
  ( ! [X261,X263,X260,X262,X264] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X262,fun('type.pair.prod'(X260,X261),bool)),X263),s('type.pair.prod'(X262,'type.pair.prod'(X260,X261)),X264)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X260,X261),bool),happ(s(fun(X262,fun('type.pair.prod'(X260,X261),bool)),X263),s(X262,sK28(X264,'type.pair.prod'(X260,X261),X262)))),s(X260,sK20(happ(s(fun(X262,fun('type.pair.prod'(X260,X261),bool)),X263),s(X262,sK28(X264,'type.pair.prod'(X260,X261),X262))),X261,X260)),s(X261,sK21(happ(s(fun(X262,fun('type.pair.prod'(X260,X261),bool)),X263),s(X262,sK28(X264,'type.pair.prod'(X260,X261),X262))),X261,X260)))))) )),
  inference(superposition,[],[f379,f712])).
fof(f8034,plain,(
  ( ! [X259] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s('type.pair.prod'(bool,bool),X259)))) | p(s(bool,sK28(X259,bool,bool)))) )),
  inference(superposition,[],[f304,f712])).
fof(f8033,plain,(
  ( ! [X258] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s('type.pair.prod'(bool,bool),X258)))) | p(s(bool,sK29(X258,bool,bool)))) )),
  inference(superposition,[],[f305,f712])).
fof(f8032,plain,(
  ( ! [X255,X257,X254,X256] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X255,fun(X254,bool)),X256),s('type.pair.prod'(X255,X254),X257)))) | p(s(bool,'|dollar|exists_1'(s(fun(X254,bool),happ(s(fun(X255,fun(X254,bool)),X256),s(X255,sK28(X257,X254,X255)))))))) )),
  inference(superposition,[],[f289,f712])).
fof(f8031,plain,(
  ( ! [X253,X251,X252,X250] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X251,fun(X250,bool)),X252),s('type.pair.prod'(X251,X250),X253)))) | p(s(bool,happ(s(fun(fun(X250,bool),bool),'?_1'),s(fun(X250,bool),happ(s(fun(X251,fun(X250,bool)),X252),s(X251,sK28(X253,X250,X251)))))))) )),
  inference(superposition,[],[f303,f712])).
fof(f8029,plain,(
  ( ! [X243,X244,X242,X245] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X243,fun(X242,bool)),X244),s('type.pair.prod'(X243,X242),X245)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X243,fun(X242,bool)),X244),s('type.pair.prod'(X243,X242),sK25(X244,X242,X243)))))) )),
  inference(superposition,[],[f374,f712])).
fof(f8028,plain,(
  ( ! [X241,X239,X240] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X240,fun(X239,bool)),sK32(X239,X240)),s('type.pair.prod'(X240,X239),X241)))) | s(X240,sK33(X239,X240)) = s(X240,sK28(X241,X239,X240))) )),
  inference(superposition,[],[f327,f712])).
fof(f8027,plain,(
  ( ! [X237,X238,X236] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X237,fun(X236,bool)),sK32(X236,X237)),s('type.pair.prod'(X237,X236),X238)))) | s(X236,sK34(X236,X237)) = s(X236,sK29(X238,X236,X237))) )),
  inference(superposition,[],[f328,f712])).
fof(f8026,plain,(
  ( ! [X235,X234] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X234,fun(X234,bool)),'=_0'),s('type.pair.prod'(X234,X234),X235)))) | s(X234,sK28(X235,X234,X234)) = s(X234,sK29(X235,X234,X234))) )),
  inference(superposition,[],[f330,f712])).
fof(f8025,plain,(
  ( ! [X233,X232] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X232,fun(X232,bool)),'=_0'),s('type.pair.prod'(X232,X232),X233)))) | s(X232,sK28(X233,X232,X232)) != s(X232,sK29(X233,X232,X232))) )),
  inference(superposition,[],[f331,f712])).
fof(f7995,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X16] : (s(X15,happ(s(fun(X10,X15),X16),s(X10,'const.pair.UNCURRY_2'(s(fun(X12,fun(X11,X10)),X13),s('type.pair.prod'(X12,X11),X14))))) = s(X15,happ(s(fun(X11,X15),'_dst_x0x1_2'(s(fun(X10,X15),X16),s(fun(X11,X10),happ(s(fun(X12,fun(X11,X10)),X13),s(X12,sK28(X14,X11,X12)))))),s(X11,sK29(X14,X11,X12))))) )),
  inference(superposition,[],[f344,f712])).
fof(f7992,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X17,fun(X16,bool)),X18),s('type.pair.prod'(X17,X16),X19))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),happ(s(fun(X17,fun(X16,bool)),X18),s(X17,sK28(X19,X16,X17)))),s(X16,sK29(X19,X16,X17))))) )),
  inference(superposition,[],[f712,f264])).
fof(f7989,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X5,'const.pair.UNCURRY_2'(s(fun(X7,fun(X6,X5)),X8),s('type.pair.prod'(X7,X6),X9))) = s(X5,happ(s(fun(X7,X5),'_dst_x1x0_2'(s(fun(X7,fun(X6,X5)),X8),s(X6,sK29(X9,X6,X7)))),s(X7,sK28(X9,X6,X7))))) )),
  inference(superposition,[],[f712,f346])).
fof(f7988,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,'const.pair.UNCURRY_2'(s(fun(X1,fun(bool,X2)),X3),s('type.pair.prod'(X1,bool),X0))) = s(X2,happ(s(fun(bool,X2),happ(s(fun(X1,fun(bool,X2)),X3),s(X1,sK28(X0,bool,X1)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK29(X0,bool,X1))) )),
  inference(superposition,[],[f712,f264])).
fof(f7987,plain,(
  ( ! [X30,X35,X33,X31,X36,X34,X32] : (s(X31,'const.pair.UNCURRY_2'(s(fun(X32,fun(X30,X31)),'_dst_x1x0_2'(s(fun(X32,fun(X33,fun(X30,X31))),X34),s(X33,X35))),s('type.pair.prod'(X32,X30),X36))) = s(X31,happ(s(fun(X30,X31),happ(s(fun(X33,fun(X30,X31)),happ(s(fun(X32,fun(X33,fun(X30,X31))),X34),s(X32,sK28(X36,X30,X32)))),s(X33,X35))),s(X30,sK29(X36,X30,X32))))) )),
  inference(superposition,[],[f712,f346])).
fof(f7986,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (s(X24,'const.pair.UNCURRY_2'(s(fun(X25,fun(X23,X24)),'_dst_x0x1_2'(s(fun(X26,fun(X23,X24)),X27),s(fun(X25,X26),X28))),s('type.pair.prod'(X25,X23),X29))) = s(X24,happ(s(fun(X23,X24),happ(s(fun(X26,fun(X23,X24)),X27),s(X26,happ(s(fun(X25,X26),X28),s(X25,sK28(X29,X23,X25)))))),s(X23,sK29(X29,X23,X25))))) )),
  inference(superposition,[],[f712,f344])).
fof(f7985,plain,(
  ( ! [X21,X19,X17,X22,X20,X18,X16] : (s(X17,'const.pair.UNCURRY_2'(s(fun(X18,fun(X16,X17)),happ(s(fun(X19,fun(X18,fun(X16,X17))),X20),s(X19,X21))),s('type.pair.prod'(X18,X16),X22))) = s(X17,happ(s(fun(X16,X17),'const.pair.UNCURRY_2'(s(fun(X19,fun(X18,fun(X16,X17))),X20),s('type.pair.prod'(X19,X18),'const.pair.,_2'(s(X19,X21),s(X18,sK28(X22,X16,X18)))))),s(X16,sK29(X22,X16,X18))))) )),
  inference(superposition,[],[f712,f345])).
fof(f7984,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(X10,'const.pair.UNCURRY_2'(s(fun(X11,fun(X9,X10)),happ(s(fun(X12,fun(X11,fun(X9,X10))),X13),s(X12,X14))),s('type.pair.prod'(X11,X9),X15))) = s(X10,happ(s(fun(X9,X10),happ(s(fun(X12,fun(X9,X10)),'_dst_x1x0_2'(s(fun(X12,fun(X11,fun(X9,X10))),X13),s(X11,sK28(X15,X9,X11)))),s(X12,X14))),s(X9,sK29(X15,X9,X11))))) )),
  inference(superposition,[],[f712,f346])).
fof(f7983,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(X4,'const.pair.UNCURRY_2'(s(fun('type.pair.prod'(X5,X6),fun(X3,X4)),'const.pair.UNCURRY_1'(s(fun(X5,fun(X6,fun(X3,X4))),X7))),s('type.pair.prod'('type.pair.prod'(X5,X6),X3),X8))) = s(X4,happ(s(fun(X3,X4),'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,fun(X3,X4))),X7),s('type.pair.prod'(X5,X6),sK28(X8,X3,'type.pair.prod'(X5,X6))))),s(X3,sK29(X8,X3,'type.pair.prod'(X5,X6)))))) )),
  inference(superposition,[],[f712,f335])).
fof(f7982,plain,(
  ( ! [X2,X0,X1] : (s(X1,'const.pair.UNCURRY_2'(s(fun(fun(X0,X1),fun(X0,X1)),'_i_0_0'),s('type.pair.prod'(fun(X0,X1),X0),X2))) = s(X1,happ(s(fun(X0,X1),sK28(X2,X0,fun(X0,X1))),s(X0,sK29(X2,X0,fun(X0,X1)))))) )),
  inference(superposition,[],[f712,f279])).
fof(f7981,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,'const.pair.UNCURRY_2'(s(fun(bool,fun(X1,X2)),X3),s('type.pair.prod'(bool,X1),X0))) = s(X2,happ(s(fun(X1,X2),happ(s(fun(bool,fun(X1,X2)),X3),s(bool,'F_0'))),s(X1,sK29(X0,X1,bool)))) | s(bool,'T_0') = s(bool,sK28(X0,X1,bool))) )),
  inference(superposition,[],[f712,f264])).
fof(f7865,plain,(
  ( ! [X668,X670,X673,X667,X669,X671,X672] : (s(X667,sK28(happ(s(fun(X669,'type.pair.prod'(X667,X668)),happ(s(fun(X670,fun(X669,'type.pair.prod'(X667,X668))),X671),s(X670,X672))),s(X669,X673)),X668,X667)) = s(X667,'const.pair.FST_1'(s('type.pair.prod'(X667,X668),happ(s(fun(X670,'type.pair.prod'(X667,X668)),'_dst_x1x0_2'(s(fun(X670,fun(X669,'type.pair.prod'(X667,X668))),X671),s(X669,X673))),s(X670,X672)))))) )),
  inference(superposition,[],[f701,f346])).
fof(f7864,plain,(
  ( ! [X666,X660,X662,X664,X661,X663,X665] : (s(X661,sK29(happ(s(fun(X662,'type.pair.prod'(X660,X661)),happ(s(fun(X663,fun(X662,'type.pair.prod'(X660,X661))),X664),s(X663,X665))),s(X662,X666)),X661,X660)) = s(X661,'const.pair.SND_1'(s('type.pair.prod'(X660,X661),happ(s(fun(X663,'type.pair.prod'(X660,X661)),'_dst_x1x0_2'(s(fun(X663,fun(X662,'type.pair.prod'(X660,X661))),X664),s(X662,X666))),s(X663,X665)))))) )),
  inference(superposition,[],[f702,f346])).
fof(f7854,plain,(
  ( ! [X573,X575,X569,X571,X576,X572,X574,X570,X577] : (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X569,X570),bool),'const.pair.UNCURRY_1'(s(fun(X569,fun(X570,bool)),happ(s(fun(X572,fun(X569,fun(X570,bool))),'_dst_x1x0_2'(s(fun(X572,fun(X571,fun(X569,fun(X570,bool)))),X573),s(X571,X575))),s(X572,X574)))))))) | p(s(bool,happ(s(fun(X570,bool),happ(s(fun(X569,fun(X570,bool)),happ(s(fun(X571,fun(X569,fun(X570,bool))),happ(s(fun(X572,fun(X571,fun(X569,fun(X570,bool)))),X573),s(X572,X574))),s(X571,X575))),s(X569,X576))),s(X570,X577))))) )),
  inference(superposition,[],[f314,f346])).
fof(f7853,plain,(
  ( ! [X565,X567,X561,X563,X568,X564,X566,X560,X562] : (~p(s(bool,happ(s(fun(X561,bool),happ(s(fun(X560,fun(X561,bool)),happ(s(fun(X563,fun(X560,fun(X561,bool))),'_dst_x1x0_2'(s(fun(X563,fun(X562,fun(X560,fun(X561,bool)))),X564),s(X562,X566))),s(X563,X565))),s(X560,X567))),s(X561,X568)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X560,X561),bool),'const.pair.UNCURRY_1'(s(fun(X560,fun(X561,bool)),happ(s(fun(X562,fun(X560,fun(X561,bool))),happ(s(fun(X563,fun(X562,fun(X560,fun(X561,bool)))),X564),s(X563,X565))),s(X562,X566))))))))) )),
  inference(superposition,[],[f319,f346])).
fof(f7852,plain,(
  ( ! [X557,X559,X553,X555,X551,X556,X558,X552,X554] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X551,fun(X552,bool)),happ(s(fun(X554,fun(X551,fun(X552,bool))),'_dst_x1x0_2'(s(fun(X554,fun(X553,fun(X551,fun(X552,bool)))),X555),s(X553,X557))),s(X554,X556))),s('type.pair.prod'(X551,X552),sK15(happ(s(fun(X553,fun(X551,fun(X552,bool))),happ(s(fun(X554,fun(X553,fun(X551,fun(X552,bool)))),X555),s(X554,X556))),s(X553,X557)),X552,X551))))) | p(s(bool,happ(s(fun(X552,bool),happ(s(fun(X551,fun(X552,bool)),happ(s(fun(X553,fun(X551,fun(X552,bool))),happ(s(fun(X554,fun(X553,fun(X551,fun(X552,bool)))),X555),s(X554,X556))),s(X553,X557))),s(X551,X558))),s(X552,X559))))) )),
  inference(superposition,[],[f373,f346])).
fof(f7851,plain,(
  ( ! [X546,X542,X549,X545,X547,X543,X548,X550,X544] : (~p(s(bool,happ(s(fun(X543,bool),happ(s(fun(X542,fun(X543,bool)),happ(s(fun(X545,fun(X542,fun(X543,bool))),'_dst_x1x0_2'(s(fun(X545,fun(X544,fun(X542,fun(X543,bool)))),X546),s(X544,X548))),s(X545,X547))),s(X542,X549))),s(X543,X550)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X542,fun(X543,bool)),happ(s(fun(X544,fun(X542,fun(X543,bool))),happ(s(fun(X545,fun(X544,fun(X542,fun(X543,bool)))),X546),s(X545,X547))),s(X544,X548))),s('type.pair.prod'(X542,X543),sK25(happ(s(fun(X544,fun(X542,fun(X543,bool))),happ(s(fun(X545,fun(X544,fun(X542,fun(X543,bool)))),X546),s(X545,X547))),s(X544,X548)),X543,X542)))))) )),
  inference(superposition,[],[f374,f346])).
fof(f7850,plain,(
  ( ! [X540,X536,X538,X534,X541,X537,X539,X533,X535] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X533,X534),bool),happ(s(fun(X536,fun('type.pair.prod'(X533,X534),bool)),'_dst_x1x0_2'(s(fun(X536,fun(X535,fun('type.pair.prod'(X533,X534),bool))),X537),s(X535,X539))),s(X536,X538))),s(X533,X540),s(X534,X541)))) | p(s(bool,happ(s(fun('type.pair.prod'(X533,X534),bool),happ(s(fun(X535,fun('type.pair.prod'(X533,X534),bool)),happ(s(fun(X536,fun(X535,fun('type.pair.prod'(X533,X534),bool))),X537),s(X536,X538))),s(X535,X539))),s('type.pair.prod'(X533,X534),sK22(happ(s(fun(X535,fun('type.pair.prod'(X533,X534),bool)),happ(s(fun(X536,fun(X535,fun('type.pair.prod'(X533,X534),bool))),X537),s(X536,X538))),s(X535,X539)),X534,X533)))))) )),
  inference(superposition,[],[f375,f346])).
fof(f7849,plain,(
  ( ! [X532,X528,X530,X524,X526,X529,X531,X525,X527] : (~p(s(bool,happ(s(fun('type.pair.prod'(X524,X525),bool),happ(s(fun(X527,fun('type.pair.prod'(X524,X525),bool)),'_dst_x1x0_2'(s(fun(X527,fun(X526,fun('type.pair.prod'(X524,X525),bool))),X528),s(X526,X530))),s(X527,X529))),s('type.pair.prod'(X524,X525),sK12(happ(s(fun(X526,fun('type.pair.prod'(X524,X525),bool)),happ(s(fun(X527,fun(X526,fun('type.pair.prod'(X524,X525),bool))),X528),s(X527,X529))),s(X526,X530)),X525,X524))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X524,X525),bool),happ(s(fun(X526,fun('type.pair.prod'(X524,X525),bool)),happ(s(fun(X527,fun(X526,fun('type.pair.prod'(X524,X525),bool))),X528),s(X527,X529))),s(X526,X530))),s(X524,X531),s(X525,X532))))) )),
  inference(superposition,[],[f376,f346])).
fof(f7848,plain,(
  ( ! [X520,X522,X516,X518,X521,X523,X517,X519] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X516,X517),bool),happ(s(fun(X519,fun('type.pair.prod'(X516,X517),bool)),'_dst_x1x0_2'(s(fun(X519,fun(X518,fun('type.pair.prod'(X516,X517),bool))),X520),s(X518,X522))),s(X519,X521))),s(X516,sK10(happ(s(fun(X518,fun('type.pair.prod'(X516,X517),bool)),happ(s(fun(X519,fun(X518,fun('type.pair.prod'(X516,X517),bool))),X520),s(X519,X521))),s(X518,X522)),X517,X516)),s(X517,sK11(happ(s(fun(X518,fun('type.pair.prod'(X516,X517),bool)),happ(s(fun(X519,fun(X518,fun('type.pair.prod'(X516,X517),bool))),X520),s(X519,X521))),s(X518,X522)),X517,X516))))) | p(s(bool,happ(s(fun('type.pair.prod'(X516,X517),bool),happ(s(fun(X518,fun('type.pair.prod'(X516,X517),bool)),happ(s(fun(X519,fun(X518,fun('type.pair.prod'(X516,X517),bool))),X520),s(X519,X521))),s(X518,X522))),s('type.pair.prod'(X516,X517),X523))))) )),
  inference(superposition,[],[f377,f346])).
fof(f7847,plain,(
  ( ! [X509,X511,X512,X514,X508,X510,X513,X515] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X508,X509),bool),happ(s(fun(X511,fun('type.pair.prod'(X508,X509),bool)),'_dst_x1x0_2'(s(fun(X511,fun(X510,fun('type.pair.prod'(X508,X509),bool))),X512),s(X510,X514))),s(X511,X513))),s(X508,sK13(happ(s(fun(X510,fun('type.pair.prod'(X508,X509),bool)),happ(s(fun(X511,fun(X510,fun('type.pair.prod'(X508,X509),bool))),X512),s(X511,X513))),s(X510,X514)),X509,X508)),s(X509,sK14(happ(s(fun(X510,fun('type.pair.prod'(X508,X509),bool)),happ(s(fun(X511,fun(X510,fun('type.pair.prod'(X508,X509),bool))),X512),s(X511,X513))),s(X510,X514)),X509,X508))))) | p(s(bool,happ(s(fun('type.pair.prod'(X508,X509),bool),happ(s(fun(X510,fun('type.pair.prod'(X508,X509),bool)),happ(s(fun(X511,fun(X510,fun('type.pair.prod'(X508,X509),bool))),X512),s(X511,X513))),s(X510,X514))),s('type.pair.prod'(X508,X509),X515))))) )),
  inference(superposition,[],[f378,f346])).
fof(f7846,plain,(
  ( ! [X505,X507,X500,X502,X504,X506,X501,X503] : (~p(s(bool,happ(s(fun('type.pair.prod'(X500,X501),bool),happ(s(fun(X503,fun('type.pair.prod'(X500,X501),bool)),'_dst_x1x0_2'(s(fun(X503,fun(X502,fun('type.pair.prod'(X500,X501),bool))),X504),s(X502,X506))),s(X503,X505))),s('type.pair.prod'(X500,X501),X507)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X500,X501),bool),happ(s(fun(X502,fun('type.pair.prod'(X500,X501),bool)),happ(s(fun(X503,fun(X502,fun('type.pair.prod'(X500,X501),bool))),X504),s(X503,X505))),s(X502,X506))),s(X500,sK20(happ(s(fun(X502,fun('type.pair.prod'(X500,X501),bool)),happ(s(fun(X503,fun(X502,fun('type.pair.prod'(X500,X501),bool))),X504),s(X503,X505))),s(X502,X506)),X501,X500)),s(X501,sK21(happ(s(fun(X502,fun('type.pair.prod'(X500,X501),bool)),happ(s(fun(X503,fun(X502,fun('type.pair.prod'(X500,X501),bool))),X504),s(X503,X505))),s(X502,X506)),X501,X500)))))) )),
  inference(superposition,[],[f379,f346])).
fof(f7845,plain,(
  ( ! [X499,X494,X496,X498,X493,X495,X497] : (~p(s(bool,happ(s(fun(X493,bool),happ(s(fun(X495,fun(X493,bool)),'_dst_x1x0_2'(s(fun(X495,fun(X494,fun(X493,bool))),X496),s(X494,X498))),s(X495,X497))),s(X493,X499)))) | p(s(bool,happ(s(fun(fun(X493,bool),bool),'?_1'),s(fun(X493,bool),happ(s(fun(X494,fun(X493,bool)),happ(s(fun(X495,fun(X494,fun(X493,bool))),X496),s(X495,X497))),s(X494,X498))))))) )),
  inference(superposition,[],[f303,f346])).
fof(f7844,plain,(
  ( ! [X492,X488,X490,X487,X489,X491] : (~p(s(bool,happ(s(fun(fun(X487,bool),bool),'?_1'),s(fun(X487,bool),happ(s(fun(X489,fun(X487,bool)),'_dst_x1x0_2'(s(fun(X489,fun(X488,fun(X487,bool))),X490),s(X488,X492))),s(X489,X491)))))) | p(s(bool,happ(s(fun(X487,bool),happ(s(fun(X488,fun(X487,bool)),happ(s(fun(X489,fun(X488,fun(X487,bool))),X490),s(X489,X491))),s(X488,X492))),s(X487,sK9(happ(s(fun(X488,fun(X487,bool)),happ(s(fun(X489,fun(X488,fun(X487,bool))),X490),s(X489,X491))),s(X488,X492)),X487)))))) )),
  inference(superposition,[],[f302,f346])).
fof(f7843,plain,(
  ( ! [X484,X486,X482,X485,X481,X483] : (~p(s(bool,happ(s(fun(X481,bool),happ(s(fun(X483,fun(X481,bool)),'_dst_x1x0_2'(s(fun(X483,fun(X482,fun(X481,bool))),X484),s(X482,X486))),s(X483,X485))),s(X481,sK8(happ(s(fun(X482,fun(X481,bool)),happ(s(fun(X483,fun(X482,fun(X481,bool))),X484),s(X483,X485))),s(X482,X486)),X481))))) | p(s(bool,happ(s(fun(fun(X481,bool),bool),'!_1'),s(fun(X481,bool),happ(s(fun(X482,fun(X481,bool)),happ(s(fun(X483,fun(X482,fun(X481,bool))),X484),s(X483,X485))),s(X482,X486))))))) )),
  inference(superposition,[],[f301,f346])).
fof(f7842,plain,(
  ( ! [X476,X478,X474,X480,X475,X477,X479] : (~p(s(bool,happ(s(fun(fun(X474,bool),bool),'!_1'),s(fun(X474,bool),happ(s(fun(X476,fun(X474,bool)),'_dst_x1x0_2'(s(fun(X476,fun(X475,fun(X474,bool))),X477),s(X475,X479))),s(X476,X478)))))) | p(s(bool,happ(s(fun(X474,bool),happ(s(fun(X475,fun(X474,bool)),happ(s(fun(X476,fun(X475,fun(X474,bool))),X477),s(X476,X478))),s(X475,X479))),s(X474,X480))))) )),
  inference(superposition,[],[f300,f346])).
fof(f7841,plain,(
  ( ! [X472,X469,X471,X467,X473,X468,X470] : (p(s(bool,'|dollar|exists_1'(s(fun(X467,bool),happ(s(fun(X469,fun(X467,bool)),'_dst_x1x0_2'(s(fun(X469,fun(X468,fun(X467,bool))),X470),s(X468,X472))),s(X469,X471)))))) | s(X467,X473) = s(X467,sK7(X473,happ(s(fun(X468,fun(X467,bool)),happ(s(fun(X469,fun(X468,fun(X467,bool))),X470),s(X469,X471))),s(X468,X472)),X467))) )),
  inference(superposition,[],[f292,f346])).
fof(f7840,plain,(
  ( ! [X466,X461,X463,X465,X460,X462,X464] : (~p(s(bool,happ(s(fun(X460,bool),happ(s(fun(X462,fun(X460,bool)),'_dst_x1x0_2'(s(fun(X462,fun(X461,fun(X460,bool))),X463),s(X461,X465))),s(X462,X464))),s(X460,X466)))) | p(s(bool,'|dollar|exists_1'(s(fun(X460,bool),happ(s(fun(X461,fun(X460,bool)),happ(s(fun(X462,fun(X461,fun(X460,bool))),X463),s(X462,X464))),s(X461,X465))))))) )),
  inference(superposition,[],[f289,f346])).
fof(f7839,plain,(
  ( ! [X455,X457,X459,X454,X456,X458] : (p(s(bool,happ(s(fun(X454,bool),happ(s(fun(X456,fun(X454,bool)),'_dst_x1x0_2'(s(fun(X456,fun(X455,fun(X454,bool))),X457),s(X455,X459))),s(X456,X458))),s(X454,sK6(happ(s(fun(X455,fun(X454,bool)),happ(s(fun(X456,fun(X455,fun(X454,bool))),X457),s(X456,X458))),s(X455,X459)),X454))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X454,bool),happ(s(fun(X455,fun(X454,bool)),happ(s(fun(X456,fun(X455,fun(X454,bool))),X457),s(X456,X458))),s(X455,X459))))))) )),
  inference(superposition,[],[f288,f346])).
fof(f7838,plain,(
  ( ! [X453,X449,X451,X452,X448,X450] : (~p(s(bool,'|dollar|exists_1'(s(fun(X448,bool),happ(s(fun(X450,fun(X448,bool)),'_dst_x1x0_2'(s(fun(X450,fun(X449,fun(X448,bool))),X451),s(X449,X453))),s(X450,X452)))))) | p(s(bool,happ(s(fun(X448,bool),happ(s(fun(X449,fun(X448,bool)),happ(s(fun(X450,fun(X449,fun(X448,bool))),X451),s(X450,X452))),s(X449,X453))),s(X448,sK5(happ(s(fun(X449,fun(X448,bool)),happ(s(fun(X450,fun(X449,fun(X448,bool))),X451),s(X450,X452))),s(X449,X453)),X448))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X448,bool),happ(s(fun(X449,fun(X448,bool)),happ(s(fun(X450,fun(X449,fun(X448,bool))),X451),s(X450,X452))),s(X449,X453))))))) )),
  inference(superposition,[],[f286,f346])).
fof(f7837,plain,(
  ( ! [X445,X447,X443,X444,X446,X442] : (~p(s(bool,'|dollar|exists_1'(s(fun(X442,bool),happ(s(fun(X444,fun(X442,bool)),'_dst_x1x0_2'(s(fun(X444,fun(X443,fun(X442,bool))),X445),s(X443,X447))),s(X444,X446)))))) | p(s(bool,happ(s(fun(X442,bool),happ(s(fun(X443,fun(X442,bool)),happ(s(fun(X444,fun(X443,fun(X442,bool))),X445),s(X444,X446))),s(X443,X447))),s(X442,sK4(happ(s(fun(X443,fun(X442,bool)),happ(s(fun(X444,fun(X443,fun(X442,bool))),X445),s(X444,X446))),s(X443,X447)),X442))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X442,bool),happ(s(fun(X443,fun(X442,bool)),happ(s(fun(X444,fun(X443,fun(X442,bool))),X445),s(X444,X446))),s(X443,X447))))))) )),
  inference(superposition,[],[f285,f346])).
fof(f7836,plain,(
  ( ! [X441,X436,X438,X434,X440,X437,X439,X435] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X434,bool),happ(s(fun(X436,fun(X434,bool)),'_dst_x1x0_2'(s(fun(X436,fun(X435,fun(X434,bool))),X437),s(X435,X439))),s(X436,X438)))))) | ~p(s(bool,happ(s(fun(X434,bool),happ(s(fun(X435,fun(X434,bool)),happ(s(fun(X436,fun(X435,fun(X434,bool))),X437),s(X436,X438))),s(X435,X439))),s(X434,X440)))) | ~p(s(bool,happ(s(fun(X434,bool),happ(s(fun(X435,fun(X434,bool)),happ(s(fun(X436,fun(X435,fun(X434,bool))),X437),s(X436,X438))),s(X435,X439))),s(X434,X441)))) | s(X434,X440) = s(X434,X441)) )),
  inference(superposition,[],[f284,f346])).
fof(f7835,plain,(
  ( ! [X428,X430,X432,X429,X431,X433] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X428,bool),happ(s(fun(X430,fun(X428,bool)),'_dst_x1x0_2'(s(fun(X430,fun(X429,fun(X428,bool))),X431),s(X429,X433))),s(X430,X432)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X428,bool),happ(s(fun(X429,fun(X428,bool)),happ(s(fun(X430,fun(X429,fun(X428,bool))),X431),s(X430,X432))),s(X429,X433))))))) )),
  inference(superposition,[],[f283,f346])).
fof(f7833,plain,(
  ( ! [X414,X420,X416,X418,X415,X421,X417,X419] : (s(X415,happ(s(fun(X414,X415),happ(s(fun(X417,fun(X414,X415)),'_dst_x1x0_2'(s(fun(X417,fun(X416,fun(X414,X415))),X418),s(X416,X420))),s(X417,X419))),s(X414,X421))) = s(X415,'const.pair.UNCURRY_2'(s(fun(X416,fun(X414,X415)),happ(s(fun(X417,fun(X416,fun(X414,X415))),X418),s(X417,X419))),s('type.pair.prod'(X416,X414),'const.pair.,_2'(s(X416,X420),s(X414,X421)))))) )),
  inference(superposition,[],[f345,f346])).
fof(f7832,plain,(
  ( ! [X412,X408,X410,X407,X413,X409,X411] : (~p(s(bool,happ(s(fun(X407,bool),happ(s(fun(X409,fun(X407,bool)),'_dst_x1x0_2'(s(fun(X409,fun(X408,fun(X407,bool))),X410),s(X408,X412))),s(X409,X411))),s(X407,X413)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X408,X407),bool),'const.pair.UNCURRY_1'(s(fun(X408,fun(X407,bool)),happ(s(fun(X409,fun(X408,fun(X407,bool))),X410),s(X409,X411))))))))) )),
  inference(superposition,[],[f319,f346])).
fof(f7831,plain,(
  ( ! [X405,X401,X403,X404,X406,X400,X402] : (~p(s(bool,happ(s(fun(X400,bool),happ(s(fun(X402,fun(X400,bool)),'_dst_x1x0_2'(s(fun(X402,fun(X401,fun(X400,bool))),X403),s(X401,X405))),s(X402,X404))),s(X400,X406)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X401,fun(X400,bool)),happ(s(fun(X402,fun(X401,fun(X400,bool))),X403),s(X402,X404))),s('type.pair.prod'(X401,X400),sK25(happ(s(fun(X402,fun(X401,fun(X400,bool))),X403),s(X402,X404)),X400,X401)))))) )),
  inference(superposition,[],[f374,f346])).
fof(f7830,plain,(
  ( ! [X397,X399,X395,X396,X398,X394] : (p(s(bool,happ(s(fun(X395,bool),'_dst_x1x0_2'(s(fun(X395,fun(X394,bool)),X396),s(X394,X398))),s(X395,X397)))) | s(bool,X399) = s(bool,happ(s(fun(X394,bool),happ(s(fun(X395,fun(X394,bool)),X396),s(X395,X397))),s(X394,X398))) | s(bool,'T_0') = s(bool,X399)) )),
  inference(superposition,[],[f468,f346])).
fof(f7829,plain,(
  ( ! [X389,X391,X393,X388,X390,X392] : (~p(s(bool,happ(s(fun(X389,bool),'_dst_x1x0_2'(s(fun(X389,fun(X388,bool)),X390),s(X388,X392))),s(X389,X391)))) | ~p(s(bool,X393)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun(X388,bool),happ(s(fun(X389,fun(X388,bool)),X390),s(X389,X391))),s(X388,X392))))),s(bool,X393))))) )),
  inference(superposition,[],[f306,f346])).
fof(f7828,plain,(
  ( ! [X383,X385,X387,X382,X384,X386] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X387))),s(bool,happ(s(fun(X383,bool),'_dst_x1x0_2'(s(fun(X383,fun(X382,bool)),X384),s(X382,X386))),s(X383,X385)))))) | p(s(bool,happ(s(fun(X382,bool),happ(s(fun(X383,fun(X382,bool)),X384),s(X383,X385))),s(X382,X386))))) )),
  inference(superposition,[],[f305,f346])).
fof(f7825,plain,(
  ( ! [X364,X366,X368,X365,X367,X369] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun(X365,bool),'_dst_x1x0_2'(s(fun(X365,fun(X364,bool)),X366),s(X364,X368))),s(X365,X367))))),s(bool,X369)))) | p(s(bool,happ(s(fun(X364,bool),happ(s(fun(X365,fun(X364,bool)),X366),s(X365,X367))),s(X364,X368))))) )),
  inference(superposition,[],[f304,f346])).
fof(f7824,plain,(
  ( ! [X358,X360,X362,X359,X361,X363] : (~p(s(bool,happ(s(fun(X359,bool),'_dst_x1x0_2'(s(fun(X359,fun(X358,bool)),X360),s(X358,X362))),s(X359,X361)))) | ~p(s(bool,X363)) | s(bool,X363) = s(bool,happ(s(fun(X358,bool),happ(s(fun(X359,fun(X358,bool)),X360),s(X359,X361))),s(X358,X362)))) )),
  inference(superposition,[],[f299,f346])).
fof(f7822,plain,(
  ( ! [X350,X346,X349,X351,X347,X348] : (p(s(bool,happ(s(fun(X347,bool),'_dst_x1x0_2'(s(fun(X347,fun(X346,bool)),X348),s(X346,X350))),s(X347,X349)))) | p(s(bool,X351)) | s(bool,X351) = s(bool,happ(s(fun(X346,bool),happ(s(fun(X347,fun(X346,bool)),X348),s(X347,X349))),s(X346,X350)))) )),
  inference(superposition,[],[f296,f346])).
fof(f7821,plain,(
  ( ! [X344,X341,X343,X345,X342] : (p(s(bool,happ(s(fun(X342,bool),'_dst_x1x0_2'(s(fun(X342,fun(X341,bool)),X343),s(X341,X345))),s(X342,X344)))) | s(bool,'F_0') = s(bool,happ(s(fun(X341,bool),happ(s(fun(X342,fun(X341,bool)),X343),s(X342,X344))),s(X341,X345)))) )),
  inference(superposition,[],[f270,f346])).
fof(f7820,plain,(
  ( ! [X337,X339,X340,X336,X338] : (~p(s(bool,happ(s(fun(X337,bool),'_dst_x1x0_2'(s(fun(X337,fun(X336,bool)),X338),s(X336,X340))),s(X337,X339)))) | s(bool,'F_0') != s(bool,happ(s(fun(X336,bool),happ(s(fun(X337,fun(X336,bool)),X338),s(X337,X339))),s(X336,X340)))) )),
  inference(superposition,[],[f269,f346])).
fof(f7819,plain,(
  ( ! [X333,X335,X331,X332,X334] : (~p(s(bool,happ(s(fun(X332,bool),'_dst_x1x0_2'(s(fun(X332,fun(X331,bool)),X333),s(X331,X335))),s(X332,X334)))) | s(bool,'T_0') = s(bool,happ(s(fun(X331,bool),happ(s(fun(X332,fun(X331,bool)),X333),s(X332,X334))),s(X331,X335)))) )),
  inference(superposition,[],[f266,f346])).
fof(f7818,plain,(
  ( ! [X327,X329,X326,X328,X330] : (s(bool,'T_0') != s(bool,happ(s(fun(X327,bool),'_dst_x1x0_2'(s(fun(X327,fun(X326,bool)),X328),s(X326,X330))),s(X327,X329))) | p(s(bool,happ(s(fun(X326,bool),happ(s(fun(X327,fun(X326,bool)),X328),s(X327,X329))),s(X326,X330))))) )),
  inference(superposition,[],[f265,f346])).
fof(f7814,plain,(
  ( ! [X313,X315,X316,X312,X314,X311] : (~p(s(bool,happ(s(fun(X313,bool),'_dst_x1x0_2'(s(fun(X313,fun('type.pair.prod'(X311,X312),bool)),X314),s('type.pair.prod'(X311,X312),X316))),s(X313,X315)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X311,X312),bool),happ(s(fun(X313,fun('type.pair.prod'(X311,X312),bool)),X314),s(X313,X315))),s(X311,sK20(happ(s(fun(X313,fun('type.pair.prod'(X311,X312),bool)),X314),s(X313,X315)),X312,X311)),s(X312,sK21(happ(s(fun(X313,fun('type.pair.prod'(X311,X312),bool)),X314),s(X313,X315)),X312,X311)))))) )),
  inference(superposition,[],[f379,f346])).
fof(f7813,plain,(
  ( ! [X308,X310,X304,X306,X309,X305,X307] : (~p(s(bool,happ(s(fun(X306,bool),'_dst_x1x0_2'(s(fun(X306,fun('type.pair.prod'(X304,X305),bool)),X307),s('type.pair.prod'(X304,X305),sK12(happ(s(fun(X306,fun('type.pair.prod'(X304,X305),bool)),X307),s(X306,X308)),X305,X304)))),s(X306,X308)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X304,X305),bool),happ(s(fun(X306,fun('type.pair.prod'(X304,X305),bool)),X307),s(X306,X308))),s(X304,X309),s(X305,X310))))) )),
  inference(superposition,[],[f376,f346])).
fof(f7812,plain,(
  ( ! [X302,X303] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X303))),s(bool,X302)))) | p(s(bool,X302))) )),
  inference(superposition,[],[f304,f346])).
fof(f7811,plain,(
  ( ! [X300,X301] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X301))),s(bool,X300)))) | p(s(bool,X301))) )),
  inference(superposition,[],[f305,f346])).
fof(f7810,plain,(
  ( ! [X296,X298,X297,X299] : (~p(s(bool,happ(s(fun(X297,bool),'_dst_x1x0_2'(s(fun(X297,fun(X296,bool)),X298),s(X296,sK8(happ(s(fun(X297,fun(X296,bool)),X298),s(X297,X299)),X296)))),s(X297,X299)))) | p(s(bool,happ(s(fun(fun(X296,bool),bool),'!_1'),s(fun(X296,bool),happ(s(fun(X297,fun(X296,bool)),X298),s(X297,X299))))))) )),
  inference(superposition,[],[f301,f346])).
fof(f7809,plain,(
  ( ! [X292,X294,X293,X295] : (p(s(bool,happ(s(fun(X293,bool),'_dst_x1x0_2'(s(fun(X293,fun(X292,bool)),X294),s(X292,sK6(happ(s(fun(X293,fun(X292,bool)),X294),s(X293,X295)),X292)))),s(X293,X295)))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X292,bool),happ(s(fun(X293,fun(X292,bool)),X294),s(X293,X295))))))) )),
  inference(superposition,[],[f288,f346])).
fof(f7807,plain,(
  ( ! [X285,X287,X283,X284,X286] : (~p(s(bool,happ(s(fun(X284,bool),'_dst_x1x0_2'(s(fun(X284,fun(X283,bool)),X285),s(X283,X287))),s(X284,X286)))) | p(s(bool,'|dollar|exists_1'(s(fun(X283,bool),happ(s(fun(X284,fun(X283,bool)),X285),s(X284,X286))))))) )),
  inference(superposition,[],[f289,f346])).
fof(f7806,plain,(
  ( ! [X280,X282,X279,X281,X278] : (~p(s(bool,happ(s(fun(X279,bool),'_dst_x1x0_2'(s(fun(X279,fun(X278,bool)),X280),s(X278,X282))),s(X279,X281)))) | p(s(bool,happ(s(fun(fun(X278,bool),bool),'?_1'),s(fun(X278,bool),happ(s(fun(X279,fun(X278,bool)),X280),s(X279,X281))))))) )),
  inference(superposition,[],[f303,f346])).
fof(f7805,plain,(
  ( ! [X277,X273,X275,X276,X274] : (~p(s(bool,happ(s(fun(X274,bool),'_dst_x1x0_2'(s(fun(X274,fun(X273,bool)),X275),s(X273,X277))),s(X274,X276)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X274,X273),bool),'const.pair.UNCURRY_1'(s(fun(X274,fun(X273,bool)),X275))))))) )),
  inference(superposition,[],[f319,f346])).
fof(f7804,plain,(
  ( ! [X269,X271,X268,X270,X272] : (~p(s(bool,happ(s(fun(X269,bool),'_dst_x1x0_2'(s(fun(X269,fun(X268,bool)),X270),s(X268,X272))),s(X269,X271)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X269,fun(X268,bool)),X270),s('type.pair.prod'(X269,X268),sK25(X270,X268,X269)))))) )),
  inference(superposition,[],[f374,f346])).
fof(f7803,plain,(
  ( ! [X265,X267,X264,X266] : (~p(s(bool,happ(s(fun(X265,bool),'_dst_x1x0_2'(s(fun(X265,fun(X264,bool)),sK32(X264,X265)),s(X264,X267))),s(X265,X266)))) | s(X265,X266) = s(X265,sK33(X264,X265))) )),
  inference(superposition,[],[f327,f346])).
fof(f7802,plain,(
  ( ! [X261,X263,X260,X262] : (~p(s(bool,happ(s(fun(X261,bool),'_dst_x1x0_2'(s(fun(X261,fun(X260,bool)),sK32(X260,X261)),s(X260,X263))),s(X261,X262)))) | s(X260,X263) = s(X260,sK34(X260,X261))) )),
  inference(superposition,[],[f328,f346])).
fof(f7801,plain,(
  ( ! [X257,X259,X258] : (~p(s(bool,happ(s(fun(X257,bool),'_dst_x1x0_2'(s(fun(X257,fun(X257,bool)),'=_0'),s(X257,X259))),s(X257,X258)))) | s(X257,X258) = s(X257,X259)) )),
  inference(superposition,[],[f330,f346])).
fof(f7800,plain,(
  ( ! [X255,X254,X256] : (p(s(bool,happ(s(fun(X254,bool),'_dst_x1x0_2'(s(fun(X254,fun(X254,bool)),'=_0'),s(X254,X256))),s(X254,X255)))) | s(X254,X255) != s(X254,X256)) )),
  inference(superposition,[],[f331,f346])).
fof(f7790,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162] : (p(s(bool,happ(s(fun(X161,bool),happ(s(fun(X161,fun(X161,bool)),'=_0'),s(X161,X167))),s(X161,happ(s(fun(X163,X161),'_dst_x1x0_2'(s(fun(X163,fun(X162,X161)),X164),s(X162,X166))),s(X163,X165)))))) | s(X161,X167) != s(X161,happ(s(fun(X162,X161),happ(s(fun(X163,fun(X162,X161)),X164),s(X163,X165))),s(X162,X166)))) )),
  inference(superposition,[],[f331,f346])).
fof(f7789,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X160] : (p(s(bool,happ(s(fun(X154,bool),happ(s(fun(X154,fun(X154,bool)),'=_0'),s(X154,happ(s(fun(X156,X154),'_dst_x1x0_2'(s(fun(X156,fun(X155,X154)),X157),s(X155,X159))),s(X156,X158))))),s(X154,X160)))) | s(X154,X160) != s(X154,happ(s(fun(X155,X154),happ(s(fun(X156,fun(X155,X154)),X157),s(X156,X158))),s(X155,X159)))) )),
  inference(superposition,[],[f331,f346])).
fof(f7788,plain,(
  ( ! [X152,X151,X149,X147,X153,X150,X148] : (~p(s(bool,happ(s(fun(X147,bool),happ(s(fun(X147,fun(X147,bool)),'=_0'),s(X147,X153))),s(X147,happ(s(fun(X149,X147),'_dst_x1x0_2'(s(fun(X149,fun(X148,X147)),X150),s(X148,X152))),s(X149,X151)))))) | s(X147,X153) = s(X147,happ(s(fun(X148,X147),happ(s(fun(X149,fun(X148,X147)),X150),s(X149,X151))),s(X148,X152)))) )),
  inference(superposition,[],[f330,f346])).
fof(f7787,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X140,bool),happ(s(fun(X140,fun(X140,bool)),'=_0'),s(X140,happ(s(fun(X142,X140),'_dst_x1x0_2'(s(fun(X142,fun(X141,X140)),X143),s(X141,X145))),s(X142,X144))))),s(X140,X146)))) | s(X140,X146) = s(X140,happ(s(fun(X141,X140),happ(s(fun(X142,fun(X141,X140)),X143),s(X142,X144))),s(X141,X145)))) )),
  inference(superposition,[],[f330,f346])).
fof(f7786,plain,(
  ( ! [X134,X132,X138,X136,X135,X133,X139,X137] : (s(X132,sK34(X132,X138)) != s(X132,happ(s(fun(X134,X132),'_dst_x1x0_2'(s(fun(X134,fun(X133,X132)),X135),s(X133,X137))),s(X134,X136))) | p(s(bool,happ(s(fun(X132,bool),happ(s(fun(X138,fun(X132,bool)),sK32(X132,X138)),s(X138,X139))),s(X132,happ(s(fun(X133,X132),happ(s(fun(X134,fun(X133,X132)),X135),s(X134,X136))),s(X133,X137)))))) | s(X138,X139) != s(X138,sK33(X132,X138))) )),
  inference(superposition,[],[f329,f346])).
fof(f7785,plain,(
  ( ! [X127,X125,X130,X128,X126,X124,X131,X129] : (~p(s(bool,happ(s(fun(X124,bool),happ(s(fun(X130,fun(X124,bool)),sK32(X124,X130)),s(X130,X131))),s(X124,happ(s(fun(X126,X124),'_dst_x1x0_2'(s(fun(X126,fun(X125,X124)),X127),s(X125,X129))),s(X126,X128)))))) | s(X124,sK34(X124,X130)) = s(X124,happ(s(fun(X125,X124),happ(s(fun(X126,fun(X125,X124)),X127),s(X126,X128))),s(X125,X129)))) )),
  inference(superposition,[],[f328,f346])).
fof(f7782,plain,(
  ( ! [X103,X101,X107,X105,X102,X100,X106,X104] : (~p(s(bool,happ(s(fun(X106,bool),happ(s(fun(X100,fun(X106,bool)),sK32(X106,X100)),s(X100,happ(s(fun(X102,X100),'_dst_x1x0_2'(s(fun(X102,fun(X101,X100)),X103),s(X101,X105))),s(X102,X104))))),s(X106,X107)))) | s(X100,sK33(X106,X100)) = s(X100,happ(s(fun(X101,X100),happ(s(fun(X102,fun(X101,X100)),X103),s(X102,X104))),s(X101,X105)))) )),
  inference(superposition,[],[f327,f346])).
fof(f7771,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (s(X20,happ(s(fun(X14,X20),X21),s(X14,happ(s(fun(X16,X14),'_dst_x1x0_2'(s(fun(X16,fun(X15,X14)),X17),s(X15,X19))),s(X16,X18))))) = s(X20,happ(s(fun(X15,X20),'_dst_x0x1_2'(s(fun(X14,X20),X21),s(fun(X15,X14),happ(s(fun(X16,fun(X15,X14)),X17),s(X16,X18))))),s(X15,X19)))) )),
  inference(superposition,[],[f344,f346])).
fof(f7757,plain,(
  ( ! [X633,X635,X629,X631,X632,X634,X630] : (s(X629,sK28(happ(s(fun(X631,'type.pair.prod'(X629,X630)),'_dst_x1x0_2'(s(fun(X631,fun(X632,'type.pair.prod'(X629,X630))),X633),s(X632,X634))),s(X631,X635)),X630,X629)) = s(X629,'const.pair.FST_1'(s('type.pair.prod'(X629,X630),happ(s(fun(X632,'type.pair.prod'(X629,X630)),happ(s(fun(X631,fun(X632,'type.pair.prod'(X629,X630))),X633),s(X631,X635))),s(X632,X634)))))) )),
  inference(superposition,[],[f701,f346])).
fof(f7756,plain,(
  ( ! [X625,X627,X623,X628,X624,X626,X622] : (s(X623,sK29(happ(s(fun(X624,'type.pair.prod'(X622,X623)),'_dst_x1x0_2'(s(fun(X624,fun(X625,'type.pair.prod'(X622,X623))),X626),s(X625,X627))),s(X624,X628)),X623,X622)) = s(X623,'const.pair.SND_1'(s('type.pair.prod'(X622,X623),happ(s(fun(X625,'type.pair.prod'(X622,X623)),happ(s(fun(X624,fun(X625,'type.pair.prod'(X622,X623))),X626),s(X624,X628))),s(X625,X627)))))) )),
  inference(superposition,[],[f702,f346])).
fof(f7746,plain,(
  ( ! [X536,X538,X532,X534,X537,X539,X533,X535,X531] : (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X531,X532),bool),'const.pair.UNCURRY_1'(s(fun(X531,fun(X532,bool)),happ(s(fun(X534,fun(X531,fun(X532,bool))),happ(s(fun(X533,fun(X534,fun(X531,fun(X532,bool)))),X535),s(X533,X537))),s(X534,X536)))))))) | p(s(bool,happ(s(fun(X532,bool),happ(s(fun(X531,fun(X532,bool)),happ(s(fun(X533,fun(X531,fun(X532,bool))),'_dst_x1x0_2'(s(fun(X533,fun(X534,fun(X531,fun(X532,bool)))),X535),s(X534,X536))),s(X533,X537))),s(X531,X538))),s(X532,X539))))) )),
  inference(superposition,[],[f314,f346])).
fof(f7745,plain,(
  ( ! [X528,X530,X524,X526,X522,X529,X525,X527,X523] : (~p(s(bool,happ(s(fun(X523,bool),happ(s(fun(X522,fun(X523,bool)),happ(s(fun(X525,fun(X522,fun(X523,bool))),happ(s(fun(X524,fun(X525,fun(X522,fun(X523,bool)))),X526),s(X524,X528))),s(X525,X527))),s(X522,X529))),s(X523,X530)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X522,X523),bool),'const.pair.UNCURRY_1'(s(fun(X522,fun(X523,bool)),happ(s(fun(X524,fun(X522,fun(X523,bool))),'_dst_x1x0_2'(s(fun(X524,fun(X525,fun(X522,fun(X523,bool)))),X526),s(X525,X527))),s(X524,X528))))))))) )),
  inference(superposition,[],[f319,f346])).
fof(f7744,plain,(
  ( ! [X520,X516,X518,X514,X521,X517,X519,X513,X515] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X513,fun(X514,bool)),happ(s(fun(X516,fun(X513,fun(X514,bool))),happ(s(fun(X515,fun(X516,fun(X513,fun(X514,bool)))),X517),s(X515,X519))),s(X516,X518))),s('type.pair.prod'(X513,X514),sK15(happ(s(fun(X515,fun(X513,fun(X514,bool))),'_dst_x1x0_2'(s(fun(X515,fun(X516,fun(X513,fun(X514,bool)))),X517),s(X516,X518))),s(X515,X519)),X514,X513))))) | p(s(bool,happ(s(fun(X514,bool),happ(s(fun(X513,fun(X514,bool)),happ(s(fun(X515,fun(X513,fun(X514,bool))),'_dst_x1x0_2'(s(fun(X515,fun(X516,fun(X513,fun(X514,bool)))),X517),s(X516,X518))),s(X515,X519))),s(X513,X520))),s(X514,X521))))) )),
  inference(superposition,[],[f373,f346])).
fof(f7743,plain,(
  ( ! [X509,X511,X505,X507,X512,X508,X510,X504,X506] : (~p(s(bool,happ(s(fun(X505,bool),happ(s(fun(X504,fun(X505,bool)),happ(s(fun(X507,fun(X504,fun(X505,bool))),happ(s(fun(X506,fun(X507,fun(X504,fun(X505,bool)))),X508),s(X506,X510))),s(X507,X509))),s(X504,X511))),s(X505,X512)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X504,fun(X505,bool)),happ(s(fun(X506,fun(X504,fun(X505,bool))),'_dst_x1x0_2'(s(fun(X506,fun(X507,fun(X504,fun(X505,bool)))),X508),s(X507,X509))),s(X506,X510))),s('type.pair.prod'(X504,X505),sK25(happ(s(fun(X506,fun(X504,fun(X505,bool))),'_dst_x1x0_2'(s(fun(X506,fun(X507,fun(X504,fun(X505,bool)))),X508),s(X507,X509))),s(X506,X510)),X505,X504)))))) )),
  inference(superposition,[],[f374,f346])).
fof(f7742,plain,(
  ( ! [X499,X500,X502,X496,X498,X495,X501,X503,X497] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X495,X496),bool),happ(s(fun(X498,fun('type.pair.prod'(X495,X496),bool)),happ(s(fun(X497,fun(X498,fun('type.pair.prod'(X495,X496),bool))),X499),s(X497,X501))),s(X498,X500))),s(X495,X502),s(X496,X503)))) | p(s(bool,happ(s(fun('type.pair.prod'(X495,X496),bool),happ(s(fun(X497,fun('type.pair.prod'(X495,X496),bool)),'_dst_x1x0_2'(s(fun(X497,fun(X498,fun('type.pair.prod'(X495,X496),bool))),X499),s(X498,X500))),s(X497,X501))),s('type.pair.prod'(X495,X496),sK22(happ(s(fun(X497,fun('type.pair.prod'(X495,X496),bool)),'_dst_x1x0_2'(s(fun(X497,fun(X498,fun('type.pair.prod'(X495,X496),bool))),X499),s(X498,X500))),s(X497,X501)),X496,X495)))))) )),
  inference(superposition,[],[f375,f346])).
fof(f7741,plain,(
  ( ! [X486,X492,X494,X488,X490,X487,X493,X489,X491] : (~p(s(bool,happ(s(fun('type.pair.prod'(X486,X487),bool),happ(s(fun(X489,fun('type.pair.prod'(X486,X487),bool)),happ(s(fun(X488,fun(X489,fun('type.pair.prod'(X486,X487),bool))),X490),s(X488,X492))),s(X489,X491))),s('type.pair.prod'(X486,X487),sK12(happ(s(fun(X488,fun('type.pair.prod'(X486,X487),bool)),'_dst_x1x0_2'(s(fun(X488,fun(X489,fun('type.pair.prod'(X486,X487),bool))),X490),s(X489,X491))),s(X488,X492)),X487,X486))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X486,X487),bool),happ(s(fun(X488,fun('type.pair.prod'(X486,X487),bool)),'_dst_x1x0_2'(s(fun(X488,fun(X489,fun('type.pair.prod'(X486,X487),bool))),X490),s(X489,X491))),s(X488,X492))),s(X486,X493),s(X487,X494))))) )),
  inference(superposition,[],[f376,f346])).
fof(f7740,plain,(
  ( ! [X478,X484,X480,X482,X479,X485,X481,X483] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X478,X479),bool),happ(s(fun(X481,fun('type.pair.prod'(X478,X479),bool)),happ(s(fun(X480,fun(X481,fun('type.pair.prod'(X478,X479),bool))),X482),s(X480,X484))),s(X481,X483))),s(X478,sK10(happ(s(fun(X480,fun('type.pair.prod'(X478,X479),bool)),'_dst_x1x0_2'(s(fun(X480,fun(X481,fun('type.pair.prod'(X478,X479),bool))),X482),s(X481,X483))),s(X480,X484)),X479,X478)),s(X479,sK11(happ(s(fun(X480,fun('type.pair.prod'(X478,X479),bool)),'_dst_x1x0_2'(s(fun(X480,fun(X481,fun('type.pair.prod'(X478,X479),bool))),X482),s(X481,X483))),s(X480,X484)),X479,X478))))) | p(s(bool,happ(s(fun('type.pair.prod'(X478,X479),bool),happ(s(fun(X480,fun('type.pair.prod'(X478,X479),bool)),'_dst_x1x0_2'(s(fun(X480,fun(X481,fun('type.pair.prod'(X478,X479),bool))),X482),s(X481,X483))),s(X480,X484))),s('type.pair.prod'(X478,X479),X485))))) )),
  inference(superposition,[],[f377,f346])).
fof(f7739,plain,(
  ( ! [X476,X472,X474,X471,X475,X477,X473,X470] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X470,X471),bool),happ(s(fun(X473,fun('type.pair.prod'(X470,X471),bool)),happ(s(fun(X472,fun(X473,fun('type.pair.prod'(X470,X471),bool))),X474),s(X472,X476))),s(X473,X475))),s(X470,sK13(happ(s(fun(X472,fun('type.pair.prod'(X470,X471),bool)),'_dst_x1x0_2'(s(fun(X472,fun(X473,fun('type.pair.prod'(X470,X471),bool))),X474),s(X473,X475))),s(X472,X476)),X471,X470)),s(X471,sK14(happ(s(fun(X472,fun('type.pair.prod'(X470,X471),bool)),'_dst_x1x0_2'(s(fun(X472,fun(X473,fun('type.pair.prod'(X470,X471),bool))),X474),s(X473,X475))),s(X472,X476)),X471,X470))))) | p(s(bool,happ(s(fun('type.pair.prod'(X470,X471),bool),happ(s(fun(X472,fun('type.pair.prod'(X470,X471),bool)),'_dst_x1x0_2'(s(fun(X472,fun(X473,fun('type.pair.prod'(X470,X471),bool))),X474),s(X473,X475))),s(X472,X476))),s('type.pair.prod'(X470,X471),X477))))) )),
  inference(superposition,[],[f378,f346])).
fof(f7738,plain,(
  ( ! [X466,X463,X469,X465,X467,X462,X468,X464] : (~p(s(bool,happ(s(fun('type.pair.prod'(X462,X463),bool),happ(s(fun(X465,fun('type.pair.prod'(X462,X463),bool)),happ(s(fun(X464,fun(X465,fun('type.pair.prod'(X462,X463),bool))),X466),s(X464,X468))),s(X465,X467))),s('type.pair.prod'(X462,X463),X469)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X462,X463),bool),happ(s(fun(X464,fun('type.pair.prod'(X462,X463),bool)),'_dst_x1x0_2'(s(fun(X464,fun(X465,fun('type.pair.prod'(X462,X463),bool))),X466),s(X465,X467))),s(X464,X468))),s(X462,sK20(happ(s(fun(X464,fun('type.pair.prod'(X462,X463),bool)),'_dst_x1x0_2'(s(fun(X464,fun(X465,fun('type.pair.prod'(X462,X463),bool))),X466),s(X465,X467))),s(X464,X468)),X463,X462)),s(X463,sK21(happ(s(fun(X464,fun('type.pair.prod'(X462,X463),bool)),'_dst_x1x0_2'(s(fun(X464,fun(X465,fun('type.pair.prod'(X462,X463),bool))),X466),s(X465,X467))),s(X464,X468)),X463,X462)))))) )),
  inference(superposition,[],[f379,f346])).
fof(f7737,plain,(
  ( ! [X455,X461,X457,X459,X460,X456,X458] : (~p(s(bool,happ(s(fun(X455,bool),happ(s(fun(X457,fun(X455,bool)),happ(s(fun(X456,fun(X457,fun(X455,bool))),X458),s(X456,X460))),s(X457,X459))),s(X455,X461)))) | p(s(bool,happ(s(fun(fun(X455,bool),bool),'?_1'),s(fun(X455,bool),happ(s(fun(X456,fun(X455,bool)),'_dst_x1x0_2'(s(fun(X456,fun(X457,fun(X455,bool))),X458),s(X457,X459))),s(X456,X460))))))) )),
  inference(superposition,[],[f303,f346])).
fof(f7736,plain,(
  ( ! [X453,X449,X451,X452,X454,X450] : (~p(s(bool,happ(s(fun(fun(X449,bool),bool),'?_1'),s(fun(X449,bool),happ(s(fun(X451,fun(X449,bool)),happ(s(fun(X450,fun(X451,fun(X449,bool))),X452),s(X450,X454))),s(X451,X453)))))) | p(s(bool,happ(s(fun(X449,bool),happ(s(fun(X450,fun(X449,bool)),'_dst_x1x0_2'(s(fun(X450,fun(X451,fun(X449,bool))),X452),s(X451,X453))),s(X450,X454))),s(X449,sK9(happ(s(fun(X450,fun(X449,bool)),'_dst_x1x0_2'(s(fun(X450,fun(X451,fun(X449,bool))),X452),s(X451,X453))),s(X450,X454)),X449)))))) )),
  inference(superposition,[],[f302,f346])).
fof(f7735,plain,(
  ( ! [X445,X447,X443,X444,X446,X448] : (~p(s(bool,happ(s(fun(X443,bool),happ(s(fun(X445,fun(X443,bool)),happ(s(fun(X444,fun(X445,fun(X443,bool))),X446),s(X444,X448))),s(X445,X447))),s(X443,sK8(happ(s(fun(X444,fun(X443,bool)),'_dst_x1x0_2'(s(fun(X444,fun(X445,fun(X443,bool))),X446),s(X445,X447))),s(X444,X448)),X443))))) | p(s(bool,happ(s(fun(fun(X443,bool),bool),'!_1'),s(fun(X443,bool),happ(s(fun(X444,fun(X443,bool)),'_dst_x1x0_2'(s(fun(X444,fun(X445,fun(X443,bool))),X446),s(X445,X447))),s(X444,X448))))))) )),
  inference(superposition,[],[f301,f346])).
fof(f7734,plain,(
  ( ! [X441,X436,X438,X440,X442,X437,X439] : (~p(s(bool,happ(s(fun(fun(X436,bool),bool),'!_1'),s(fun(X436,bool),happ(s(fun(X438,fun(X436,bool)),happ(s(fun(X437,fun(X438,fun(X436,bool))),X439),s(X437,X441))),s(X438,X440)))))) | p(s(bool,happ(s(fun(X436,bool),happ(s(fun(X437,fun(X436,bool)),'_dst_x1x0_2'(s(fun(X437,fun(X438,fun(X436,bool))),X439),s(X438,X440))),s(X437,X441))),s(X436,X442))))) )),
  inference(superposition,[],[f300,f346])).
fof(f7733,plain,(
  ( ! [X430,X432,X434,X429,X431,X433,X435] : (p(s(bool,'|dollar|exists_1'(s(fun(X429,bool),happ(s(fun(X431,fun(X429,bool)),happ(s(fun(X430,fun(X431,fun(X429,bool))),X432),s(X430,X434))),s(X431,X433)))))) | s(X429,X435) = s(X429,sK7(X435,happ(s(fun(X430,fun(X429,bool)),'_dst_x1x0_2'(s(fun(X430,fun(X431,fun(X429,bool))),X432),s(X431,X433))),s(X430,X434)),X429))) )),
  inference(superposition,[],[f292,f346])).
fof(f7732,plain,(
  ( ! [X422,X428,X424,X426,X423,X425,X427] : (~p(s(bool,happ(s(fun(X422,bool),happ(s(fun(X424,fun(X422,bool)),happ(s(fun(X423,fun(X424,fun(X422,bool))),X425),s(X423,X427))),s(X424,X426))),s(X422,X428)))) | p(s(bool,'|dollar|exists_1'(s(fun(X422,bool),happ(s(fun(X423,fun(X422,bool)),'_dst_x1x0_2'(s(fun(X423,fun(X424,fun(X422,bool))),X425),s(X424,X426))),s(X423,X427))))))) )),
  inference(superposition,[],[f289,f346])).
fof(f7731,plain,(
  ( ! [X420,X416,X418,X421,X417,X419] : (p(s(bool,happ(s(fun(X416,bool),happ(s(fun(X418,fun(X416,bool)),happ(s(fun(X417,fun(X418,fun(X416,bool))),X419),s(X417,X421))),s(X418,X420))),s(X416,sK6(happ(s(fun(X417,fun(X416,bool)),'_dst_x1x0_2'(s(fun(X417,fun(X418,fun(X416,bool))),X419),s(X418,X420))),s(X417,X421)),X416))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X416,bool),happ(s(fun(X417,fun(X416,bool)),'_dst_x1x0_2'(s(fun(X417,fun(X418,fun(X416,bool))),X419),s(X418,X420))),s(X417,X421))))))) )),
  inference(superposition,[],[f288,f346])).
fof(f7730,plain,(
  ( ! [X412,X414,X410,X413,X415,X411] : (~p(s(bool,'|dollar|exists_1'(s(fun(X410,bool),happ(s(fun(X412,fun(X410,bool)),happ(s(fun(X411,fun(X412,fun(X410,bool))),X413),s(X411,X415))),s(X412,X414)))))) | p(s(bool,happ(s(fun(X410,bool),happ(s(fun(X411,fun(X410,bool)),'_dst_x1x0_2'(s(fun(X411,fun(X412,fun(X410,bool))),X413),s(X412,X414))),s(X411,X415))),s(X410,sK5(happ(s(fun(X411,fun(X410,bool)),'_dst_x1x0_2'(s(fun(X411,fun(X412,fun(X410,bool))),X413),s(X412,X414))),s(X411,X415)),X410))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X410,bool),happ(s(fun(X411,fun(X410,bool)),'_dst_x1x0_2'(s(fun(X411,fun(X412,fun(X410,bool))),X413),s(X412,X414))),s(X411,X415))))))) )),
  inference(superposition,[],[f286,f346])).
fof(f7729,plain,(
  ( ! [X408,X405,X407,X409,X404,X406] : (~p(s(bool,'|dollar|exists_1'(s(fun(X404,bool),happ(s(fun(X406,fun(X404,bool)),happ(s(fun(X405,fun(X406,fun(X404,bool))),X407),s(X405,X409))),s(X406,X408)))))) | p(s(bool,happ(s(fun(X404,bool),happ(s(fun(X405,fun(X404,bool)),'_dst_x1x0_2'(s(fun(X405,fun(X406,fun(X404,bool))),X407),s(X406,X408))),s(X405,X409))),s(X404,sK4(happ(s(fun(X405,fun(X404,bool)),'_dst_x1x0_2'(s(fun(X405,fun(X406,fun(X404,bool))),X407),s(X406,X408))),s(X405,X409)),X404))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X404,bool),happ(s(fun(X405,fun(X404,bool)),'_dst_x1x0_2'(s(fun(X405,fun(X406,fun(X404,bool))),X407),s(X406,X408))),s(X405,X409))))))) )),
  inference(superposition,[],[f285,f346])).
fof(f7728,plain,(
  ( ! [X397,X399,X401,X403,X396,X398,X400,X402] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X396,bool),happ(s(fun(X398,fun(X396,bool)),happ(s(fun(X397,fun(X398,fun(X396,bool))),X399),s(X397,X401))),s(X398,X400)))))) | ~p(s(bool,happ(s(fun(X396,bool),happ(s(fun(X397,fun(X396,bool)),'_dst_x1x0_2'(s(fun(X397,fun(X398,fun(X396,bool))),X399),s(X398,X400))),s(X397,X401))),s(X396,X402)))) | ~p(s(bool,happ(s(fun(X396,bool),happ(s(fun(X397,fun(X396,bool)),'_dst_x1x0_2'(s(fun(X397,fun(X398,fun(X396,bool))),X399),s(X398,X400))),s(X397,X401))),s(X396,X403)))) | s(X396,X402) = s(X396,X403)) )),
  inference(superposition,[],[f284,f346])).
fof(f7727,plain,(
  ( ! [X391,X393,X395,X390,X392,X394] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X390,bool),happ(s(fun(X392,fun(X390,bool)),happ(s(fun(X391,fun(X392,fun(X390,bool))),X393),s(X391,X395))),s(X392,X394)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X390,bool),happ(s(fun(X391,fun(X390,bool)),'_dst_x1x0_2'(s(fun(X391,fun(X392,fun(X390,bool))),X393),s(X392,X394))),s(X391,X395))))))) )),
  inference(superposition,[],[f283,f346])).
fof(f7725,plain,(
  ( ! [X383,X377,X379,X380,X382,X376,X378,X381] : (s(X377,happ(s(fun(X376,X377),happ(s(fun(X379,fun(X376,X377)),happ(s(fun(X378,fun(X379,fun(X376,X377))),X380),s(X378,X382))),s(X379,X381))),s(X376,X383))) = s(X377,'const.pair.UNCURRY_2'(s(fun(X378,fun(X376,X377)),'_dst_x1x0_2'(s(fun(X378,fun(X379,fun(X376,X377))),X380),s(X379,X381))),s('type.pair.prod'(X378,X376),'const.pair.,_2'(s(X378,X382),s(X376,X383)))))) )),
  inference(superposition,[],[f345,f346])).
fof(f7724,plain,(
  ( ! [X372,X374,X370,X373,X375,X369,X371] : (~p(s(bool,happ(s(fun(X369,bool),happ(s(fun(X371,fun(X369,bool)),happ(s(fun(X370,fun(X371,fun(X369,bool))),X372),s(X370,X374))),s(X371,X373))),s(X369,X375)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X370,X369),bool),'const.pair.UNCURRY_1'(s(fun(X370,fun(X369,bool)),'_dst_x1x0_2'(s(fun(X370,fun(X371,fun(X369,bool))),X372),s(X371,X373))))))))) )),
  inference(superposition,[],[f319,f346])).
fof(f7723,plain,(
  ( ! [X364,X366,X362,X368,X365,X367,X363] : (~p(s(bool,happ(s(fun(X362,bool),happ(s(fun(X364,fun(X362,bool)),happ(s(fun(X363,fun(X364,fun(X362,bool))),X365),s(X363,X367))),s(X364,X366))),s(X362,X368)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X363,fun(X362,bool)),'_dst_x1x0_2'(s(fun(X363,fun(X364,fun(X362,bool))),X365),s(X364,X366))),s('type.pair.prod'(X363,X362),sK25('_dst_x1x0_2'(s(fun(X363,fun(X364,fun(X362,bool))),X365),s(X364,X366)),X362,X363)))))) )),
  inference(superposition,[],[f374,f346])).
fof(f7722,plain,(
  ( ! [X356,X358,X360,X357,X359,X361] : (p(s(bool,happ(s(fun(X357,bool),happ(s(fun(X356,fun(X357,bool)),X358),s(X356,X360))),s(X357,X359)))) | s(bool,X361) = s(bool,happ(s(fun(X356,bool),'_dst_x1x0_2'(s(fun(X356,fun(X357,bool)),X358),s(X357,X359))),s(X356,X360))) | s(bool,'T_0') = s(bool,X361)) )),
  inference(superposition,[],[f468,f346])).
fof(f7721,plain,(
  ( ! [X350,X352,X354,X351,X353,X355] : (~p(s(bool,happ(s(fun(X351,bool),happ(s(fun(X350,fun(X351,bool)),X352),s(X350,X354))),s(X351,X353)))) | ~p(s(bool,X355)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun(X350,bool),'_dst_x1x0_2'(s(fun(X350,fun(X351,bool)),X352),s(X351,X353))),s(X350,X354))))),s(bool,X355))))) )),
  inference(superposition,[],[f306,f346])).
fof(f7720,plain,(
  ( ! [X344,X346,X349,X345,X347,X348] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X349))),s(bool,happ(s(fun(X345,bool),happ(s(fun(X344,fun(X345,bool)),X346),s(X344,X348))),s(X345,X347)))))) | p(s(bool,happ(s(fun(X344,bool),'_dst_x1x0_2'(s(fun(X344,fun(X345,bool)),X346),s(X345,X347))),s(X344,X348))))) )),
  inference(superposition,[],[f305,f346])).
fof(f7717,plain,(
  ( ! [X327,X329,X331,X326,X328,X330] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun(X327,bool),happ(s(fun(X326,fun(X327,bool)),X328),s(X326,X330))),s(X327,X329))))),s(bool,X331)))) | p(s(bool,happ(s(fun(X326,bool),'_dst_x1x0_2'(s(fun(X326,fun(X327,bool)),X328),s(X327,X329))),s(X326,X330))))) )),
  inference(superposition,[],[f304,f346])).
fof(f7716,plain,(
  ( ! [X325,X321,X323,X324,X320,X322] : (~p(s(bool,happ(s(fun(X321,bool),happ(s(fun(X320,fun(X321,bool)),X322),s(X320,X324))),s(X321,X323)))) | ~p(s(bool,X325)) | s(bool,X325) = s(bool,happ(s(fun(X320,bool),'_dst_x1x0_2'(s(fun(X320,fun(X321,bool)),X322),s(X321,X323))),s(X320,X324)))) )),
  inference(superposition,[],[f299,f346])).
fof(f7714,plain,(
  ( ! [X313,X308,X310,X312,X309,X311] : (p(s(bool,happ(s(fun(X309,bool),happ(s(fun(X308,fun(X309,bool)),X310),s(X308,X312))),s(X309,X311)))) | p(s(bool,X313)) | s(bool,X313) = s(bool,happ(s(fun(X308,bool),'_dst_x1x0_2'(s(fun(X308,fun(X309,bool)),X310),s(X309,X311))),s(X308,X312)))) )),
  inference(superposition,[],[f296,f346])).
fof(f7713,plain,(
  ( ! [X304,X306,X303,X305,X307] : (p(s(bool,happ(s(fun(X304,bool),happ(s(fun(X303,fun(X304,bool)),X305),s(X303,X307))),s(X304,X306)))) | s(bool,'F_0') = s(bool,happ(s(fun(X303,bool),'_dst_x1x0_2'(s(fun(X303,fun(X304,bool)),X305),s(X304,X306))),s(X303,X307)))) )),
  inference(superposition,[],[f270,f346])).
fof(f7712,plain,(
  ( ! [X300,X302,X298,X301,X299] : (~p(s(bool,happ(s(fun(X299,bool),happ(s(fun(X298,fun(X299,bool)),X300),s(X298,X302))),s(X299,X301)))) | s(bool,'F_0') != s(bool,happ(s(fun(X298,bool),'_dst_x1x0_2'(s(fun(X298,fun(X299,bool)),X300),s(X299,X301))),s(X298,X302)))) )),
  inference(superposition,[],[f269,f346])).
fof(f7711,plain,(
  ( ! [X294,X296,X293,X295,X297] : (~p(s(bool,happ(s(fun(X294,bool),happ(s(fun(X293,fun(X294,bool)),X295),s(X293,X297))),s(X294,X296)))) | s(bool,'T_0') = s(bool,happ(s(fun(X293,bool),'_dst_x1x0_2'(s(fun(X293,fun(X294,bool)),X295),s(X294,X296))),s(X293,X297)))) )),
  inference(superposition,[],[f266,f346])).
fof(f7710,plain,(
  ( ! [X292,X288,X290,X289,X291] : (s(bool,'T_0') != s(bool,happ(s(fun(X289,bool),happ(s(fun(X288,fun(X289,bool)),X290),s(X288,X292))),s(X289,X291))) | p(s(bool,happ(s(fun(X288,bool),'_dst_x1x0_2'(s(fun(X288,fun(X289,bool)),X290),s(X289,X291))),s(X288,X292))))) )),
  inference(superposition,[],[f265,f346])).
fof(f7708,plain,(
  ( ! [X280,X282,X277,X279,X281,X278] : (~p(s(bool,happ(s(fun(X279,bool),happ(s(fun('type.pair.prod'(X277,X278),fun(X279,bool)),X280),s('type.pair.prod'(X277,X278),X282))),s(X279,X281)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X277,X278),bool),'_dst_x1x0_2'(s(fun('type.pair.prod'(X277,X278),fun(X279,bool)),X280),s(X279,X281))),s(X277,sK20('_dst_x1x0_2'(s(fun('type.pair.prod'(X277,X278),fun(X279,bool)),X280),s(X279,X281)),X278,X277)),s(X278,sK21('_dst_x1x0_2'(s(fun('type.pair.prod'(X277,X278),fun(X279,bool)),X280),s(X279,X281)),X278,X277)))))) )),
  inference(superposition,[],[f379,f346])).
fof(f7707,plain,(
  ( ! [X271,X273,X275,X270,X276,X272,X274] : (~p(s(bool,happ(s(fun(X272,bool),happ(s(fun('type.pair.prod'(X270,X271),fun(X272,bool)),X273),s('type.pair.prod'(X270,X271),sK12('_dst_x1x0_2'(s(fun('type.pair.prod'(X270,X271),fun(X272,bool)),X273),s(X272,X274)),X271,X270)))),s(X272,X274)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X270,X271),bool),'_dst_x1x0_2'(s(fun('type.pair.prod'(X270,X271),fun(X272,bool)),X273),s(X272,X274))),s(X270,X275),s(X271,X276))))) )),
  inference(superposition,[],[f376,f346])).
fof(f7706,plain,(
  ( ! [X269,X267,X268,X266] : (~p(s(bool,happ(s(fun(X267,bool),happ(s(fun(X266,fun(X267,bool)),X268),s(X266,sK8('_dst_x1x0_2'(s(fun(X266,fun(X267,bool)),X268),s(X267,X269)),X266)))),s(X267,X269)))) | p(s(bool,happ(s(fun(fun(X266,bool),bool),'!_1'),s(fun(X266,bool),'_dst_x1x0_2'(s(fun(X266,fun(X267,bool)),X268),s(X267,X269))))))) )),
  inference(superposition,[],[f301,f346])).
fof(f7705,plain,(
  ( ! [X263,X265,X262,X264] : (p(s(bool,happ(s(fun(X263,bool),happ(s(fun(X262,fun(X263,bool)),X264),s(X262,sK6('_dst_x1x0_2'(s(fun(X262,fun(X263,bool)),X264),s(X263,X265)),X262)))),s(X263,X265)))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X262,bool),'_dst_x1x0_2'(s(fun(X262,fun(X263,bool)),X264),s(X263,X265))))))) )),
  inference(superposition,[],[f288,f346])).
fof(f7703,plain,(
  ( ! [X255,X253,X257,X254,X256] : (~p(s(bool,happ(s(fun(X254,bool),happ(s(fun(X253,fun(X254,bool)),X255),s(X253,X257))),s(X254,X256)))) | p(s(bool,'|dollar|exists_1'(s(fun(X253,bool),'_dst_x1x0_2'(s(fun(X253,fun(X254,bool)),X255),s(X254,X256))))))) )),
  inference(superposition,[],[f289,f346])).
fof(f7702,plain,(
  ( ! [X251,X249,X252,X250,X248] : (~p(s(bool,happ(s(fun(X249,bool),happ(s(fun(X248,fun(X249,bool)),X250),s(X248,X252))),s(X249,X251)))) | p(s(bool,happ(s(fun(fun(X248,bool),bool),'?_1'),s(fun(X248,bool),'_dst_x1x0_2'(s(fun(X248,fun(X249,bool)),X250),s(X249,X251))))))) )),
  inference(superposition,[],[f303,f346])).
fof(f7692,plain,(
  ( ! [X158,X156,X161,X159,X157,X155,X160] : (p(s(bool,happ(s(fun(X155,bool),happ(s(fun(X155,fun(X155,bool)),'=_0'),s(X155,X161))),s(X155,happ(s(fun(X157,X155),happ(s(fun(X156,fun(X157,X155)),X158),s(X156,X160))),s(X157,X159)))))) | s(X155,X161) != s(X155,happ(s(fun(X156,X155),'_dst_x1x0_2'(s(fun(X156,fun(X157,X155)),X158),s(X157,X159))),s(X156,X160)))) )),
  inference(superposition,[],[f331,f346])).
fof(f7691,plain,(
  ( ! [X154,X152,X151,X149,X153,X150,X148] : (p(s(bool,happ(s(fun(X148,bool),happ(s(fun(X148,fun(X148,bool)),'=_0'),s(X148,happ(s(fun(X150,X148),happ(s(fun(X149,fun(X150,X148)),X151),s(X149,X153))),s(X150,X152))))),s(X148,X154)))) | s(X148,X154) != s(X148,happ(s(fun(X149,X148),'_dst_x1x0_2'(s(fun(X149,fun(X150,X148)),X151),s(X150,X152))),s(X149,X153)))) )),
  inference(superposition,[],[f331,f346])).
fof(f7690,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141] : (~p(s(bool,happ(s(fun(X141,bool),happ(s(fun(X141,fun(X141,bool)),'=_0'),s(X141,X147))),s(X141,happ(s(fun(X143,X141),happ(s(fun(X142,fun(X143,X141)),X144),s(X142,X146))),s(X143,X145)))))) | s(X141,X147) = s(X141,happ(s(fun(X142,X141),'_dst_x1x0_2'(s(fun(X142,fun(X143,X141)),X144),s(X143,X145))),s(X142,X146)))) )),
  inference(superposition,[],[f330,f346])).
fof(f7689,plain,(
  ( ! [X134,X140,X138,X136,X135,X139,X137] : (~p(s(bool,happ(s(fun(X134,bool),happ(s(fun(X134,fun(X134,bool)),'=_0'),s(X134,happ(s(fun(X136,X134),happ(s(fun(X135,fun(X136,X134)),X137),s(X135,X139))),s(X136,X138))))),s(X134,X140)))) | s(X134,X140) = s(X134,happ(s(fun(X135,X134),'_dst_x1x0_2'(s(fun(X135,fun(X136,X134)),X137),s(X136,X138))),s(X135,X139)))) )),
  inference(superposition,[],[f330,f346])).
fof(f7688,plain,(
  ( ! [X127,X132,X130,X128,X126,X133,X131,X129] : (s(X126,sK34(X126,X132)) != s(X126,happ(s(fun(X128,X126),happ(s(fun(X127,fun(X128,X126)),X129),s(X127,X131))),s(X128,X130))) | p(s(bool,happ(s(fun(X126,bool),happ(s(fun(X132,fun(X126,bool)),sK32(X126,X132)),s(X132,X133))),s(X126,happ(s(fun(X127,X126),'_dst_x1x0_2'(s(fun(X127,fun(X128,X126)),X129),s(X128,X130))),s(X127,X131)))))) | s(X132,X133) != s(X132,sK33(X126,X132))) )),
  inference(superposition,[],[f329,f346])).
fof(f7687,plain,(
  ( ! [X125,X123,X121,X118,X124,X122,X120,X119] : (~p(s(bool,happ(s(fun(X118,bool),happ(s(fun(X124,fun(X118,bool)),sK32(X118,X124)),s(X124,X125))),s(X118,happ(s(fun(X120,X118),happ(s(fun(X119,fun(X120,X118)),X121),s(X119,X123))),s(X120,X122)))))) | s(X118,sK34(X118,X124)) = s(X118,happ(s(fun(X119,X118),'_dst_x1x0_2'(s(fun(X119,fun(X120,X118)),X121),s(X120,X122))),s(X119,X123)))) )),
  inference(superposition,[],[f328,f346])).
fof(f7684,plain,(
  ( ! [X94,X101,X99,X97,X95,X100,X98,X96] : (~p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X94,fun(X100,bool)),sK32(X100,X94)),s(X94,happ(s(fun(X96,X94),happ(s(fun(X95,fun(X96,X94)),X97),s(X95,X99))),s(X96,X98))))),s(X100,X101)))) | s(X94,sK33(X100,X94)) = s(X94,happ(s(fun(X95,X94),'_dst_x1x0_2'(s(fun(X95,fun(X96,X94)),X97),s(X96,X98))),s(X95,X99)))) )),
  inference(superposition,[],[f327,f346])).
fof(f7673,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(X14,happ(s(fun(X8,X14),X15),s(X8,happ(s(fun(X10,X8),happ(s(fun(X9,fun(X10,X8)),X11),s(X9,X13))),s(X10,X12))))) = s(X14,happ(s(fun(X9,X14),'_dst_x0x1_2'(s(fun(X8,X14),X15),s(fun(X9,X8),'_dst_x1x0_2'(s(fun(X9,fun(X10,X8)),X11),s(X10,X12))))),s(X9,X13)))) )),
  inference(superposition,[],[f344,f346])).
fof(f7671,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),'_dst_x1x0_2'(s(fun(X35,fun(X34,bool)),X36),s(X34,X38))),s(X35,X37))) | s(bool,'T_0') = s(bool,happ(s(fun(X34,bool),happ(s(fun(X35,fun(X34,bool)),X36),s(X35,X37))),s(X34,X38)))) )),
  inference(superposition,[],[f346,f264])).
fof(f7669,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'|dollar|exists_1'(s(fun(X28,bool),happ(s(fun(X29,fun(X28,bool)),X30),s(X29,X31))))) = s(bool,happ(s(fun(X29,bool),'_dst_x1x0_2'(s(fun(X29,fun(X28,bool)),X30),s(X28,'const.min.@_1'(s(fun(X28,bool),happ(s(fun(X29,fun(X28,bool)),X30),s(X29,X31))))))),s(X29,X31)))) )),
  inference(superposition,[],[f346,f280])).
fof(f7668,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (s(X20,happ(s(fun(X22,X20),'_dst_x1x0_2'(s(fun(X22,fun(X21,X20)),X23),s(X21,happ(s(fun(X25,X21),X26),s(X25,X27))))),s(X22,X24))) = s(X20,happ(s(fun(X25,X20),'_dst_x0x1_2'(s(fun(X21,X20),happ(s(fun(X22,fun(X21,X20)),X23),s(X22,X24))),s(fun(X25,X21),X26))),s(X25,X27)))) )),
  inference(superposition,[],[f346,f344])).
fof(f7667,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s(X14,happ(s(fun(X16,X14),'_dst_x1x0_2'(s(fun(X16,fun(X15,X14)),X17),s(X15,X19))),s(X16,X18))) = s(X14,'const.pair.UNCURRY_2'(s(fun(X16,fun(X15,X14)),X17),s('type.pair.prod'(X16,X15),'const.pair.,_2'(s(X16,X18),s(X15,X19)))))) )),
  inference(superposition,[],[f346,f345])).
fof(f7666,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),happ(s(fun(X3,fun('type.pair.prod'(X1,X2),X0)),X4),s(X3,X5))),s(X1,X6),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),'_dst_x1x0_2'(s(fun(X3,fun('type.pair.prod'(X1,X2),X0)),X4),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X6),s(X2,X7))))),s(X3,X5)))) )),
  inference(superposition,[],[f346,f343])).
fof(f7627,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s(X52,happ(s(fun(X51,X52),happ(s(fun(X54,fun(X51,X52)),happ(s(fun(X53,fun(X54,fun(X51,X52))),X55),s(X53,X57))),s(X54,X56))),s(X51,X58))) = s(X52,happ(s(fun(X53,X52),'_dst_x1x0_2'(s(fun(X53,fun(X51,X52)),'_dst_x1x0_2'(s(fun(X53,fun(X54,fun(X51,X52))),X55),s(X54,X56))),s(X51,X58))),s(X53,X57)))) )),
  inference(superposition,[],[f346,f346])).
fof(f7626,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X49] : (s(X44,happ(s(fun(X43,X44),happ(s(fun(X46,fun(X43,X44)),X47),s(X46,happ(s(fun(X45,X46),X48),s(X45,X49))))),s(X43,X50))) = s(X44,happ(s(fun(X45,X44),'_dst_x1x0_2'(s(fun(X45,fun(X43,X44)),'_dst_x0x1_2'(s(fun(X46,fun(X43,X44)),X47),s(fun(X45,X46),X48))),s(X43,X50))),s(X45,X49)))) )),
  inference(superposition,[],[f346,f344])).
fof(f7625,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X42,X40] : (s(X36,happ(s(fun(X35,X36),'const.pair.UNCURRY_2'(s(fun(X38,fun(X37,fun(X35,X36))),X39),s('type.pair.prod'(X38,X37),'const.pair.,_2'(s(X38,X40),s(X37,X41))))),s(X35,X42))) = s(X36,happ(s(fun(X37,X36),'_dst_x1x0_2'(s(fun(X37,fun(X35,X36)),happ(s(fun(X38,fun(X37,fun(X35,X36))),X39),s(X38,X40))),s(X35,X42))),s(X37,X41)))) )),
  inference(superposition,[],[f346,f345])).
fof(f7624,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X34,X32] : (s(X28,happ(s(fun(X27,X28),happ(s(fun(X30,fun(X27,X28)),'_dst_x1x0_2'(s(fun(X30,fun(X29,fun(X27,X28))),X31),s(X29,X33))),s(X30,X32))),s(X27,X34))) = s(X28,happ(s(fun(X29,X28),'_dst_x1x0_2'(s(fun(X29,fun(X27,X28)),happ(s(fun(X30,fun(X29,fun(X27,X28))),X31),s(X30,X32))),s(X27,X34))),s(X29,X33)))) )),
  inference(superposition,[],[f346,f346])).
fof(f7623,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (s(X20,happ(s(fun(X21,X20),'_dst_x1x0_2'(s(fun(X21,fun(X19,X20)),X22),s(X19,X26))),s(X21,happ(s(fun(X23,X21),X24),s(X23,X25))))) = s(X20,happ(s(fun(X19,X20),happ(s(fun(X23,fun(X19,X20)),'_dst_x0x1_2'(s(fun(X21,fun(X19,X20)),X22),s(fun(X23,X21),X24))),s(X23,X25))),s(X19,X26)))) )),
  inference(superposition,[],[f346,f344])).
fof(f7622,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X18,X16] : (s(X12,happ(s(fun(X11,X12),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X13,X14),fun(X11,X12)),X15),s(X13,X16),s(X14,X17))),s(X11,X18))) = s(X12,happ(s(fun('type.pair.prod'(X13,X14),X12),'_dst_x1x0_2'(s(fun('type.pair.prod'(X13,X14),fun(X11,X12)),X15),s(X11,X18))),s('type.pair.prod'(X13,X14),'const.pair.,_2'(s(X13,X16),s(X14,X17)))))) )),
  inference(superposition,[],[f346,f343])).
fof(f7621,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s(X5,happ(s(fun(X4,X5),'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,fun(X4,X5))),X8),s('type.pair.prod'(X6,X7),X9))),s(X4,X10))) = s(X5,happ(s(fun('type.pair.prod'(X6,X7),X5),'_dst_x1x0_2'(s(fun('type.pair.prod'(X6,X7),fun(X4,X5)),'const.pair.UNCURRY_1'(s(fun(X6,fun(X7,fun(X4,X5))),X8))),s(X4,X10))),s('type.pair.prod'(X6,X7),X9)))) )),
  inference(superposition,[],[f346,f335])).
fof(f7567,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,happ(s(fun(X27,bool),happ(s(fun(X26,fun(X27,bool)),X28),s(X26,X30))),s(X27,X29))) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),'_dst_x1x0_2'(s(fun(X26,fun(X27,bool)),X28),s(X27,X29))),s(X26,X30)))) )),
  inference(superposition,[],[f346,f264])).
fof(f7566,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'|dollar|exists_1'(s(fun(X22,bool),'_dst_x1x0_2'(s(fun(X22,fun(X23,bool)),X24),s(X23,X25))))) = s(bool,happ(s(fun(X23,bool),happ(s(fun(X22,fun(X23,bool)),X24),s(X22,'const.min.@_1'(s(fun(X22,bool),'_dst_x1x0_2'(s(fun(X22,fun(X23,bool)),X24),s(X23,X25))))))),s(X23,X25)))) )),
  inference(superposition,[],[f346,f280])).
fof(f7565,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (s(X14,happ(s(fun(X16,X14),happ(s(fun(X15,fun(X16,X14)),X17),s(X15,happ(s(fun(X19,X15),X20),s(X19,X21))))),s(X16,X18))) = s(X14,happ(s(fun(X19,X14),'_dst_x0x1_2'(s(fun(X15,X14),'_dst_x1x0_2'(s(fun(X15,fun(X16,X14)),X17),s(X16,X18))),s(fun(X19,X15),X20))),s(X19,X21)))) )),
  inference(superposition,[],[f346,f344])).
fof(f7877,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),fun(X3,X0)),X4),s(X1,X6),s(X2,X7))),s(X3,X5))) = s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),'_dst_x1x0_2'(s(fun('type.pair.prod'(X1,X2),fun(X3,X0)),X4),s(X3,X5))),s(X1,X6),s(X2,X7)))) )),
  inference(forward_demodulation,[],[f7564,f343])).
fof(f7564,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),'_dst_x1x0_2'(s(fun('type.pair.prod'(X1,X2),fun(X3,X0)),X4),s(X3,X5))),s(X1,X6),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun('type.pair.prod'(X1,X2),fun(X3,X0)),X4),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X6),s(X2,X7))))),s(X3,X5)))) )),
  inference(superposition,[],[f346,f343])).
fof(f7562,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))))),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4))))))) )),
  inference(unit_resulting_resolution,[],[f346,f331])).
fof(f7563,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))),s(X1,X5))))),s(X0,happ(s(fun(X2,X0),'_dst_x1x0_2'(s(fun(X2,fun(X1,X0)),X3),s(X1,X5))),s(X2,X4))))))) )),
  inference(unit_resulting_resolution,[],[f346,f331])).
fof(f7458,plain,(
  ( ! [X524,X520,X522,X518,X521,X523,X519] : (s(X518,sK28('const.pair.UNCURRY_2'(s(fun(X520,fun(X521,'type.pair.prod'(X518,X519))),X522),s('type.pair.prod'(X520,X521),'const.pair.,_2'(s(X520,X523),s(X521,X524)))),X519,X518)) = s(X518,'const.pair.FST_1'(s('type.pair.prod'(X518,X519),happ(s(fun(X521,'type.pair.prod'(X518,X519)),happ(s(fun(X520,fun(X521,'type.pair.prod'(X518,X519))),X522),s(X520,X523))),s(X521,X524)))))) )),
  inference(superposition,[],[f701,f345])).
fof(f7457,plain,(
  ( ! [X511,X516,X512,X514,X517,X513,X515] : (s(X512,sK29('const.pair.UNCURRY_2'(s(fun(X513,fun(X514,'type.pair.prod'(X511,X512))),X515),s('type.pair.prod'(X513,X514),'const.pair.,_2'(s(X513,X516),s(X514,X517)))),X512,X511)) = s(X512,'const.pair.SND_1'(s('type.pair.prod'(X511,X512),happ(s(fun(X514,'type.pair.prod'(X511,X512)),happ(s(fun(X513,fun(X514,'type.pair.prod'(X511,X512))),X515),s(X513,X516))),s(X514,X517)))))) )),
  inference(superposition,[],[f702,f345])).
fof(f7448,plain,(
  ( ! [X430,X436,X438,X432,X434,X431,X437,X433,X435] : (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X430,X431),bool),'const.pair.UNCURRY_1'(s(fun(X430,fun(X431,bool)),happ(s(fun(X433,fun(X430,fun(X431,bool))),happ(s(fun(X432,fun(X433,fun(X430,fun(X431,bool)))),X434),s(X432,X435))),s(X433,X436)))))))) | p(s(bool,happ(s(fun(X431,bool),happ(s(fun(X430,fun(X431,bool)),'const.pair.UNCURRY_2'(s(fun(X432,fun(X433,fun(X430,fun(X431,bool)))),X434),s('type.pair.prod'(X432,X433),'const.pair.,_2'(s(X432,X435),s(X433,X436))))),s(X430,X437))),s(X431,X438))))) )),
  inference(superposition,[],[f314,f345])).
fof(f7447,plain,(
  ( ! [X422,X428,X424,X426,X421,X423,X429,X425,X427] : (~p(s(bool,happ(s(fun(X422,bool),happ(s(fun(X421,fun(X422,bool)),happ(s(fun(X424,fun(X421,fun(X422,bool))),happ(s(fun(X423,fun(X424,fun(X421,fun(X422,bool)))),X425),s(X423,X426))),s(X424,X427))),s(X421,X428))),s(X422,X429)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X421,X422),bool),'const.pair.UNCURRY_1'(s(fun(X421,fun(X422,bool)),'const.pair.UNCURRY_2'(s(fun(X423,fun(X424,fun(X421,fun(X422,bool)))),X425),s('type.pair.prod'(X423,X424),'const.pair.,_2'(s(X423,X426),s(X424,X427))))))))))) )),
  inference(superposition,[],[f319,f345])).
fof(f7446,plain,(
  ( ! [X412,X414,X420,X416,X418,X413,X415,X417,X419] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X412,fun(X413,bool)),happ(s(fun(X415,fun(X412,fun(X413,bool))),happ(s(fun(X414,fun(X415,fun(X412,fun(X413,bool)))),X416),s(X414,X417))),s(X415,X418))),s('type.pair.prod'(X412,X413),sK15('const.pair.UNCURRY_2'(s(fun(X414,fun(X415,fun(X412,fun(X413,bool)))),X416),s('type.pair.prod'(X414,X415),'const.pair.,_2'(s(X414,X417),s(X415,X418)))),X413,X412))))) | p(s(bool,happ(s(fun(X413,bool),happ(s(fun(X412,fun(X413,bool)),'const.pair.UNCURRY_2'(s(fun(X414,fun(X415,fun(X412,fun(X413,bool)))),X416),s('type.pair.prod'(X414,X415),'const.pair.,_2'(s(X414,X417),s(X415,X418))))),s(X412,X419))),s(X413,X420))))) )),
  inference(superposition,[],[f373,f345])).
fof(f7445,plain,(
  ( ! [X408,X410,X405,X407,X403,X409,X411,X404,X406] : (~p(s(bool,happ(s(fun(X404,bool),happ(s(fun(X403,fun(X404,bool)),happ(s(fun(X406,fun(X403,fun(X404,bool))),happ(s(fun(X405,fun(X406,fun(X403,fun(X404,bool)))),X407),s(X405,X408))),s(X406,X409))),s(X403,X410))),s(X404,X411)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X403,fun(X404,bool)),'const.pair.UNCURRY_2'(s(fun(X405,fun(X406,fun(X403,fun(X404,bool)))),X407),s('type.pair.prod'(X405,X406),'const.pair.,_2'(s(X405,X408),s(X406,X409))))),s('type.pair.prod'(X403,X404),sK25('const.pair.UNCURRY_2'(s(fun(X405,fun(X406,fun(X403,fun(X404,bool)))),X407),s('type.pair.prod'(X405,X406),'const.pair.,_2'(s(X405,X408),s(X406,X409)))),X404,X403)))))) )),
  inference(superposition,[],[f374,f345])).
fof(f7444,plain,(
  ( ! [X397,X399,X395,X401,X396,X398,X394,X400,X402] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X394,X395),bool),happ(s(fun(X397,fun('type.pair.prod'(X394,X395),bool)),happ(s(fun(X396,fun(X397,fun('type.pair.prod'(X394,X395),bool))),X398),s(X396,X399))),s(X397,X400))),s(X394,X401),s(X395,X402)))) | p(s(bool,happ(s(fun('type.pair.prod'(X394,X395),bool),'const.pair.UNCURRY_2'(s(fun(X396,fun(X397,fun('type.pair.prod'(X394,X395),bool))),X398),s('type.pair.prod'(X396,X397),'const.pair.,_2'(s(X396,X399),s(X397,X400))))),s('type.pair.prod'(X394,X395),sK22('const.pair.UNCURRY_2'(s(fun(X396,fun(X397,fun('type.pair.prod'(X394,X395),bool))),X398),s('type.pair.prod'(X396,X397),'const.pair.,_2'(s(X396,X399),s(X397,X400)))),X395,X394)))))) )),
  inference(superposition,[],[f375,f345])).
fof(f7443,plain,(
  ( ! [X389,X391,X385,X387,X393,X388,X390,X386,X392] : (~p(s(bool,happ(s(fun('type.pair.prod'(X385,X386),bool),happ(s(fun(X388,fun('type.pair.prod'(X385,X386),bool)),happ(s(fun(X387,fun(X388,fun('type.pair.prod'(X385,X386),bool))),X389),s(X387,X390))),s(X388,X391))),s('type.pair.prod'(X385,X386),sK12('const.pair.UNCURRY_2'(s(fun(X387,fun(X388,fun('type.pair.prod'(X385,X386),bool))),X389),s('type.pair.prod'(X387,X388),'const.pair.,_2'(s(X387,X390),s(X388,X391)))),X386,X385))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X385,X386),bool),'const.pair.UNCURRY_2'(s(fun(X387,fun(X388,fun('type.pair.prod'(X385,X386),bool))),X389),s('type.pair.prod'(X387,X388),'const.pair.,_2'(s(X387,X390),s(X388,X391))))),s(X385,X392),s(X386,X393))))) )),
  inference(superposition,[],[f376,f345])).
fof(f7442,plain,(
  ( ! [X383,X377,X379,X380,X382,X378,X384,X381] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X377,X378),bool),happ(s(fun(X380,fun('type.pair.prod'(X377,X378),bool)),happ(s(fun(X379,fun(X380,fun('type.pair.prod'(X377,X378),bool))),X381),s(X379,X382))),s(X380,X383))),s(X377,sK10('const.pair.UNCURRY_2'(s(fun(X379,fun(X380,fun('type.pair.prod'(X377,X378),bool))),X381),s('type.pair.prod'(X379,X380),'const.pair.,_2'(s(X379,X382),s(X380,X383)))),X378,X377)),s(X378,sK11('const.pair.UNCURRY_2'(s(fun(X379,fun(X380,fun('type.pair.prod'(X377,X378),bool))),X381),s('type.pair.prod'(X379,X380),'const.pair.,_2'(s(X379,X382),s(X380,X383)))),X378,X377))))) | p(s(bool,happ(s(fun('type.pair.prod'(X377,X378),bool),'const.pair.UNCURRY_2'(s(fun(X379,fun(X380,fun('type.pair.prod'(X377,X378),bool))),X381),s('type.pair.prod'(X379,X380),'const.pair.,_2'(s(X379,X382),s(X380,X383))))),s('type.pair.prod'(X377,X378),X384))))) )),
  inference(superposition,[],[f377,f345])).
fof(f7441,plain,(
  ( ! [X372,X374,X370,X376,X373,X375,X369,X371] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X369,X370),bool),happ(s(fun(X372,fun('type.pair.prod'(X369,X370),bool)),happ(s(fun(X371,fun(X372,fun('type.pair.prod'(X369,X370),bool))),X373),s(X371,X374))),s(X372,X375))),s(X369,sK13('const.pair.UNCURRY_2'(s(fun(X371,fun(X372,fun('type.pair.prod'(X369,X370),bool))),X373),s('type.pair.prod'(X371,X372),'const.pair.,_2'(s(X371,X374),s(X372,X375)))),X370,X369)),s(X370,sK14('const.pair.UNCURRY_2'(s(fun(X371,fun(X372,fun('type.pair.prod'(X369,X370),bool))),X373),s('type.pair.prod'(X371,X372),'const.pair.,_2'(s(X371,X374),s(X372,X375)))),X370,X369))))) | p(s(bool,happ(s(fun('type.pair.prod'(X369,X370),bool),'const.pair.UNCURRY_2'(s(fun(X371,fun(X372,fun('type.pair.prod'(X369,X370),bool))),X373),s('type.pair.prod'(X371,X372),'const.pair.,_2'(s(X371,X374),s(X372,X375))))),s('type.pair.prod'(X369,X370),X376))))) )),
  inference(superposition,[],[f378,f345])).
fof(f7440,plain,(
  ( ! [X364,X366,X362,X368,X365,X367,X361,X363] : (~p(s(bool,happ(s(fun('type.pair.prod'(X361,X362),bool),happ(s(fun(X364,fun('type.pair.prod'(X361,X362),bool)),happ(s(fun(X363,fun(X364,fun('type.pair.prod'(X361,X362),bool))),X365),s(X363,X366))),s(X364,X367))),s('type.pair.prod'(X361,X362),X368)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X361,X362),bool),'const.pair.UNCURRY_2'(s(fun(X363,fun(X364,fun('type.pair.prod'(X361,X362),bool))),X365),s('type.pair.prod'(X363,X364),'const.pair.,_2'(s(X363,X366),s(X364,X367))))),s(X361,sK20('const.pair.UNCURRY_2'(s(fun(X363,fun(X364,fun('type.pair.prod'(X361,X362),bool))),X365),s('type.pair.prod'(X363,X364),'const.pair.,_2'(s(X363,X366),s(X364,X367)))),X362,X361)),s(X362,sK21('const.pair.UNCURRY_2'(s(fun(X363,fun(X364,fun('type.pair.prod'(X361,X362),bool))),X365),s('type.pair.prod'(X363,X364),'const.pair.,_2'(s(X363,X366),s(X364,X367)))),X362,X361)))))) )),
  inference(superposition,[],[f379,f345])).
fof(f7439,plain,(
  ( ! [X356,X358,X354,X360,X357,X359,X355] : (~p(s(bool,happ(s(fun(X354,bool),happ(s(fun(X356,fun(X354,bool)),happ(s(fun(X355,fun(X356,fun(X354,bool))),X357),s(X355,X358))),s(X356,X359))),s(X354,X360)))) | p(s(bool,happ(s(fun(fun(X354,bool),bool),'?_1'),s(fun(X354,bool),'const.pair.UNCURRY_2'(s(fun(X355,fun(X356,fun(X354,bool))),X357),s('type.pair.prod'(X355,X356),'const.pair.,_2'(s(X355,X358),s(X356,X359))))))))) )),
  inference(superposition,[],[f303,f345])).
fof(f7438,plain,(
  ( ! [X350,X352,X349,X351,X353,X348] : (~p(s(bool,happ(s(fun(fun(X348,bool),bool),'?_1'),s(fun(X348,bool),happ(s(fun(X350,fun(X348,bool)),happ(s(fun(X349,fun(X350,fun(X348,bool))),X351),s(X349,X352))),s(X350,X353)))))) | p(s(bool,happ(s(fun(X348,bool),'const.pair.UNCURRY_2'(s(fun(X349,fun(X350,fun(X348,bool))),X351),s('type.pair.prod'(X349,X350),'const.pair.,_2'(s(X349,X352),s(X350,X353))))),s(X348,sK9('const.pair.UNCURRY_2'(s(fun(X349,fun(X350,fun(X348,bool))),X351),s('type.pair.prod'(X349,X350),'const.pair.,_2'(s(X349,X352),s(X350,X353)))),X348)))))) )),
  inference(superposition,[],[f302,f345])).
fof(f7437,plain,(
  ( ! [X344,X346,X343,X345,X347,X342] : (~p(s(bool,happ(s(fun(X342,bool),happ(s(fun(X344,fun(X342,bool)),happ(s(fun(X343,fun(X344,fun(X342,bool))),X345),s(X343,X346))),s(X344,X347))),s(X342,sK8('const.pair.UNCURRY_2'(s(fun(X343,fun(X344,fun(X342,bool))),X345),s('type.pair.prod'(X343,X344),'const.pair.,_2'(s(X343,X346),s(X344,X347)))),X342))))) | p(s(bool,happ(s(fun(fun(X342,bool),bool),'!_1'),s(fun(X342,bool),'const.pair.UNCURRY_2'(s(fun(X343,fun(X344,fun(X342,bool))),X345),s('type.pair.prod'(X343,X344),'const.pair.,_2'(s(X343,X346),s(X344,X347))))))))) )),
  inference(superposition,[],[f301,f345])).
fof(f7436,plain,(
  ( ! [X335,X341,X337,X339,X340,X336,X338] : (~p(s(bool,happ(s(fun(fun(X335,bool),bool),'!_1'),s(fun(X335,bool),happ(s(fun(X337,fun(X335,bool)),happ(s(fun(X336,fun(X337,fun(X335,bool))),X338),s(X336,X339))),s(X337,X340)))))) | p(s(bool,happ(s(fun(X335,bool),'const.pair.UNCURRY_2'(s(fun(X336,fun(X337,fun(X335,bool))),X338),s('type.pair.prod'(X336,X337),'const.pair.,_2'(s(X336,X339),s(X337,X340))))),s(X335,X341))))) )),
  inference(superposition,[],[f300,f345])).
fof(f7435,plain,(
  ( ! [X333,X329,X331,X332,X334,X328,X330] : (p(s(bool,'|dollar|exists_1'(s(fun(X328,bool),happ(s(fun(X330,fun(X328,bool)),happ(s(fun(X329,fun(X330,fun(X328,bool))),X331),s(X329,X332))),s(X330,X333)))))) | s(X328,X334) = s(X328,sK7(X334,'const.pair.UNCURRY_2'(s(fun(X329,fun(X330,fun(X328,bool))),X331),s('type.pair.prod'(X329,X330),'const.pair.,_2'(s(X329,X332),s(X330,X333)))),X328))) )),
  inference(superposition,[],[f292,f345])).
fof(f7434,plain,(
  ( ! [X325,X327,X321,X323,X324,X326,X322] : (~p(s(bool,happ(s(fun(X321,bool),happ(s(fun(X323,fun(X321,bool)),happ(s(fun(X322,fun(X323,fun(X321,bool))),X324),s(X322,X325))),s(X323,X326))),s(X321,X327)))) | p(s(bool,'|dollar|exists_1'(s(fun(X321,bool),'const.pair.UNCURRY_2'(s(fun(X322,fun(X323,fun(X321,bool))),X324),s('type.pair.prod'(X322,X323),'const.pair.,_2'(s(X322,X325),s(X323,X326))))))))) )),
  inference(superposition,[],[f289,f345])).
fof(f7433,plain,(
  ( ! [X315,X316,X318,X320,X317,X319] : (p(s(bool,happ(s(fun(X315,bool),happ(s(fun(X317,fun(X315,bool)),happ(s(fun(X316,fun(X317,fun(X315,bool))),X318),s(X316,X319))),s(X317,X320))),s(X315,sK6('const.pair.UNCURRY_2'(s(fun(X316,fun(X317,fun(X315,bool))),X318),s('type.pair.prod'(X316,X317),'const.pair.,_2'(s(X316,X319),s(X317,X320)))),X315))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X315,bool),'const.pair.UNCURRY_2'(s(fun(X316,fun(X317,fun(X315,bool))),X318),s('type.pair.prod'(X316,X317),'const.pair.,_2'(s(X316,X319),s(X317,X320))))))))) )),
  inference(superposition,[],[f288,f345])).
fof(f7432,plain,(
  ( ! [X313,X310,X312,X314,X309,X311] : (~p(s(bool,'|dollar|exists_1'(s(fun(X309,bool),happ(s(fun(X311,fun(X309,bool)),happ(s(fun(X310,fun(X311,fun(X309,bool))),X312),s(X310,X313))),s(X311,X314)))))) | p(s(bool,happ(s(fun(X309,bool),'const.pair.UNCURRY_2'(s(fun(X310,fun(X311,fun(X309,bool))),X312),s('type.pair.prod'(X310,X311),'const.pair.,_2'(s(X310,X313),s(X311,X314))))),s(X309,sK5('const.pair.UNCURRY_2'(s(fun(X310,fun(X311,fun(X309,bool))),X312),s('type.pair.prod'(X310,X311),'const.pair.,_2'(s(X310,X313),s(X311,X314)))),X309))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X309,bool),'const.pair.UNCURRY_2'(s(fun(X310,fun(X311,fun(X309,bool))),X312),s('type.pair.prod'(X310,X311),'const.pair.,_2'(s(X310,X313),s(X311,X314))))))))) )),
  inference(superposition,[],[f286,f345])).
fof(f7431,plain,(
  ( ! [X308,X304,X306,X303,X305,X307] : (~p(s(bool,'|dollar|exists_1'(s(fun(X303,bool),happ(s(fun(X305,fun(X303,bool)),happ(s(fun(X304,fun(X305,fun(X303,bool))),X306),s(X304,X307))),s(X305,X308)))))) | p(s(bool,happ(s(fun(X303,bool),'const.pair.UNCURRY_2'(s(fun(X304,fun(X305,fun(X303,bool))),X306),s('type.pair.prod'(X304,X305),'const.pair.,_2'(s(X304,X307),s(X305,X308))))),s(X303,sK4('const.pair.UNCURRY_2'(s(fun(X304,fun(X305,fun(X303,bool))),X306),s('type.pair.prod'(X304,X305),'const.pair.,_2'(s(X304,X307),s(X305,X308)))),X303))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X303,bool),'const.pair.UNCURRY_2'(s(fun(X304,fun(X305,fun(X303,bool))),X306),s('type.pair.prod'(X304,X305),'const.pair.,_2'(s(X304,X307),s(X305,X308))))))))) )),
  inference(superposition,[],[f285,f345])).
fof(f7430,plain,(
  ( ! [X300,X302,X296,X298,X295,X301,X297,X299] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X295,bool),happ(s(fun(X297,fun(X295,bool)),happ(s(fun(X296,fun(X297,fun(X295,bool))),X298),s(X296,X299))),s(X297,X300)))))) | ~p(s(bool,happ(s(fun(X295,bool),'const.pair.UNCURRY_2'(s(fun(X296,fun(X297,fun(X295,bool))),X298),s('type.pair.prod'(X296,X297),'const.pair.,_2'(s(X296,X299),s(X297,X300))))),s(X295,X301)))) | ~p(s(bool,happ(s(fun(X295,bool),'const.pair.UNCURRY_2'(s(fun(X296,fun(X297,fun(X295,bool))),X298),s('type.pair.prod'(X296,X297),'const.pair.,_2'(s(X296,X299),s(X297,X300))))),s(X295,X302)))) | s(X295,X301) = s(X295,X302)) )),
  inference(superposition,[],[f284,f345])).
fof(f7429,plain,(
  ( ! [X292,X294,X290,X293,X289,X291] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X289,bool),happ(s(fun(X291,fun(X289,bool)),happ(s(fun(X290,fun(X291,fun(X289,bool))),X292),s(X290,X293))),s(X291,X294)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X289,bool),'const.pair.UNCURRY_2'(s(fun(X290,fun(X291,fun(X289,bool))),X292),s('type.pair.prod'(X290,X291),'const.pair.,_2'(s(X290,X293),s(X291,X294))))))))) )),
  inference(superposition,[],[f283,f345])).
fof(f7406,plain,(
  ( ! [X134,X132,X136,X135,X133,X131,X137] : (p(s(bool,happ(s(fun(X131,bool),happ(s(fun(X131,fun(X131,bool)),'=_0'),s(X131,X137))),s(X131,happ(s(fun(X133,X131),happ(s(fun(X132,fun(X133,X131)),X134),s(X132,X135))),s(X133,X136)))))) | s(X131,X137) != s(X131,'const.pair.UNCURRY_2'(s(fun(X132,fun(X133,X131)),X134),s('type.pair.prod'(X132,X133),'const.pair.,_2'(s(X132,X135),s(X133,X136)))))) )),
  inference(superposition,[],[f331,f345])).
fof(f7405,plain,(
  ( ! [X127,X125,X130,X128,X126,X124,X129] : (p(s(bool,happ(s(fun(X124,bool),happ(s(fun(X124,fun(X124,bool)),'=_0'),s(X124,happ(s(fun(X126,X124),happ(s(fun(X125,fun(X126,X124)),X127),s(X125,X128))),s(X126,X129))))),s(X124,X130)))) | s(X124,X130) != s(X124,'const.pair.UNCURRY_2'(s(fun(X125,fun(X126,X124)),X127),s('type.pair.prod'(X125,X126),'const.pair.,_2'(s(X125,X128),s(X126,X129)))))) )),
  inference(superposition,[],[f331,f345])).
fof(f7404,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (~p(s(bool,happ(s(fun(X117,bool),happ(s(fun(X117,fun(X117,bool)),'=_0'),s(X117,X123))),s(X117,happ(s(fun(X119,X117),happ(s(fun(X118,fun(X119,X117)),X120),s(X118,X121))),s(X119,X122)))))) | s(X117,X123) = s(X117,'const.pair.UNCURRY_2'(s(fun(X118,fun(X119,X117)),X120),s('type.pair.prod'(X118,X119),'const.pair.,_2'(s(X118,X121),s(X119,X122)))))) )),
  inference(superposition,[],[f330,f345])).
fof(f7403,plain,(
  ( ! [X111,X116,X114,X112,X110,X115,X113] : (~p(s(bool,happ(s(fun(X110,bool),happ(s(fun(X110,fun(X110,bool)),'=_0'),s(X110,happ(s(fun(X112,X110),happ(s(fun(X111,fun(X112,X110)),X113),s(X111,X114))),s(X112,X115))))),s(X110,X116)))) | s(X110,X116) = s(X110,'const.pair.UNCURRY_2'(s(fun(X111,fun(X112,X110)),X113),s('type.pair.prod'(X111,X112),'const.pair.,_2'(s(X111,X114),s(X112,X115)))))) )),
  inference(superposition,[],[f330,f345])).
fof(f7402,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (s(X102,sK34(X102,X108)) != s(X102,happ(s(fun(X104,X102),happ(s(fun(X103,fun(X104,X102)),X105),s(X103,X106))),s(X104,X107))) | p(s(bool,happ(s(fun(X102,bool),happ(s(fun(X108,fun(X102,bool)),sK32(X102,X108)),s(X108,X109))),s(X102,'const.pair.UNCURRY_2'(s(fun(X103,fun(X104,X102)),X105),s('type.pair.prod'(X103,X104),'const.pair.,_2'(s(X103,X106),s(X104,X107)))))))) | s(X108,X109) != s(X108,sK33(X102,X108))) )),
  inference(superposition,[],[f329,f345])).
fof(f7401,plain,(
  ( ! [X94,X101,X99,X97,X95,X100,X98,X96] : (~p(s(bool,happ(s(fun(X94,bool),happ(s(fun(X100,fun(X94,bool)),sK32(X94,X100)),s(X100,X101))),s(X94,happ(s(fun(X96,X94),happ(s(fun(X95,fun(X96,X94)),X97),s(X95,X98))),s(X96,X99)))))) | s(X94,sK34(X94,X100)) = s(X94,'const.pair.UNCURRY_2'(s(fun(X95,fun(X96,X94)),X97),s('type.pair.prod'(X95,X96),'const.pair.,_2'(s(X95,X98),s(X96,X99)))))) )),
  inference(superposition,[],[f328,f345])).
fof(f7398,plain,(
  ( ! [X70,X76,X74,X72,X71,X77,X75,X73] : (~p(s(bool,happ(s(fun(X76,bool),happ(s(fun(X70,fun(X76,bool)),sK32(X76,X70)),s(X70,happ(s(fun(X72,X70),happ(s(fun(X71,fun(X72,X70)),X73),s(X71,X74))),s(X72,X75))))),s(X76,X77)))) | s(X70,sK33(X76,X70)) = s(X70,'const.pair.UNCURRY_2'(s(fun(X71,fun(X72,X70)),X73),s('type.pair.prod'(X71,X72),'const.pair.,_2'(s(X71,X74),s(X72,X75)))))) )),
  inference(superposition,[],[f327,f345])).
fof(f7388,plain,(
  ( ! [X617,X613,X615,X616,X618,X612,X614] : (s(X612,'const.pair.FST_1'(s('type.pair.prod'(X612,X613),'const.pair.UNCURRY_2'(s(fun(X615,fun(X614,'type.pair.prod'(X612,X613))),X616),s('type.pair.prod'(X615,X614),'const.pair.,_2'(s(X615,X617),s(X614,X618))))))) = s(X612,sK28(happ(s(fun(X614,'type.pair.prod'(X612,X613)),happ(s(fun(X615,fun(X614,'type.pair.prod'(X612,X613))),X616),s(X615,X617))),s(X614,X618)),X613,X612))) )),
  inference(superposition,[],[f701,f345])).
fof(f7387,plain,(
  ( ! [X606,X609,X605,X607,X611,X608,X610] : (s(X606,'const.pair.SND_1'(s('type.pair.prod'(X605,X606),'const.pair.UNCURRY_2'(s(fun(X608,fun(X607,'type.pair.prod'(X605,X606))),X609),s('type.pair.prod'(X608,X607),'const.pair.,_2'(s(X608,X610),s(X607,X611))))))) = s(X606,sK29(happ(s(fun(X607,'type.pair.prod'(X605,X606)),happ(s(fun(X608,fun(X607,'type.pair.prod'(X605,X606))),X609),s(X608,X610))),s(X607,X611)),X606,X605))) )),
  inference(superposition,[],[f702,f345])).
fof(f7378,plain,(
  ( ! [X532,X528,X530,X524,X526,X529,X531,X525,X527] : (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X524,X525),bool),'const.pair.UNCURRY_1'(s(fun(X524,fun(X525,bool)),'const.pair.UNCURRY_2'(s(fun(X527,fun(X526,fun(X524,fun(X525,bool)))),X528),s('type.pair.prod'(X527,X526),'const.pair.,_2'(s(X527,X529),s(X526,X530)))))))))) | p(s(bool,happ(s(fun(X525,bool),happ(s(fun(X524,fun(X525,bool)),happ(s(fun(X526,fun(X524,fun(X525,bool))),happ(s(fun(X527,fun(X526,fun(X524,fun(X525,bool)))),X528),s(X527,X529))),s(X526,X530))),s(X524,X531))),s(X525,X532))))) )),
  inference(superposition,[],[f314,f345])).
fof(f7377,plain,(
  ( ! [X520,X522,X516,X518,X521,X523,X517,X519,X515] : (~p(s(bool,happ(s(fun(X516,bool),happ(s(fun(X515,fun(X516,bool)),'const.pair.UNCURRY_2'(s(fun(X518,fun(X517,fun(X515,fun(X516,bool)))),X519),s('type.pair.prod'(X518,X517),'const.pair.,_2'(s(X518,X520),s(X517,X521))))),s(X515,X522))),s(X516,X523)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X515,X516),bool),'const.pair.UNCURRY_1'(s(fun(X515,fun(X516,bool)),happ(s(fun(X517,fun(X515,fun(X516,bool))),happ(s(fun(X518,fun(X517,fun(X515,fun(X516,bool)))),X519),s(X518,X520))),s(X517,X521))))))))) )),
  inference(superposition,[],[f319,f345])).
fof(f7376,plain,(
  ( ! [X509,X511,X507,X512,X514,X508,X510,X506,X513] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X506,fun(X507,bool)),'const.pair.UNCURRY_2'(s(fun(X509,fun(X508,fun(X506,fun(X507,bool)))),X510),s('type.pair.prod'(X509,X508),'const.pair.,_2'(s(X509,X511),s(X508,X512))))),s('type.pair.prod'(X506,X507),sK15(happ(s(fun(X508,fun(X506,fun(X507,bool))),happ(s(fun(X509,fun(X508,fun(X506,fun(X507,bool)))),X510),s(X509,X511))),s(X508,X512)),X507,X506))))) | p(s(bool,happ(s(fun(X507,bool),happ(s(fun(X506,fun(X507,bool)),happ(s(fun(X508,fun(X506,fun(X507,bool))),happ(s(fun(X509,fun(X508,fun(X506,fun(X507,bool)))),X510),s(X509,X511))),s(X508,X512))),s(X506,X513))),s(X507,X514))))) )),
  inference(superposition,[],[f373,f345])).
fof(f7375,plain,(
  ( ! [X499,X505,X500,X502,X498,X504,X501,X503,X497] : (~p(s(bool,happ(s(fun(X498,bool),happ(s(fun(X497,fun(X498,bool)),'const.pair.UNCURRY_2'(s(fun(X500,fun(X499,fun(X497,fun(X498,bool)))),X501),s('type.pair.prod'(X500,X499),'const.pair.,_2'(s(X500,X502),s(X499,X503))))),s(X497,X504))),s(X498,X505)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X497,fun(X498,bool)),happ(s(fun(X499,fun(X497,fun(X498,bool))),happ(s(fun(X500,fun(X499,fun(X497,fun(X498,bool)))),X501),s(X500,X502))),s(X499,X503))),s('type.pair.prod'(X497,X498),sK25(happ(s(fun(X499,fun(X497,fun(X498,bool))),happ(s(fun(X500,fun(X499,fun(X497,fun(X498,bool)))),X501),s(X500,X502))),s(X499,X503)),X498,X497)))))) )),
  inference(superposition,[],[f374,f345])).
fof(f7374,plain,(
  ( ! [X492,X494,X488,X490,X496,X493,X495,X489,X491] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X488,X489),bool),'const.pair.UNCURRY_2'(s(fun(X491,fun(X490,fun('type.pair.prod'(X488,X489),bool))),X492),s('type.pair.prod'(X491,X490),'const.pair.,_2'(s(X491,X493),s(X490,X494))))),s(X488,X495),s(X489,X496)))) | p(s(bool,happ(s(fun('type.pair.prod'(X488,X489),bool),happ(s(fun(X490,fun('type.pair.prod'(X488,X489),bool)),happ(s(fun(X491,fun(X490,fun('type.pair.prod'(X488,X489),bool))),X492),s(X491,X493))),s(X490,X494))),s('type.pair.prod'(X488,X489),sK22(happ(s(fun(X490,fun('type.pair.prod'(X488,X489),bool)),happ(s(fun(X491,fun(X490,fun('type.pair.prod'(X488,X489),bool))),X492),s(X491,X493))),s(X490,X494)),X489,X488)))))) )),
  inference(superposition,[],[f375,f345])).
fof(f7373,plain,(
  ( ! [X484,X486,X480,X482,X479,X485,X487,X481,X483] : (~p(s(bool,happ(s(fun('type.pair.prod'(X479,X480),bool),'const.pair.UNCURRY_2'(s(fun(X482,fun(X481,fun('type.pair.prod'(X479,X480),bool))),X483),s('type.pair.prod'(X482,X481),'const.pair.,_2'(s(X482,X484),s(X481,X485))))),s('type.pair.prod'(X479,X480),sK12(happ(s(fun(X481,fun('type.pair.prod'(X479,X480),bool)),happ(s(fun(X482,fun(X481,fun('type.pair.prod'(X479,X480),bool))),X483),s(X482,X484))),s(X481,X485)),X480,X479))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X479,X480),bool),happ(s(fun(X481,fun('type.pair.prod'(X479,X480),bool)),happ(s(fun(X482,fun(X481,fun('type.pair.prod'(X479,X480),bool))),X483),s(X482,X484))),s(X481,X485))),s(X479,X486),s(X480,X487))))) )),
  inference(superposition,[],[f376,f345])).
fof(f7372,plain,(
  ( ! [X476,X478,X472,X474,X471,X475,X477,X473] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X471,X472),bool),'const.pair.UNCURRY_2'(s(fun(X474,fun(X473,fun('type.pair.prod'(X471,X472),bool))),X475),s('type.pair.prod'(X474,X473),'const.pair.,_2'(s(X474,X476),s(X473,X477))))),s(X471,sK10(happ(s(fun(X473,fun('type.pair.prod'(X471,X472),bool)),happ(s(fun(X474,fun(X473,fun('type.pair.prod'(X471,X472),bool))),X475),s(X474,X476))),s(X473,X477)),X472,X471)),s(X472,sK11(happ(s(fun(X473,fun('type.pair.prod'(X471,X472),bool)),happ(s(fun(X474,fun(X473,fun('type.pair.prod'(X471,X472),bool))),X475),s(X474,X476))),s(X473,X477)),X472,X471))))) | p(s(bool,happ(s(fun('type.pair.prod'(X471,X472),bool),happ(s(fun(X473,fun('type.pair.prod'(X471,X472),bool)),happ(s(fun(X474,fun(X473,fun('type.pair.prod'(X471,X472),bool))),X475),s(X474,X476))),s(X473,X477))),s('type.pair.prod'(X471,X472),X478))))) )),
  inference(superposition,[],[f377,f345])).
fof(f7371,plain,(
  ( ! [X466,X463,X469,X465,X467,X468,X470,X464] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X463,X464),bool),'const.pair.UNCURRY_2'(s(fun(X466,fun(X465,fun('type.pair.prod'(X463,X464),bool))),X467),s('type.pair.prod'(X466,X465),'const.pair.,_2'(s(X466,X468),s(X465,X469))))),s(X463,sK13(happ(s(fun(X465,fun('type.pair.prod'(X463,X464),bool)),happ(s(fun(X466,fun(X465,fun('type.pair.prod'(X463,X464),bool))),X467),s(X466,X468))),s(X465,X469)),X464,X463)),s(X464,sK14(happ(s(fun(X465,fun('type.pair.prod'(X463,X464),bool)),happ(s(fun(X466,fun(X465,fun('type.pair.prod'(X463,X464),bool))),X467),s(X466,X468))),s(X465,X469)),X464,X463))))) | p(s(bool,happ(s(fun('type.pair.prod'(X463,X464),bool),happ(s(fun(X465,fun('type.pair.prod'(X463,X464),bool)),happ(s(fun(X466,fun(X465,fun('type.pair.prod'(X463,X464),bool))),X467),s(X466,X468))),s(X465,X469))),s('type.pair.prod'(X463,X464),X470))))) )),
  inference(superposition,[],[f378,f345])).
fof(f7370,plain,(
  ( ! [X455,X461,X457,X459,X460,X462,X456,X458] : (~p(s(bool,happ(s(fun('type.pair.prod'(X455,X456),bool),'const.pair.UNCURRY_2'(s(fun(X458,fun(X457,fun('type.pair.prod'(X455,X456),bool))),X459),s('type.pair.prod'(X458,X457),'const.pair.,_2'(s(X458,X460),s(X457,X461))))),s('type.pair.prod'(X455,X456),X462)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X455,X456),bool),happ(s(fun(X457,fun('type.pair.prod'(X455,X456),bool)),happ(s(fun(X458,fun(X457,fun('type.pair.prod'(X455,X456),bool))),X459),s(X458,X460))),s(X457,X461))),s(X455,sK20(happ(s(fun(X457,fun('type.pair.prod'(X455,X456),bool)),happ(s(fun(X458,fun(X457,fun('type.pair.prod'(X455,X456),bool))),X459),s(X458,X460))),s(X457,X461)),X456,X455)),s(X456,sK21(happ(s(fun(X457,fun('type.pair.prod'(X455,X456),bool)),happ(s(fun(X458,fun(X457,fun('type.pair.prod'(X455,X456),bool))),X459),s(X458,X460))),s(X457,X461)),X456,X455)))))) )),
  inference(superposition,[],[f379,f345])).
fof(f7369,plain,(
  ( ! [X453,X449,X451,X452,X454,X448,X450] : (~p(s(bool,happ(s(fun(X448,bool),'const.pair.UNCURRY_2'(s(fun(X450,fun(X449,fun(X448,bool))),X451),s('type.pair.prod'(X450,X449),'const.pair.,_2'(s(X450,X452),s(X449,X453))))),s(X448,X454)))) | p(s(bool,happ(s(fun(fun(X448,bool),bool),'?_1'),s(fun(X448,bool),happ(s(fun(X449,fun(X448,bool)),happ(s(fun(X450,fun(X449,fun(X448,bool))),X451),s(X450,X452))),s(X449,X453))))))) )),
  inference(superposition,[],[f303,f345])).
fof(f7368,plain,(
  ( ! [X445,X447,X443,X444,X446,X442] : (~p(s(bool,happ(s(fun(fun(X442,bool),bool),'?_1'),s(fun(X442,bool),'const.pair.UNCURRY_2'(s(fun(X444,fun(X443,fun(X442,bool))),X445),s('type.pair.prod'(X444,X443),'const.pair.,_2'(s(X444,X446),s(X443,X447)))))))) | p(s(bool,happ(s(fun(X442,bool),happ(s(fun(X443,fun(X442,bool)),happ(s(fun(X444,fun(X443,fun(X442,bool))),X445),s(X444,X446))),s(X443,X447))),s(X442,sK9(happ(s(fun(X443,fun(X442,bool)),happ(s(fun(X444,fun(X443,fun(X442,bool))),X445),s(X444,X446))),s(X443,X447)),X442)))))) )),
  inference(superposition,[],[f302,f345])).
fof(f7367,plain,(
  ( ! [X441,X436,X438,X440,X437,X439] : (~p(s(bool,happ(s(fun(X436,bool),'const.pair.UNCURRY_2'(s(fun(X438,fun(X437,fun(X436,bool))),X439),s('type.pair.prod'(X438,X437),'const.pair.,_2'(s(X438,X440),s(X437,X441))))),s(X436,sK8(happ(s(fun(X437,fun(X436,bool)),happ(s(fun(X438,fun(X437,fun(X436,bool))),X439),s(X438,X440))),s(X437,X441)),X436))))) | p(s(bool,happ(s(fun(fun(X436,bool),bool),'!_1'),s(fun(X436,bool),happ(s(fun(X437,fun(X436,bool)),happ(s(fun(X438,fun(X437,fun(X436,bool))),X439),s(X438,X440))),s(X437,X441))))))) )),
  inference(superposition,[],[f301,f345])).
fof(f7366,plain,(
  ( ! [X430,X432,X434,X429,X431,X433,X435] : (~p(s(bool,happ(s(fun(fun(X429,bool),bool),'!_1'),s(fun(X429,bool),'const.pair.UNCURRY_2'(s(fun(X431,fun(X430,fun(X429,bool))),X432),s('type.pair.prod'(X431,X430),'const.pair.,_2'(s(X431,X433),s(X430,X434)))))))) | p(s(bool,happ(s(fun(X429,bool),happ(s(fun(X430,fun(X429,bool)),happ(s(fun(X431,fun(X430,fun(X429,bool))),X432),s(X431,X433))),s(X430,X434))),s(X429,X435))))) )),
  inference(superposition,[],[f300,f345])).
fof(f7365,plain,(
  ( ! [X422,X428,X424,X426,X423,X425,X427] : (p(s(bool,'|dollar|exists_1'(s(fun(X422,bool),'const.pair.UNCURRY_2'(s(fun(X424,fun(X423,fun(X422,bool))),X425),s('type.pair.prod'(X424,X423),'const.pair.,_2'(s(X424,X426),s(X423,X427)))))))) | s(X422,X428) = s(X422,sK7(X428,happ(s(fun(X423,fun(X422,bool)),happ(s(fun(X424,fun(X423,fun(X422,bool))),X425),s(X424,X426))),s(X423,X427)),X422))) )),
  inference(superposition,[],[f292,f345])).
fof(f7364,plain,(
  ( ! [X420,X416,X418,X415,X421,X417,X419] : (~p(s(bool,happ(s(fun(X415,bool),'const.pair.UNCURRY_2'(s(fun(X417,fun(X416,fun(X415,bool))),X418),s('type.pair.prod'(X417,X416),'const.pair.,_2'(s(X417,X419),s(X416,X420))))),s(X415,X421)))) | p(s(bool,'|dollar|exists_1'(s(fun(X415,bool),happ(s(fun(X416,fun(X415,bool)),happ(s(fun(X417,fun(X416,fun(X415,bool))),X418),s(X417,X419))),s(X416,X420))))))) )),
  inference(superposition,[],[f289,f345])).
fof(f7363,plain,(
  ( ! [X412,X414,X410,X413,X409,X411] : (p(s(bool,happ(s(fun(X409,bool),'const.pair.UNCURRY_2'(s(fun(X411,fun(X410,fun(X409,bool))),X412),s('type.pair.prod'(X411,X410),'const.pair.,_2'(s(X411,X413),s(X410,X414))))),s(X409,sK6(happ(s(fun(X410,fun(X409,bool)),happ(s(fun(X411,fun(X410,fun(X409,bool))),X412),s(X411,X413))),s(X410,X414)),X409))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X409,bool),happ(s(fun(X410,fun(X409,bool)),happ(s(fun(X411,fun(X410,fun(X409,bool))),X412),s(X411,X413))),s(X410,X414))))))) )),
  inference(superposition,[],[f288,f345])).
fof(f7362,plain,(
  ( ! [X408,X405,X407,X403,X404,X406] : (~p(s(bool,'|dollar|exists_1'(s(fun(X403,bool),'const.pair.UNCURRY_2'(s(fun(X405,fun(X404,fun(X403,bool))),X406),s('type.pair.prod'(X405,X404),'const.pair.,_2'(s(X405,X407),s(X404,X408)))))))) | p(s(bool,happ(s(fun(X403,bool),happ(s(fun(X404,fun(X403,bool)),happ(s(fun(X405,fun(X404,fun(X403,bool))),X406),s(X405,X407))),s(X404,X408))),s(X403,sK5(happ(s(fun(X404,fun(X403,bool)),happ(s(fun(X405,fun(X404,fun(X403,bool))),X406),s(X405,X407))),s(X404,X408)),X403))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X403,bool),happ(s(fun(X404,fun(X403,bool)),happ(s(fun(X405,fun(X404,fun(X403,bool))),X406),s(X405,X407))),s(X404,X408))))))) )),
  inference(superposition,[],[f286,f345])).
fof(f7361,plain,(
  ( ! [X397,X399,X401,X398,X400,X402] : (~p(s(bool,'|dollar|exists_1'(s(fun(X397,bool),'const.pair.UNCURRY_2'(s(fun(X399,fun(X398,fun(X397,bool))),X400),s('type.pair.prod'(X399,X398),'const.pair.,_2'(s(X399,X401),s(X398,X402)))))))) | p(s(bool,happ(s(fun(X397,bool),happ(s(fun(X398,fun(X397,bool)),happ(s(fun(X399,fun(X398,fun(X397,bool))),X400),s(X399,X401))),s(X398,X402))),s(X397,sK4(happ(s(fun(X398,fun(X397,bool)),happ(s(fun(X399,fun(X398,fun(X397,bool))),X400),s(X399,X401))),s(X398,X402)),X397))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X397,bool),happ(s(fun(X398,fun(X397,bool)),happ(s(fun(X399,fun(X398,fun(X397,bool))),X400),s(X399,X401))),s(X398,X402))))))) )),
  inference(superposition,[],[f285,f345])).
fof(f7360,plain,(
  ( ! [X389,X391,X393,X395,X390,X396,X392,X394] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X389,bool),'const.pair.UNCURRY_2'(s(fun(X391,fun(X390,fun(X389,bool))),X392),s('type.pair.prod'(X391,X390),'const.pair.,_2'(s(X391,X393),s(X390,X394)))))))) | ~p(s(bool,happ(s(fun(X389,bool),happ(s(fun(X390,fun(X389,bool)),happ(s(fun(X391,fun(X390,fun(X389,bool))),X392),s(X391,X393))),s(X390,X394))),s(X389,X395)))) | ~p(s(bool,happ(s(fun(X389,bool),happ(s(fun(X390,fun(X389,bool)),happ(s(fun(X391,fun(X390,fun(X389,bool))),X392),s(X391,X393))),s(X390,X394))),s(X389,X396)))) | s(X389,X395) = s(X389,X396)) )),
  inference(superposition,[],[f284,f345])).
fof(f7359,plain,(
  ( ! [X383,X385,X387,X388,X384,X386] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X383,bool),'const.pair.UNCURRY_2'(s(fun(X385,fun(X384,fun(X383,bool))),X386),s('type.pair.prod'(X385,X384),'const.pair.,_2'(s(X385,X387),s(X384,X388)))))))) | p(s(bool,'|dollar|exists_1'(s(fun(X383,bool),happ(s(fun(X384,fun(X383,bool)),happ(s(fun(X385,fun(X384,fun(X383,bool))),X386),s(X385,X387))),s(X384,X388))))))) )),
  inference(superposition,[],[f283,f345])).
fof(f7357,plain,(
  ( ! [X372,X374,X370,X376,X373,X375,X371] : (~p(s(bool,happ(s(fun(X370,bool),'const.pair.UNCURRY_2'(s(fun(X372,fun(X371,fun(X370,bool))),X373),s('type.pair.prod'(X372,X371),'const.pair.,_2'(s(X372,X374),s(X371,X375))))),s(X370,X376)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X371,X370),bool),'const.pair.UNCURRY_1'(s(fun(X371,fun(X370,bool)),happ(s(fun(X372,fun(X371,fun(X370,bool))),X373),s(X372,X374))))))))) )),
  inference(superposition,[],[f319,f345])).
fof(f7356,plain,(
  ( ! [X364,X366,X368,X365,X367,X363,X369] : (~p(s(bool,happ(s(fun(X363,bool),'const.pair.UNCURRY_2'(s(fun(X365,fun(X364,fun(X363,bool))),X366),s('type.pair.prod'(X365,X364),'const.pair.,_2'(s(X365,X367),s(X364,X368))))),s(X363,X369)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X364,fun(X363,bool)),happ(s(fun(X365,fun(X364,fun(X363,bool))),X366),s(X365,X367))),s('type.pair.prod'(X364,X363),sK25(happ(s(fun(X365,fun(X364,fun(X363,bool))),X366),s(X365,X367)),X363,X364)))))) )),
  inference(superposition,[],[f374,f345])).
fof(f7348,plain,(
  ( ! [X158,X156,X161,X159,X157,X155,X160] : (p(s(bool,happ(s(fun(X155,bool),happ(s(fun(X155,fun(X155,bool)),'=_0'),s(X155,X161))),s(X155,'const.pair.UNCURRY_2'(s(fun(X157,fun(X156,X155)),X158),s('type.pair.prod'(X157,X156),'const.pair.,_2'(s(X157,X159),s(X156,X160)))))))) | s(X155,X161) != s(X155,happ(s(fun(X156,X155),happ(s(fun(X157,fun(X156,X155)),X158),s(X157,X159))),s(X156,X160)))) )),
  inference(superposition,[],[f331,f345])).
fof(f7347,plain,(
  ( ! [X154,X152,X151,X149,X153,X150,X148] : (p(s(bool,happ(s(fun(X148,bool),happ(s(fun(X148,fun(X148,bool)),'=_0'),s(X148,'const.pair.UNCURRY_2'(s(fun(X150,fun(X149,X148)),X151),s('type.pair.prod'(X150,X149),'const.pair.,_2'(s(X150,X152),s(X149,X153))))))),s(X148,X154)))) | s(X148,X154) != s(X148,happ(s(fun(X149,X148),happ(s(fun(X150,fun(X149,X148)),X151),s(X150,X152))),s(X149,X153)))) )),
  inference(superposition,[],[f331,f345])).
fof(f7346,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141] : (~p(s(bool,happ(s(fun(X141,bool),happ(s(fun(X141,fun(X141,bool)),'=_0'),s(X141,X147))),s(X141,'const.pair.UNCURRY_2'(s(fun(X143,fun(X142,X141)),X144),s('type.pair.prod'(X143,X142),'const.pair.,_2'(s(X143,X145),s(X142,X146)))))))) | s(X141,X147) = s(X141,happ(s(fun(X142,X141),happ(s(fun(X143,fun(X142,X141)),X144),s(X143,X145))),s(X142,X146)))) )),
  inference(superposition,[],[f330,f345])).
fof(f7345,plain,(
  ( ! [X134,X140,X138,X136,X135,X139,X137] : (~p(s(bool,happ(s(fun(X134,bool),happ(s(fun(X134,fun(X134,bool)),'=_0'),s(X134,'const.pair.UNCURRY_2'(s(fun(X136,fun(X135,X134)),X137),s('type.pair.prod'(X136,X135),'const.pair.,_2'(s(X136,X138),s(X135,X139))))))),s(X134,X140)))) | s(X134,X140) = s(X134,happ(s(fun(X135,X134),happ(s(fun(X136,fun(X135,X134)),X137),s(X136,X138))),s(X135,X139)))) )),
  inference(superposition,[],[f330,f345])).
fof(f7344,plain,(
  ( ! [X127,X132,X130,X128,X126,X133,X131,X129] : (s(X126,sK34(X126,X132)) != s(X126,'const.pair.UNCURRY_2'(s(fun(X128,fun(X127,X126)),X129),s('type.pair.prod'(X128,X127),'const.pair.,_2'(s(X128,X130),s(X127,X131))))) | p(s(bool,happ(s(fun(X126,bool),happ(s(fun(X132,fun(X126,bool)),sK32(X126,X132)),s(X132,X133))),s(X126,happ(s(fun(X127,X126),happ(s(fun(X128,fun(X127,X126)),X129),s(X128,X130))),s(X127,X131)))))) | s(X132,X133) != s(X132,sK33(X126,X132))) )),
  inference(superposition,[],[f329,f345])).
fof(f7343,plain,(
  ( ! [X125,X123,X121,X118,X124,X122,X120,X119] : (~p(s(bool,happ(s(fun(X118,bool),happ(s(fun(X124,fun(X118,bool)),sK32(X118,X124)),s(X124,X125))),s(X118,'const.pair.UNCURRY_2'(s(fun(X120,fun(X119,X118)),X121),s('type.pair.prod'(X120,X119),'const.pair.,_2'(s(X120,X122),s(X119,X123)))))))) | s(X118,sK34(X118,X124)) = s(X118,happ(s(fun(X119,X118),happ(s(fun(X120,fun(X119,X118)),X121),s(X120,X122))),s(X119,X123)))) )),
  inference(superposition,[],[f328,f345])).
fof(f7340,plain,(
  ( ! [X94,X101,X99,X97,X95,X100,X98,X96] : (~p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X94,fun(X100,bool)),sK32(X100,X94)),s(X94,'const.pair.UNCURRY_2'(s(fun(X96,fun(X95,X94)),X97),s('type.pair.prod'(X96,X95),'const.pair.,_2'(s(X96,X98),s(X95,X99))))))),s(X100,X101)))) | s(X94,sK33(X100,X94)) = s(X94,happ(s(fun(X95,X94),happ(s(fun(X96,fun(X95,X94)),X97),s(X96,X98))),s(X95,X99)))) )),
  inference(superposition,[],[f327,f345])).
fof(f7329,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(X14,happ(s(fun(X8,X14),X15),s(X8,'const.pair.UNCURRY_2'(s(fun(X10,fun(X9,X8)),X11),s('type.pair.prod'(X10,X9),'const.pair.,_2'(s(X10,X12),s(X9,X13))))))) = s(X14,happ(s(fun(X9,X14),'_dst_x0x1_2'(s(fun(X8,X14),X15),s(fun(X9,X8),happ(s(fun(X10,fun(X9,X8)),X11),s(X10,X12))))),s(X9,X13)))) )),
  inference(superposition,[],[f344,f345])).
fof(f7327,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,'F_0') = s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X8),s(X6,X9))),s(X7,X10))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,bool)),X8),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X10)))))) )),
  inference(superposition,[],[f345,f264])).
fof(f7242,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (s(X14,'const.pair.UNCURRY_2'(s(fun(X16,fun(X15,X14)),X17),s('type.pair.prod'(X16,X15),'const.pair.,_2'(s(X16,X18),s(X15,happ(s(fun(X19,X15),X20),s(X19,X21))))))) = s(X14,happ(s(fun(X19,X14),'_dst_x0x1_2'(s(fun(X15,X14),happ(s(fun(X16,fun(X15,X14)),X17),s(X16,X18))),s(fun(X19,X15),X20))),s(X19,X21)))) )),
  inference(superposition,[],[f345,f344])).
fof(f7241,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),happ(s(fun(X3,fun('type.pair.prod'(X1,X2),X0)),X4),s(X3,X5))),s(X1,X6),s(X2,X7))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X3,fun('type.pair.prod'(X1,X2),X0)),X4),s('type.pair.prod'(X3,'type.pair.prod'(X1,X2)),'const.pair.,_2'(s(X3,X5),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X6),s(X2,X7)))))))) )),
  inference(superposition,[],[f345,f343])).
fof(f7240,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X42,X40] : (s(X36,happ(s(fun(X35,X36),'const.pair.UNCURRY_2'(s(fun(X38,fun(X37,fun(X35,X36))),X39),s('type.pair.prod'(X38,X37),'const.pair.,_2'(s(X38,X40),s(X37,X41))))),s(X35,X42))) = s(X36,'const.pair.UNCURRY_2'(s(fun(X37,fun(X35,X36)),happ(s(fun(X38,fun(X37,fun(X35,X36))),X39),s(X38,X40))),s('type.pair.prod'(X37,X35),'const.pair.,_2'(s(X37,X41),s(X35,X42)))))) )),
  inference(superposition,[],[f345,f345])).
fof(f7239,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X34,X32] : (s(X28,happ(s(fun(X27,X28),happ(s(fun(X30,fun(X27,X28)),X31),s(X30,happ(s(fun(X29,X30),X32),s(X29,X33))))),s(X27,X34))) = s(X28,'const.pair.UNCURRY_2'(s(fun(X29,fun(X27,X28)),'_dst_x0x1_2'(s(fun(X30,fun(X27,X28)),X31),s(fun(X29,X30),X32))),s('type.pair.prod'(X29,X27),'const.pair.,_2'(s(X29,X33),s(X27,X34)))))) )),
  inference(superposition,[],[f345,f344])).
fof(f7238,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (s(X20,'const.pair.UNCURRY_2'(s(fun(X21,fun(X19,X20)),X22),s('type.pair.prod'(X21,X19),'const.pair.,_2'(s(X21,happ(s(fun(X23,X21),X24),s(X23,X25))),s(X19,X26))))) = s(X20,happ(s(fun(X19,X20),happ(s(fun(X23,fun(X19,X20)),'_dst_x0x1_2'(s(fun(X21,fun(X19,X20)),X22),s(fun(X23,X21),X24))),s(X23,X25))),s(X19,X26)))) )),
  inference(superposition,[],[f345,f344])).
fof(f7237,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X18,X16] : (s(X12,happ(s(fun(X11,X12),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X13,X14),fun(X11,X12)),X15),s(X13,X16),s(X14,X17))),s(X11,X18))) = s(X12,'const.pair.UNCURRY_2'(s(fun('type.pair.prod'(X13,X14),fun(X11,X12)),X15),s('type.pair.prod'('type.pair.prod'(X13,X14),X11),'const.pair.,_2'(s('type.pair.prod'(X13,X14),'const.pair.,_2'(s(X13,X16),s(X14,X17))),s(X11,X18)))))) )),
  inference(superposition,[],[f345,f343])).
fof(f7236,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s(X5,happ(s(fun(X4,X5),'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,fun(X4,X5))),X8),s('type.pair.prod'(X6,X7),X9))),s(X4,X10))) = s(X5,'const.pair.UNCURRY_2'(s(fun('type.pair.prod'(X6,X7),fun(X4,X5)),'const.pair.UNCURRY_1'(s(fun(X6,fun(X7,fun(X4,X5))),X8))),s('type.pair.prod'('type.pair.prod'(X6,X7),X4),'const.pair.,_2'(s('type.pair.prod'(X6,X7),X9),s(X4,X10)))))) )),
  inference(superposition,[],[f345,f335])).
fof(f7234,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))),s(X1,X5))))),s(X0,'const.pair.UNCURRY_2'(s(fun(X2,fun(X1,X0)),X3),s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X4),s(X1,X5))))))))) )),
  inference(unit_resulting_resolution,[],[f345,f331])).
fof(f7235,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))))),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))))))) )),
  inference(unit_resulting_resolution,[],[f345,f331])).
fof(f6930,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),sK32(X0,X1)),s(X1,X2))),s(X0,sK34(X0,X1))))) | s(X1,X2) != s(X1,sK33(X0,X1))) )),
  inference(equality_resolution,[],[f329])).
fof(f6929,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X0,fun(bool,bool)),sK32(bool,X0)),s(X0,X2))),s(bool,X1)))) | s(X0,X2) != s(X0,sK33(bool,X0)) | s(bool,'T_0') = s(bool,sK34(bool,X0))) )),
  inference(superposition,[],[f329,f264])).
fof(f6909,plain,(
  ( ! [X59,X57,X58] : (s(bool,'F_0') != s(bool,sK34(bool,X58)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X58,fun(bool,bool)),sK32(bool,X58)),s(X58,X59))),s(bool,X57)))) | s(X58,X59) != s(X58,sK33(bool,X58)) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f329,f264])).
fof(f6904,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X34,X32] : (s(X27,sK34(X27,X33)) != s(X27,happ(s(fun(X29,X27),X30),s(X29,happ(s(fun(X28,X29),X31),s(X28,X32))))) | p(s(bool,happ(s(fun(X27,bool),happ(s(fun(X33,fun(X27,bool)),sK32(X27,X33)),s(X33,X34))),s(X27,happ(s(fun(X28,X27),'_dst_x0x1_2'(s(fun(X29,X27),X30),s(fun(X28,X29),X31))),s(X28,X32)))))) | s(X33,X34) != s(X33,sK33(X27,X33))) )),
  inference(superposition,[],[f329,f344])).
fof(f6903,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (s(X19,sK34(X19,X25)) != s(X19,happ(s(fun(X22,X19),'_dst_x0x1_2'(s(fun(X20,X19),X21),s(fun(X22,X20),X23))),s(X22,X24))) | p(s(bool,happ(s(fun(X19,bool),happ(s(fun(X25,fun(X19,bool)),sK32(X19,X25)),s(X25,X26))),s(X19,happ(s(fun(X20,X19),X21),s(X20,happ(s(fun(X22,X20),X23),s(X22,X24)))))))) | s(X25,X26) != s(X25,sK33(X19,X25))) )),
  inference(superposition,[],[f329,f344])).
fof(f7092,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun(X2,fun(X0,bool)),sK32(X0,X2)),s(X2,X3),s(fun(X0,X0),'_i_0_0'))),s(X0,X1)))) | s(X0,X1) != s(X0,sK34(X0,X2)) | s(X2,X3) != s(X2,sK33(X0,X2))) )),
  inference(forward_demodulation,[],[f6900,f352])).
fof(f6900,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) != s(X0,sK34(X0,X2)) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X2,fun(X0,bool)),sK32(X0,X2)),s(X2,X3))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))))) | s(X2,X3) != s(X2,sK33(X0,X2))) )),
  inference(superposition,[],[f329,f279])).
fof(f7066,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun(X1,fun(X0,bool)),sK32(X0,X1)),s(X1,sK33(X0,X1)),s(fun(X0,X0),'_i_0_0'))),s(X0,sK34(X0,X1)))))) )),
  inference(forward_demodulation,[],[f7065,f322])).
fof(f7065,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun(X1,fun(X0,bool)),sK32(X0,X1)),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X3),s(X1,sK33(X0,X1)))))),s(fun(X0,X0),'_i_0_0'))),s(X0,sK34(X0,X1)))))) )),
  inference(forward_demodulation,[],[f6847,f352])).
fof(f6847,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),sK32(X0,X1)),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X3),s(X1,sK33(X0,X1)))))))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK34(X0,X1)))))))) )),
  inference(unit_resulting_resolution,[],[f322,f279,f329])).
fof(f7064,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun(fun(X1,fun(X2,X3)),fun(X0,bool)),sK32(X0,fun(X1,fun(X2,X3)))),s(fun(X1,fun(X2,X3)),sK33(X0,fun(X1,fun(X2,X3)))),s(fun(X0,X0),'_i_0_0'))),s(X0,sK34(X0,fun(X1,fun(X2,X3)))))))) )),
  inference(forward_demodulation,[],[f7063,f332])).
fof(f7063,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun(fun(X1,fun(X2,X3)),fun(X0,bool)),sK32(X0,fun(X1,fun(X2,X3)))),s(fun(X1,fun(X2,X3)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X1,X2),X3),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X3)),sK33(X0,fun(X1,fun(X2,X3)))))))),s(fun(X0,X0),'_i_0_0'))),s(X0,sK34(X0,fun(X1,fun(X2,X3)))))))) )),
  inference(forward_demodulation,[],[f6848,f352])).
fof(f6848,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,fun(X2,X3)),fun(X0,bool)),sK32(X0,fun(X1,fun(X2,X3)))),s(fun(X1,fun(X2,X3)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X1,X2),X3),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X3)),sK33(X0,fun(X1,fun(X2,X3)))))))))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK34(X0,fun(X1,fun(X2,X3)))))))))) )),
  inference(unit_resulting_resolution,[],[f332,f279,f329])).
fof(f7062,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun('type.pair.prod'(X1,X2),fun(X0,bool)),sK32(X0,'type.pair.prod'(X1,X2))),s('type.pair.prod'(X1,X2),sK33(X0,'type.pair.prod'(X1,X2))),s(fun(X0,X0),'_i_0_0'))),s(X0,sK34(X0,'type.pair.prod'(X1,X2))))))) )),
  inference(forward_demodulation,[],[f7061,f321])).
fof(f7061,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun('type.pair.prod'(X1,X2),fun(X0,bool)),sK32(X0,'type.pair.prod'(X1,X2))),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,sK28(sK33(X0,'type.pair.prod'(X1,X2)),X2,X1)),s(X2,sK29(sK33(X0,'type.pair.prod'(X1,X2)),X2,X1)))),s(fun(X0,X0),'_i_0_0'))),s(X0,sK34(X0,'type.pair.prod'(X1,X2))))))) )),
  inference(forward_demodulation,[],[f6849,f352])).
fof(f6849,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun('type.pair.prod'(X1,X2),fun(X0,bool)),sK32(X0,'type.pair.prod'(X1,X2))),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,sK28(sK33(X0,'type.pair.prod'(X1,X2)),X2,X1)),s(X2,sK29(sK33(X0,'type.pair.prod'(X1,X2)),X2,X1)))))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK34(X0,'type.pair.prod'(X1,X2))))))))) )),
  inference(unit_resulting_resolution,[],[f321,f279,f329])).
fof(f7031,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),sK32(X0,X1)),s(X1,sK33(X0,X1)))),s(X0,sK34(X0,X1)))))) )),
  inference(forward_demodulation,[],[f7030,f322])).
fof(f7030,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),sK32(X0,X1)),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X3),s(X1,sK33(X0,X1)))))))),s(X0,sK34(X0,X1)))))) )),
  inference(forward_demodulation,[],[f6863,f322])).
fof(f6863,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),sK32(X0,X1)),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X3),s(X1,sK33(X0,X1)))))))),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X4,X0),'const.pair.,_2'(s(X4,X5),s(X0,sK34(X0,X1)))))))))) )),
  inference(unit_resulting_resolution,[],[f322,f322,f329])).
fof(f7029,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,fun(X2,X3)),fun(X0,bool)),sK32(X0,fun(X1,fun(X2,X3)))),s(fun(X1,fun(X2,X3)),sK33(X0,fun(X1,fun(X2,X3)))))),s(X0,sK34(X0,fun(X1,fun(X2,X3)))))))) )),
  inference(forward_demodulation,[],[f7028,f332])).
fof(f7028,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,fun(X2,X3)),fun(X0,bool)),sK32(X0,fun(X1,fun(X2,X3)))),s(fun(X1,fun(X2,X3)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X1,X2),X3),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X3)),sK33(X0,fun(X1,fun(X2,X3)))))))))),s(X0,sK34(X0,fun(X1,fun(X2,X3)))))))) )),
  inference(forward_demodulation,[],[f6864,f322])).
fof(f6864,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,fun(X2,X3)),fun(X0,bool)),sK32(X0,fun(X1,fun(X2,X3)))),s(fun(X1,fun(X2,X3)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X1,X2),X3),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X3)),sK33(X0,fun(X1,fun(X2,X3)))))))))),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X4,X0),'const.pair.,_2'(s(X4,X5),s(X0,sK34(X0,fun(X1,fun(X2,X3)))))))))))) )),
  inference(unit_resulting_resolution,[],[f332,f322,f329])).
fof(f7027,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun('type.pair.prod'(X1,X2),fun(X0,bool)),sK32(X0,'type.pair.prod'(X1,X2))),s('type.pair.prod'(X1,X2),sK33(X0,'type.pair.prod'(X1,X2))))),s(X0,sK34(X0,'type.pair.prod'(X1,X2))))))) )),
  inference(forward_demodulation,[],[f7026,f5470])).
fof(f7026,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),fun(X0,bool)),sK32(X0,'type.pair.prod'(X1,X2))),s(X1,sK28(sK33(X0,'type.pair.prod'(X1,X2)),X2,X1)),s(X2,sK29(sK33(X0,'type.pair.prod'(X1,X2)),X2,X1)))),s(X0,sK34(X0,'type.pair.prod'(X1,X2))))))) )),
  inference(forward_demodulation,[],[f7025,f343])).
fof(f7025,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun('type.pair.prod'(X1,X2),fun(X0,bool)),sK32(X0,'type.pair.prod'(X1,X2))),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,sK28(sK33(X0,'type.pair.prod'(X1,X2)),X2,X1)),s(X2,sK29(sK33(X0,'type.pair.prod'(X1,X2)),X2,X1)))))),s(X0,sK34(X0,'type.pair.prod'(X1,X2))))))) )),
  inference(forward_demodulation,[],[f6865,f322])).
fof(f6865,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun('type.pair.prod'(X1,X2),fun(X0,bool)),sK32(X0,'type.pair.prod'(X1,X2))),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,sK28(sK33(X0,'type.pair.prod'(X1,X2)),X2,X1)),s(X2,sK29(sK33(X0,'type.pair.prod'(X1,X2)),X2,X1)))))),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X3,X0),'const.pair.,_2'(s(X3,X4),s(X0,sK34(X0,'type.pair.prod'(X1,X2))))))))))) )),
  inference(unit_resulting_resolution,[],[f321,f322,f329])).
fof(f7012,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(X3,fun(fun(X0,fun(X1,X2)),bool)),sK32(fun(X0,fun(X1,X2)),X3)),s(X3,sK33(fun(X0,fun(X1,X2)),X3)))),s(fun(X0,fun(X1,X2)),sK34(fun(X0,fun(X1,X2)),X3)))))) )),
  inference(forward_demodulation,[],[f7011,f322])).
fof(f7011,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(X3,fun(fun(X0,fun(X1,X2)),bool)),sK32(fun(X0,fun(X1,X2)),X3)),s(X3,'const.pair.SND_1'(s('type.pair.prod'(X4,X3),'const.pair.,_2'(s(X4,X5),s(X3,sK33(fun(X0,fun(X1,X2)),X3)))))))),s(fun(X0,fun(X1,X2)),sK34(fun(X0,fun(X1,X2)),X3)))))) )),
  inference(forward_demodulation,[],[f6871,f332])).
fof(f6871,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(X3,fun(fun(X0,fun(X1,X2)),bool)),sK32(fun(X0,fun(X1,X2)),X3)),s(X3,'const.pair.SND_1'(s('type.pair.prod'(X4,X3),'const.pair.,_2'(s(X4,X5),s(X3,sK33(fun(X0,fun(X1,X2)),X3)))))))),s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),sK34(fun(X0,fun(X1,X2)),X3)))))))))) )),
  inference(unit_resulting_resolution,[],[f322,f332,f329])).
fof(f7010,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(fun(X3,fun(X4,X5)),fun(fun(X0,fun(X1,X2)),bool)),sK32(fun(X0,fun(X1,X2)),fun(X3,fun(X4,X5)))),s(fun(X3,fun(X4,X5)),sK33(fun(X0,fun(X1,X2)),fun(X3,fun(X4,X5)))))),s(fun(X0,fun(X1,X2)),sK34(fun(X0,fun(X1,X2)),fun(X3,fun(X4,X5)))))))) )),
  inference(forward_demodulation,[],[f7009,f332])).
fof(f7009,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(fun(X3,fun(X4,X5)),fun(fun(X0,fun(X1,X2)),bool)),sK32(fun(X0,fun(X1,X2)),fun(X3,fun(X4,X5)))),s(fun(X3,fun(X4,X5)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X3,X4),X5),'const.pair.UNCURRY_1'(s(fun(X3,fun(X4,X5)),sK33(fun(X0,fun(X1,X2)),fun(X3,fun(X4,X5)))))))))),s(fun(X0,fun(X1,X2)),sK34(fun(X0,fun(X1,X2)),fun(X3,fun(X4,X5)))))))) )),
  inference(forward_demodulation,[],[f6872,f332])).
fof(f6872,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(fun(X3,fun(X4,X5)),fun(fun(X0,fun(X1,X2)),bool)),sK32(fun(X0,fun(X1,X2)),fun(X3,fun(X4,X5)))),s(fun(X3,fun(X4,X5)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X3,X4),X5),'const.pair.UNCURRY_1'(s(fun(X3,fun(X4,X5)),sK33(fun(X0,fun(X1,X2)),fun(X3,fun(X4,X5)))))))))),s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),sK34(fun(X0,fun(X1,X2)),fun(X3,fun(X4,X5)))))))))))) )),
  inference(unit_resulting_resolution,[],[f332,f332,f329])).
fof(f7008,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun('type.pair.prod'(X3,X4),fun(fun(X0,fun(X1,X2)),bool)),sK32(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4))),s('type.pair.prod'(X3,X4),sK33(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4))))),s(fun(X0,fun(X1,X2)),sK34(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4))))))) )),
  inference(forward_demodulation,[],[f7007,f5470])).
fof(f7007,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X3,X4),fun(fun(X0,fun(X1,X2)),bool)),sK32(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4))),s(X3,sK28(sK33(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4)),X4,X3)),s(X4,sK29(sK33(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4)),X4,X3)))),s(fun(X0,fun(X1,X2)),sK34(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4))))))) )),
  inference(forward_demodulation,[],[f7006,f343])).
fof(f7006,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun('type.pair.prod'(X3,X4),fun(fun(X0,fun(X1,X2)),bool)),sK32(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4))),s('type.pair.prod'(X3,X4),'const.pair.,_2'(s(X3,sK28(sK33(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4)),X4,X3)),s(X4,sK29(sK33(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4)),X4,X3)))))),s(fun(X0,fun(X1,X2)),sK34(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4))))))) )),
  inference(forward_demodulation,[],[f6873,f332])).
fof(f6873,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun('type.pair.prod'(X3,X4),fun(fun(X0,fun(X1,X2)),bool)),sK32(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4))),s('type.pair.prod'(X3,X4),'const.pair.,_2'(s(X3,sK28(sK33(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4)),X4,X3)),s(X4,sK29(sK33(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4)),X4,X3)))))),s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),sK34(fun(X0,fun(X1,X2)),'type.pair.prod'(X3,X4))))))))))) )),
  inference(unit_resulting_resolution,[],[f321,f332,f329])).
fof(f6990,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(X2,fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),X2)),s(X2,sK33('type.pair.prod'(X0,X1),X2)))),s('type.pair.prod'(X0,X1),sK34('type.pair.prod'(X0,X1),X2)))))) )),
  inference(forward_demodulation,[],[f6989,f322])).
fof(f6989,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(X2,fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),X2)),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X3,X2),'const.pair.,_2'(s(X3,X4),s(X2,sK33('type.pair.prod'(X0,X1),X2)))))))),s('type.pair.prod'(X0,X1),sK34('type.pair.prod'(X0,X1),X2)))))) )),
  inference(forward_demodulation,[],[f6988,f5470])).
fof(f6988,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(X2,fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),X2)),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X3,X2),'const.pair.,_2'(s(X3,X4),s(X2,sK33('type.pair.prod'(X0,X1),X2)))))))),s(X0,sK28(sK34('type.pair.prod'(X0,X1),X2),X1,X0)),s(X1,sK29(sK34('type.pair.prod'(X0,X1),X2),X1,X0)))))) )),
  inference(forward_demodulation,[],[f6879,f343])).
fof(f6879,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(X2,fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),X2)),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X3,X2),'const.pair.,_2'(s(X3,X4),s(X2,sK33('type.pair.prod'(X0,X1),X2)))))))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK28(sK34('type.pair.prod'(X0,X1),X2),X1,X0)),s(X1,sK29(sK34('type.pair.prod'(X0,X1),X2),X1,X0)))))))) )),
  inference(unit_resulting_resolution,[],[f322,f321,f329])).
fof(f6987,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun(X2,fun(X3,X4)),fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4)))),s(fun(X2,fun(X3,X4)),sK33('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4)))))),s('type.pair.prod'(X0,X1),sK34('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4)))))))) )),
  inference(forward_demodulation,[],[f6986,f332])).
fof(f6986,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun(X2,fun(X3,X4)),fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4)))),s(fun(X2,fun(X3,X4)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X2,X3),X4),'const.pair.UNCURRY_1'(s(fun(X2,fun(X3,X4)),sK33('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4)))))))))),s('type.pair.prod'(X0,X1),sK34('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4)))))))) )),
  inference(forward_demodulation,[],[f6985,f5470])).
fof(f6985,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun(X2,fun(X3,X4)),fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4)))),s(fun(X2,fun(X3,X4)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X2,X3),X4),'const.pair.UNCURRY_1'(s(fun(X2,fun(X3,X4)),sK33('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4)))))))))),s(X0,sK28(sK34('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4))),X1,X0)),s(X1,sK29(sK34('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4))),X1,X0)))))) )),
  inference(forward_demodulation,[],[f6880,f343])).
fof(f6880,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun(X2,fun(X3,X4)),fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4)))),s(fun(X2,fun(X3,X4)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X2,X3),X4),'const.pair.UNCURRY_1'(s(fun(X2,fun(X3,X4)),sK33('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4)))))))))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK28(sK34('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4))),X1,X0)),s(X1,sK29(sK34('type.pair.prod'(X0,X1),fun(X2,fun(X3,X4))),X1,X0)))))))) )),
  inference(unit_resulting_resolution,[],[f332,f321,f329])).
fof(f6984,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun('type.pair.prod'(X2,X3),fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))),s('type.pair.prod'(X2,X3),sK33('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))))),s('type.pair.prod'(X0,X1),sK34('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))))))) )),
  inference(forward_demodulation,[],[f6983,f5470])).
fof(f6983,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X2,X3),fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))),s(X2,sK28(sK33('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),X3,X2)),s(X3,sK29(sK33('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),X3,X2)))),s('type.pair.prod'(X0,X1),sK34('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))))))) )),
  inference(forward_demodulation,[],[f6982,f343])).
fof(f6982,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun('type.pair.prod'(X2,X3),fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))),s('type.pair.prod'(X2,X3),'const.pair.,_2'(s(X2,sK28(sK33('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),X3,X2)),s(X3,sK29(sK33('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),X3,X2)))))),s('type.pair.prod'(X0,X1),sK34('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))))))) )),
  inference(forward_demodulation,[],[f6981,f5470])).
fof(f6981,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun('type.pair.prod'(X2,X3),fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))),s('type.pair.prod'(X2,X3),'const.pair.,_2'(s(X2,sK28(sK33('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),X3,X2)),s(X3,sK29(sK33('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),X3,X2)))))),s(X0,sK28(sK34('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),X1,X0)),s(X1,sK29(sK34('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),X1,X0)))))) )),
  inference(forward_demodulation,[],[f6881,f343])).
fof(f6881,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun('type.pair.prod'(X2,X3),fun('type.pair.prod'(X0,X1),bool)),sK32('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))),s('type.pair.prod'(X2,X3),'const.pair.,_2'(s(X2,sK28(sK33('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),X3,X2)),s(X3,sK29(sK33('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),X3,X2)))))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK28(sK34('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),X1,X0)),s(X1,sK29(sK34('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),X1,X0)))))))) )),
  inference(unit_resulting_resolution,[],[f321,f321,f329])).
fof(f6737,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (~p(s(bool,happ(s(fun(X22,bool),'_dst_x0x1_2'(s(fun(X18,bool),happ(s(fun(X19,fun(X18,bool)),X20),s(X19,X21))),s(fun(X22,X18),X23))),s(X22,X24)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X19,fun(X18,bool)),X20),s('type.pair.prod'(X19,X18),sK25(X20,X18,X19)))))) )),
  inference(superposition,[],[f374,f344])).
fof(f6736,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X11,X12),bool),happ(s(fun(X13,fun('type.pair.prod'(X11,X12),bool)),X14),s(X13,X15))),s(X11,X16),s(X12,X17)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X13,fun('type.pair.prod'(X11,X12),bool)),X14),s('type.pair.prod'(X13,'type.pair.prod'(X11,X12)),sK25(X14,'type.pair.prod'(X11,X12),X13)))))) )),
  inference(superposition,[],[f374,f343])).
fof(f6733,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3)))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X1,fun(X0,bool)),X2),s('type.pair.prod'(X1,X0),sK25(X2,X0,X1)))))) )),
  inference(superposition,[],[f374,f280])).
fof(f6701,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,happ(s(fun(X23,bool),happ(s(fun(X25,fun(X23,bool)),X26),s(X25,happ(s(fun(X24,X25),X27),s(X24,X28))))),s(X23,X29)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X24,fun(X23,bool)),'_dst_x0x1_2'(s(fun(X25,fun(X23,bool)),X26),s(fun(X24,X25),X27))),s('type.pair.prod'(X24,X23),sK25('_dst_x0x1_2'(s(fun(X25,fun(X23,bool)),X26),s(fun(X24,X25),X27)),X23,X24)))))) )),
  inference(superposition,[],[f374,f344])).
fof(f6700,plain,(
  ( ! [X21,X19,X17,X22,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X19,fun(X16,bool)),'_dst_x0x1_2'(s(fun(X17,fun(X16,bool)),X18),s(fun(X19,X17),X20))),s(X19,X21))),s(X16,X22)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X17,fun(X16,bool)),X18),s('type.pair.prod'(X17,X16),sK25(X18,X16,X17)))))) )),
  inference(superposition,[],[f374,f344])).
fof(f6699,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X10,X11),fun(X9,bool)),X12),s(X10,X13),s(X11,X14))),s(X9,X15)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.pair.prod'(X10,X11),fun(X9,bool)),X12),s('type.pair.prod'('type.pair.prod'(X10,X11),X9),sK25(X12,X9,'type.pair.prod'(X10,X11))))))) )),
  inference(superposition,[],[f374,f343])).
fof(f6698,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),'const.pair.UNCURRY_2'(s(fun(X4,fun(X5,fun(X3,bool))),X6),s('type.pair.prod'(X4,X5),X7))),s(X3,X8)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.pair.prod'(X4,X5),fun(X3,bool)),'const.pair.UNCURRY_1'(s(fun(X4,fun(X5,fun(X3,bool))),X6))),s('type.pair.prod'('type.pair.prod'(X4,X5),X3),sK25('const.pair.UNCURRY_1'(s(fun(X4,fun(X5,fun(X3,bool))),X6)),X3,'type.pair.prod'(X4,X5))))))) )),
  inference(superposition,[],[f374,f335])).
fof(f6697,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s('type.pair.prod'(fun(X0,bool),X0),sK25('_i_0_0',X0,fun(X0,bool))))))) )),
  inference(superposition,[],[f374,f279])).
fof(f6747,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X71,fun(X72,bool)),sK28(X74,X73,fun(X71,fun(X72,bool)))),s('type.pair.prod'(X71,X72),sK25('const.pair.FST_1'(s('type.pair.prod'(fun(X71,fun(X72,bool)),X73),X74)),X72,X71))))) | ~p(s(bool,happ(s(fun(X72,bool),happ(s(fun(X71,fun(X72,bool)),sK28(X74,X73,fun(X71,fun(X72,bool)))),s(X71,X75))),s(X72,X76))))) )),
  inference(forward_demodulation,[],[f6665,f701])).
fof(f6665,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (~p(s(bool,happ(s(fun(X72,bool),happ(s(fun(X71,fun(X72,bool)),sK28(X74,X73,fun(X71,fun(X72,bool)))),s(X71,X75))),s(X72,X76)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X71,fun(X72,bool)),'const.pair.FST_1'(s('type.pair.prod'(fun(X71,fun(X72,bool)),X73),X74))),s('type.pair.prod'(X71,X72),sK25('const.pair.FST_1'(s('type.pair.prod'(fun(X71,fun(X72,bool)),X73),X74)),X72,X71)))))) )),
  inference(superposition,[],[f374,f701])).
fof(f6746,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X64,fun(X65,bool)),X67),s('type.pair.prod'(X64,X65),sK25('const.pair.FST_1'(s('type.pair.prod'(fun(X64,fun(X65,bool)),X66),'const.pair.,_2'(s(fun(X64,fun(X65,bool)),X67),s(X66,X68)))),X65,X64))))) | ~p(s(bool,happ(s(fun(X65,bool),happ(s(fun(X64,fun(X65,bool)),X67),s(X64,X69))),s(X65,X70))))) )),
  inference(forward_demodulation,[],[f6664,f323])).
fof(f6664,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun(X65,bool),happ(s(fun(X64,fun(X65,bool)),X67),s(X64,X69))),s(X65,X70)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X64,fun(X65,bool)),'const.pair.FST_1'(s('type.pair.prod'(fun(X64,fun(X65,bool)),X66),'const.pair.,_2'(s(fun(X64,fun(X65,bool)),X67),s(X66,X68))))),s('type.pair.prod'(X64,X65),sK25('const.pair.FST_1'(s('type.pair.prod'(fun(X64,fun(X65,bool)),X66),'const.pair.,_2'(s(fun(X64,fun(X65,bool)),X67),s(X66,X68)))),X65,X64)))))) )),
  inference(superposition,[],[f374,f323])).
fof(f6745,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X58,fun(X59,bool)),sK29(X61,fun(X58,fun(X59,bool)),X60)),s('type.pair.prod'(X58,X59),sK25('const.pair.SND_1'(s('type.pair.prod'(X60,fun(X58,fun(X59,bool))),X61)),X59,X58))))) | ~p(s(bool,happ(s(fun(X59,bool),happ(s(fun(X58,fun(X59,bool)),sK29(X61,fun(X58,fun(X59,bool)),X60)),s(X58,X62))),s(X59,X63))))) )),
  inference(forward_demodulation,[],[f6663,f702])).
fof(f6663,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (~p(s(bool,happ(s(fun(X59,bool),happ(s(fun(X58,fun(X59,bool)),sK29(X61,fun(X58,fun(X59,bool)),X60)),s(X58,X62))),s(X59,X63)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X58,fun(X59,bool)),'const.pair.SND_1'(s('type.pair.prod'(X60,fun(X58,fun(X59,bool))),X61))),s('type.pair.prod'(X58,X59),sK25('const.pair.SND_1'(s('type.pair.prod'(X60,fun(X58,fun(X59,bool))),X61)),X59,X58)))))) )),
  inference(superposition,[],[f374,f702])).
fof(f6744,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X51,fun(X52,bool)),X55),s('type.pair.prod'(X51,X52),sK25('const.pair.SND_1'(s('type.pair.prod'(X53,fun(X51,fun(X52,bool))),'const.pair.,_2'(s(X53,X54),s(fun(X51,fun(X52,bool)),X55)))),X52,X51))))) | ~p(s(bool,happ(s(fun(X52,bool),happ(s(fun(X51,fun(X52,bool)),X55),s(X51,X56))),s(X52,X57))))) )),
  inference(forward_demodulation,[],[f6662,f322])).
fof(f6662,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (~p(s(bool,happ(s(fun(X52,bool),happ(s(fun(X51,fun(X52,bool)),X55),s(X51,X56))),s(X52,X57)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X51,fun(X52,bool)),'const.pair.SND_1'(s('type.pair.prod'(X53,fun(X51,fun(X52,bool))),'const.pair.,_2'(s(X53,X54),s(fun(X51,fun(X52,bool)),X55))))),s('type.pair.prod'(X51,X52),sK25('const.pair.SND_1'(s('type.pair.prod'(X53,fun(X51,fun(X52,bool))),'const.pair.,_2'(s(X53,X54),s(fun(X51,fun(X52,bool)),X55)))),X52,X51)))))) )),
  inference(superposition,[],[f374,f322])).
fof(f6661,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X42,X49] : (~p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X42,fun(X43,bool)),happ(s(fun(X45,fun(X42,fun(X43,bool))),X46),s(X45,happ(s(fun(X44,X45),X47),s(X44,X48))))),s(X42,X49))),s(X43,X50)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X42,fun(X43,bool)),happ(s(fun(X44,fun(X42,fun(X43,bool))),'_dst_x0x1_2'(s(fun(X45,fun(X42,fun(X43,bool))),X46),s(fun(X44,X45),X47))),s(X44,X48))),s('type.pair.prod'(X42,X43),sK25(happ(s(fun(X44,fun(X42,fun(X43,bool))),'_dst_x0x1_2'(s(fun(X45,fun(X42,fun(X43,bool))),X46),s(fun(X44,X45),X47))),s(X44,X48)),X43,X42)))))) )),
  inference(superposition,[],[f374,f344])).
fof(f6660,plain,(
  ( ! [X39,X37,X35,X33,X41,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X34,bool),happ(s(fun(X33,fun(X34,bool)),happ(s(fun(X37,fun(X33,fun(X34,bool))),'_dst_x0x1_2'(s(fun(X35,fun(X33,fun(X34,bool))),X36),s(fun(X37,X35),X38))),s(X37,X39))),s(X33,X40))),s(X34,X41)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X33,fun(X34,bool)),happ(s(fun(X35,fun(X33,fun(X34,bool))),X36),s(X35,happ(s(fun(X37,X35),X38),s(X37,X39))))),s('type.pair.prod'(X33,X34),sK25(happ(s(fun(X35,fun(X33,fun(X34,bool))),X36),s(X35,happ(s(fun(X37,X35),X38),s(X37,X39)))),X34,X33)))))) )),
  inference(superposition,[],[f374,f344])).
fof(f6743,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25,X32] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,bool)),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X26,X27),fun(X24,fun(X25,bool))),X28),s(X26,X29),s(X27,X30))),s('type.pair.prod'(X24,X25),sK25(happ(s(fun('type.pair.prod'(X26,X27),fun(X24,fun(X25,bool))),X28),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X29),s(X27,X30)))),X25,X24))))) | ~p(s(bool,happ(s(fun(X25,bool),happ(s(fun(X24,fun(X25,bool)),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X26,X27),fun(X24,fun(X25,bool))),X28),s(X26,X29),s(X27,X30))),s(X24,X31))),s(X25,X32))))) )),
  inference(forward_demodulation,[],[f6659,f343])).
fof(f6659,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25,X32] : (~p(s(bool,happ(s(fun(X25,bool),happ(s(fun(X24,fun(X25,bool)),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X26,X27),fun(X24,fun(X25,bool))),X28),s(X26,X29),s(X27,X30))),s(X24,X31))),s(X25,X32)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,bool)),happ(s(fun('type.pair.prod'(X26,X27),fun(X24,fun(X25,bool))),X28),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X29),s(X27,X30))))),s('type.pair.prod'(X24,X25),sK25(happ(s(fun('type.pair.prod'(X26,X27),fun(X24,fun(X25,bool))),X28),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X29),s(X27,X30)))),X25,X24)))))) )),
  inference(superposition,[],[f374,f343])).
fof(f6742,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X16,fun(X17,bool)),'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,fun(X16,fun(X17,bool)))),X20),s('type.pair.prod'(X18,X19),X21))),s('type.pair.prod'(X16,X17),sK25(happ(s(fun('type.pair.prod'(X18,X19),fun(X16,fun(X17,bool))),'const.pair.UNCURRY_1'(s(fun(X18,fun(X19,fun(X16,fun(X17,bool)))),X20))),s('type.pair.prod'(X18,X19),X21)),X17,X16))))) | ~p(s(bool,happ(s(fun(X17,bool),happ(s(fun(X16,fun(X17,bool)),'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,fun(X16,fun(X17,bool)))),X20),s('type.pair.prod'(X18,X19),X21))),s(X16,X22))),s(X17,X23))))) )),
  inference(forward_demodulation,[],[f6658,f335])).
fof(f6658,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (~p(s(bool,happ(s(fun(X17,bool),happ(s(fun(X16,fun(X17,bool)),'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,fun(X16,fun(X17,bool)))),X20),s('type.pair.prod'(X18,X19),X21))),s(X16,X22))),s(X17,X23)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X16,fun(X17,bool)),happ(s(fun('type.pair.prod'(X18,X19),fun(X16,fun(X17,bool))),'const.pair.UNCURRY_1'(s(fun(X18,fun(X19,fun(X16,fun(X17,bool)))),X20))),s('type.pair.prod'(X18,X19),X21))),s('type.pair.prod'(X16,X17),sK25(happ(s(fun('type.pair.prod'(X18,X19),fun(X16,fun(X17,bool))),'const.pair.UNCURRY_1'(s(fun(X18,fun(X19,fun(X16,fun(X17,bool)))),X20))),s('type.pair.prod'(X18,X19),X21)),X17,X16)))))) )),
  inference(superposition,[],[f374,f335])).
fof(f6741,plain,(
  ( ! [X14,X12,X15,X13,X11] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X11,fun(X12,bool)),X13),s('type.pair.prod'(X11,X12),sK25(happ(s(fun(fun(X11,fun(X12,bool)),fun(X11,fun(X12,bool))),'_i_0_0'),s(fun(X11,fun(X12,bool)),X13)),X12,X11))))) | ~p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X13),s(X11,X14))),s(X12,X15))))) )),
  inference(forward_demodulation,[],[f6657,f279])).
fof(f6657,plain,(
  ( ! [X14,X12,X15,X13,X11] : (~p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X13),s(X11,X14))),s(X12,X15)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X11,fun(X12,bool)),happ(s(fun(fun(X11,fun(X12,bool)),fun(X11,fun(X12,bool))),'_i_0_0'),s(fun(X11,fun(X12,bool)),X13))),s('type.pair.prod'(X11,X12),sK25(happ(s(fun(fun(X11,fun(X12,bool)),fun(X11,fun(X12,bool))),'_i_0_0'),s(fun(X11,fun(X12,bool)),X13)),X12,X11)))))) )),
  inference(superposition,[],[f374,f279])).
fof(f6740,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8),s('type.pair.prod'('type.pair.prod'(X5,X6),X7),sK25('const.pair.UNCURRY_1'(s(fun(X5,fun(X6,fun(X7,bool))),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8)))),X7,'type.pair.prod'(X5,X6)))))) | ~p(s(bool,happ(s(fun(X7,bool),happ(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8),s('type.pair.prod'(X5,X6),X9))),s(X7,X10))))) )),
  inference(forward_demodulation,[],[f6656,f333])).
fof(f6656,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X7,bool),happ(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8),s('type.pair.prod'(X5,X6),X9))),s(X7,X10)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),'const.pair.UNCURRY_1'(s(fun(X5,fun(X6,fun(X7,bool))),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8))))),s('type.pair.prod'('type.pair.prod'(X5,X6),X7),sK25('const.pair.UNCURRY_1'(s(fun(X5,fun(X6,fun(X7,bool))),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8)))),X7,'type.pair.prod'(X5,X6))))))) )),
  inference(superposition,[],[f374,f333])).
fof(f6739,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),sK25('const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))),X1,X0))))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))) )),
  inference(forward_demodulation,[],[f6738,f335])).
fof(f6738,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))),s('type.pair.prod'(X0,X1),sK25('const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))),X1,X0))))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))) )),
  inference(forward_demodulation,[],[f6655,f5249])).
fof(f6655,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))),s('type.pair.prod'(X0,X1),sK25('const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))),X1,X0)))))) )),
  inference(superposition,[],[f374,f332])).
fof(f6579,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (p(s(bool,happ(s(fun(X72,bool),happ(s(fun(X71,fun(X72,bool)),sK28(X74,X73,fun(X71,fun(X72,bool)))),s(X71,X75))),s(X72,X76)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X71,fun(X72,bool)),sK28(X74,X73,fun(X71,fun(X72,bool)))),s('type.pair.prod'(X71,X72),sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X71,fun(X72,bool)),X73),X74)),X72,X71)))))) )),
  inference(forward_demodulation,[],[f6567,f701])).
fof(f6567,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X71,fun(X72,bool)),sK28(X74,X73,fun(X71,fun(X72,bool)))),s('type.pair.prod'(X71,X72),sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X71,fun(X72,bool)),X73),X74)),X72,X71))))) | p(s(bool,happ(s(fun(X72,bool),happ(s(fun(X71,fun(X72,bool)),'const.pair.FST_1'(s('type.pair.prod'(fun(X71,fun(X72,bool)),X73),X74))),s(X71,X75))),s(X72,X76))))) )),
  inference(superposition,[],[f373,f701])).
fof(f6578,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X65,bool),happ(s(fun(X64,fun(X65,bool)),X67),s(X64,X69))),s(X65,X70)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X64,fun(X65,bool)),X67),s('type.pair.prod'(X64,X65),sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X64,fun(X65,bool)),X66),'const.pair.,_2'(s(fun(X64,fun(X65,bool)),X67),s(X66,X68)))),X65,X64)))))) )),
  inference(forward_demodulation,[],[f6566,f323])).
fof(f6566,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X64,fun(X65,bool)),X67),s('type.pair.prod'(X64,X65),sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X64,fun(X65,bool)),X66),'const.pair.,_2'(s(fun(X64,fun(X65,bool)),X67),s(X66,X68)))),X65,X64))))) | p(s(bool,happ(s(fun(X65,bool),happ(s(fun(X64,fun(X65,bool)),'const.pair.FST_1'(s('type.pair.prod'(fun(X64,fun(X65,bool)),X66),'const.pair.,_2'(s(fun(X64,fun(X65,bool)),X67),s(X66,X68))))),s(X64,X69))),s(X65,X70))))) )),
  inference(superposition,[],[f373,f323])).
fof(f6577,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (p(s(bool,happ(s(fun(X59,bool),happ(s(fun(X58,fun(X59,bool)),sK29(X61,fun(X58,fun(X59,bool)),X60)),s(X58,X62))),s(X59,X63)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X58,fun(X59,bool)),sK29(X61,fun(X58,fun(X59,bool)),X60)),s('type.pair.prod'(X58,X59),sK15('const.pair.SND_1'(s('type.pair.prod'(X60,fun(X58,fun(X59,bool))),X61)),X59,X58)))))) )),
  inference(forward_demodulation,[],[f6565,f702])).
fof(f6565,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X58,fun(X59,bool)),sK29(X61,fun(X58,fun(X59,bool)),X60)),s('type.pair.prod'(X58,X59),sK15('const.pair.SND_1'(s('type.pair.prod'(X60,fun(X58,fun(X59,bool))),X61)),X59,X58))))) | p(s(bool,happ(s(fun(X59,bool),happ(s(fun(X58,fun(X59,bool)),'const.pair.SND_1'(s('type.pair.prod'(X60,fun(X58,fun(X59,bool))),X61))),s(X58,X62))),s(X59,X63))))) )),
  inference(superposition,[],[f373,f702])).
fof(f6576,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (p(s(bool,happ(s(fun(X52,bool),happ(s(fun(X51,fun(X52,bool)),X55),s(X51,X56))),s(X52,X57)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X51,fun(X52,bool)),X55),s('type.pair.prod'(X51,X52),sK15('const.pair.SND_1'(s('type.pair.prod'(X53,fun(X51,fun(X52,bool))),'const.pair.,_2'(s(X53,X54),s(fun(X51,fun(X52,bool)),X55)))),X52,X51)))))) )),
  inference(forward_demodulation,[],[f6564,f322])).
fof(f6564,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X51,fun(X52,bool)),X55),s('type.pair.prod'(X51,X52),sK15('const.pair.SND_1'(s('type.pair.prod'(X53,fun(X51,fun(X52,bool))),'const.pair.,_2'(s(X53,X54),s(fun(X51,fun(X52,bool)),X55)))),X52,X51))))) | p(s(bool,happ(s(fun(X52,bool),happ(s(fun(X51,fun(X52,bool)),'const.pair.SND_1'(s('type.pair.prod'(X53,fun(X51,fun(X52,bool))),'const.pair.,_2'(s(X53,X54),s(fun(X51,fun(X52,bool)),X55))))),s(X51,X56))),s(X52,X57))))) )),
  inference(superposition,[],[f373,f322])).
fof(f6563,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X42,X49] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X42,fun(X43,bool)),happ(s(fun(X45,fun(X42,fun(X43,bool))),X46),s(X45,happ(s(fun(X44,X45),X47),s(X44,X48))))),s('type.pair.prod'(X42,X43),sK15(happ(s(fun(X44,fun(X42,fun(X43,bool))),'_dst_x0x1_2'(s(fun(X45,fun(X42,fun(X43,bool))),X46),s(fun(X44,X45),X47))),s(X44,X48)),X43,X42))))) | p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X42,fun(X43,bool)),happ(s(fun(X44,fun(X42,fun(X43,bool))),'_dst_x0x1_2'(s(fun(X45,fun(X42,fun(X43,bool))),X46),s(fun(X44,X45),X47))),s(X44,X48))),s(X42,X49))),s(X43,X50))))) )),
  inference(superposition,[],[f373,f344])).
fof(f6562,plain,(
  ( ! [X39,X37,X35,X33,X41,X38,X36,X34,X40] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X33,fun(X34,bool)),happ(s(fun(X37,fun(X33,fun(X34,bool))),'_dst_x0x1_2'(s(fun(X35,fun(X33,fun(X34,bool))),X36),s(fun(X37,X35),X38))),s(X37,X39))),s('type.pair.prod'(X33,X34),sK15(happ(s(fun(X35,fun(X33,fun(X34,bool))),X36),s(X35,happ(s(fun(X37,X35),X38),s(X37,X39)))),X34,X33))))) | p(s(bool,happ(s(fun(X34,bool),happ(s(fun(X33,fun(X34,bool)),happ(s(fun(X35,fun(X33,fun(X34,bool))),X36),s(X35,happ(s(fun(X37,X35),X38),s(X37,X39))))),s(X33,X40))),s(X34,X41))))) )),
  inference(superposition,[],[f373,f344])).
fof(f6575,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25,X32] : (p(s(bool,happ(s(fun(X25,bool),happ(s(fun(X24,fun(X25,bool)),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X26,X27),fun(X24,fun(X25,bool))),X28),s(X26,X29),s(X27,X30))),s(X24,X31))),s(X25,X32)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,bool)),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X26,X27),fun(X24,fun(X25,bool))),X28),s(X26,X29),s(X27,X30))),s('type.pair.prod'(X24,X25),sK15(happ(s(fun('type.pair.prod'(X26,X27),fun(X24,fun(X25,bool))),X28),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X29),s(X27,X30)))),X25,X24)))))) )),
  inference(forward_demodulation,[],[f6561,f343])).
fof(f6561,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25,X32] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,bool)),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X26,X27),fun(X24,fun(X25,bool))),X28),s(X26,X29),s(X27,X30))),s('type.pair.prod'(X24,X25),sK15(happ(s(fun('type.pair.prod'(X26,X27),fun(X24,fun(X25,bool))),X28),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X29),s(X27,X30)))),X25,X24))))) | p(s(bool,happ(s(fun(X25,bool),happ(s(fun(X24,fun(X25,bool)),happ(s(fun('type.pair.prod'(X26,X27),fun(X24,fun(X25,bool))),X28),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X29),s(X27,X30))))),s(X24,X31))),s(X25,X32))))) )),
  inference(superposition,[],[f373,f343])).
fof(f6574,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,happ(s(fun(X17,bool),happ(s(fun(X16,fun(X17,bool)),'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,fun(X16,fun(X17,bool)))),X20),s('type.pair.prod'(X18,X19),X21))),s(X16,X22))),s(X17,X23)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X16,fun(X17,bool)),'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,fun(X16,fun(X17,bool)))),X20),s('type.pair.prod'(X18,X19),X21))),s('type.pair.prod'(X16,X17),sK15(happ(s(fun('type.pair.prod'(X18,X19),fun(X16,fun(X17,bool))),'const.pair.UNCURRY_1'(s(fun(X18,fun(X19,fun(X16,fun(X17,bool)))),X20))),s('type.pair.prod'(X18,X19),X21)),X17,X16)))))) )),
  inference(forward_demodulation,[],[f6560,f335])).
fof(f6560,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X16,fun(X17,bool)),'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,fun(X16,fun(X17,bool)))),X20),s('type.pair.prod'(X18,X19),X21))),s('type.pair.prod'(X16,X17),sK15(happ(s(fun('type.pair.prod'(X18,X19),fun(X16,fun(X17,bool))),'const.pair.UNCURRY_1'(s(fun(X18,fun(X19,fun(X16,fun(X17,bool)))),X20))),s('type.pair.prod'(X18,X19),X21)),X17,X16))))) | p(s(bool,happ(s(fun(X17,bool),happ(s(fun(X16,fun(X17,bool)),happ(s(fun('type.pair.prod'(X18,X19),fun(X16,fun(X17,bool))),'const.pair.UNCURRY_1'(s(fun(X18,fun(X19,fun(X16,fun(X17,bool)))),X20))),s('type.pair.prod'(X18,X19),X21))),s(X16,X22))),s(X17,X23))))) )),
  inference(superposition,[],[f373,f335])).
fof(f6573,plain,(
  ( ! [X14,X12,X15,X13,X11] : (p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X13),s(X11,X14))),s(X12,X15)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X11,fun(X12,bool)),X13),s('type.pair.prod'(X11,X12),sK15(happ(s(fun(fun(X11,fun(X12,bool)),fun(X11,fun(X12,bool))),'_i_0_0'),s(fun(X11,fun(X12,bool)),X13)),X12,X11)))))) )),
  inference(forward_demodulation,[],[f6559,f279])).
fof(f6559,plain,(
  ( ! [X14,X12,X15,X13,X11] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X11,fun(X12,bool)),X13),s('type.pair.prod'(X11,X12),sK15(happ(s(fun(fun(X11,fun(X12,bool)),fun(X11,fun(X12,bool))),'_i_0_0'),s(fun(X11,fun(X12,bool)),X13)),X12,X11))))) | p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),happ(s(fun(fun(X11,fun(X12,bool)),fun(X11,fun(X12,bool))),'_i_0_0'),s(fun(X11,fun(X12,bool)),X13))),s(X11,X14))),s(X12,X15))))) )),
  inference(superposition,[],[f373,f279])).
fof(f6572,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8),s('type.pair.prod'(X5,X6),X9))),s(X7,X10)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8),s('type.pair.prod'('type.pair.prod'(X5,X6),X7),sK15('const.pair.UNCURRY_1'(s(fun(X5,fun(X6,fun(X7,bool))),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8)))),X7,'type.pair.prod'(X5,X6))))))) )),
  inference(forward_demodulation,[],[f6571,f5249])).
fof(f6571,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(X7,bool),'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,fun(X7,bool))),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8))),s('type.pair.prod'(X5,X6),X9))),s(X7,X10)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8),s('type.pair.prod'('type.pair.prod'(X5,X6),X7),sK15('const.pair.UNCURRY_1'(s(fun(X5,fun(X6,fun(X7,bool))),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8)))),X7,'type.pair.prod'(X5,X6))))))) )),
  inference(forward_demodulation,[],[f6558,f335])).
fof(f6558,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8),s('type.pair.prod'('type.pair.prod'(X5,X6),X7),sK15('const.pair.UNCURRY_1'(s(fun(X5,fun(X6,fun(X7,bool))),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8)))),X7,'type.pair.prod'(X5,X6)))))) | p(s(bool,happ(s(fun(X7,bool),happ(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),'const.pair.UNCURRY_1'(s(fun(X5,fun(X6,fun(X7,bool))),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X5,X6),fun(X7,bool)),X8))))),s('type.pair.prod'(X5,X6),X9))),s(X7,X10))))) )),
  inference(superposition,[],[f373,f333])).
fof(f6570,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),sK15('const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))),X1,X0)))))) )),
  inference(forward_demodulation,[],[f6569,f348])).
fof(f6569,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))),s(X0,X3),s(X1,X4)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),sK15('const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))),X1,X0)))))) )),
  inference(forward_demodulation,[],[f6557,f347])).
fof(f6557,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),sK15('const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))),X1,X0))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))),s(X0,X3))),s(X1,X4))))) )),
  inference(superposition,[],[f373,f332])).
fof(f6492,plain,(
  ( ! [X499,X494,X500,X496,X498,X495,X497] : (s(X494,'const.pair.FST_1'(s('type.pair.prod'(X494,X495),happ(s(fun(X497,'type.pair.prod'(X494,X495)),X498),s(X497,happ(s(fun(X496,X497),X499),s(X496,X500))))))) = s(X494,sK28(happ(s(fun(X496,'type.pair.prod'(X494,X495)),'_dst_x0x1_2'(s(fun(X497,'type.pair.prod'(X494,X495)),X498),s(fun(X496,X497),X499))),s(X496,X500)),X495,X494))) )),
  inference(superposition,[],[f701,f344])).
fof(f6491,plain,(
  ( ! [X492,X488,X490,X487,X493,X489,X491] : (s(X488,'const.pair.SND_1'(s('type.pair.prod'(X487,X488),happ(s(fun(X490,'type.pair.prod'(X487,X488)),X491),s(X490,happ(s(fun(X489,X490),X492),s(X489,X493))))))) = s(X488,sK29(happ(s(fun(X489,'type.pair.prod'(X487,X488)),'_dst_x0x1_2'(s(fun(X490,'type.pair.prod'(X487,X488)),X491),s(fun(X489,X490),X492))),s(X489,X493)),X488,X487))) )),
  inference(superposition,[],[f702,f344])).
fof(f6484,plain,(
  ( ! [X428,X430,X426,X432,X434,X429,X431,X427,X433] : (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X426,X427),bool),'const.pair.UNCURRY_1'(s(fun(X426,fun(X427,bool)),happ(s(fun(X429,fun(X426,fun(X427,bool))),X430),s(X429,happ(s(fun(X428,X429),X431),s(X428,X432)))))))))) | p(s(bool,happ(s(fun(X427,bool),happ(s(fun(X426,fun(X427,bool)),happ(s(fun(X428,fun(X426,fun(X427,bool))),'_dst_x0x1_2'(s(fun(X429,fun(X426,fun(X427,bool))),X430),s(fun(X428,X429),X431))),s(X428,X432))),s(X426,X433))),s(X427,X434))))) )),
  inference(superposition,[],[f314,f344])).
fof(f6483,plain,(
  ( ! [X420,X422,X418,X424,X421,X423,X417,X419,X425] : (~p(s(bool,happ(s(fun(X418,bool),happ(s(fun(X417,fun(X418,bool)),happ(s(fun(X420,fun(X417,fun(X418,bool))),X421),s(X420,happ(s(fun(X419,X420),X422),s(X419,X423))))),s(X417,X424))),s(X418,X425)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X417,X418),bool),'const.pair.UNCURRY_1'(s(fun(X417,fun(X418,bool)),happ(s(fun(X419,fun(X417,fun(X418,bool))),'_dst_x0x1_2'(s(fun(X420,fun(X417,fun(X418,bool))),X421),s(fun(X419,X420),X422))),s(X419,X423))))))))) )),
  inference(superposition,[],[f319,f344])).
fof(f6482,plain,(
  ( ! [X412,X414,X408,X410,X416,X413,X415,X409,X411] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X408,X409),bool),happ(s(fun(X411,fun('type.pair.prod'(X408,X409),bool)),X412),s(X411,happ(s(fun(X410,X411),X413),s(X410,X414))))),s(X408,X415),s(X409,X416)))) | p(s(bool,happ(s(fun('type.pair.prod'(X408,X409),bool),happ(s(fun(X410,fun('type.pair.prod'(X408,X409),bool)),'_dst_x0x1_2'(s(fun(X411,fun('type.pair.prod'(X408,X409),bool)),X412),s(fun(X410,X411),X413))),s(X410,X414))),s('type.pair.prod'(X408,X409),sK22(happ(s(fun(X410,fun('type.pair.prod'(X408,X409),bool)),'_dst_x0x1_2'(s(fun(X411,fun('type.pair.prod'(X408,X409),bool)),X412),s(fun(X410,X411),X413))),s(X410,X414)),X409,X408)))))) )),
  inference(superposition,[],[f375,f344])).
fof(f6481,plain,(
  ( ! [X399,X405,X407,X401,X403,X404,X406,X400,X402] : (~p(s(bool,happ(s(fun('type.pair.prod'(X399,X400),bool),happ(s(fun(X402,fun('type.pair.prod'(X399,X400),bool)),X403),s(X402,happ(s(fun(X401,X402),X404),s(X401,X405))))),s('type.pair.prod'(X399,X400),sK12(happ(s(fun(X401,fun('type.pair.prod'(X399,X400),bool)),'_dst_x0x1_2'(s(fun(X402,fun('type.pair.prod'(X399,X400),bool)),X403),s(fun(X401,X402),X404))),s(X401,X405)),X400,X399))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X399,X400),bool),happ(s(fun(X401,fun('type.pair.prod'(X399,X400),bool)),'_dst_x0x1_2'(s(fun(X402,fun('type.pair.prod'(X399,X400),bool)),X403),s(fun(X401,X402),X404))),s(X401,X405))),s(X399,X406),s(X400,X407))))) )),
  inference(superposition,[],[f376,f344])).
fof(f6480,plain,(
  ( ! [X391,X397,X393,X395,X396,X398,X392,X394] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X391,X392),bool),happ(s(fun(X394,fun('type.pair.prod'(X391,X392),bool)),X395),s(X394,happ(s(fun(X393,X394),X396),s(X393,X397))))),s(X391,sK10(happ(s(fun(X393,fun('type.pair.prod'(X391,X392),bool)),'_dst_x0x1_2'(s(fun(X394,fun('type.pair.prod'(X391,X392),bool)),X395),s(fun(X393,X394),X396))),s(X393,X397)),X392,X391)),s(X392,sK11(happ(s(fun(X393,fun('type.pair.prod'(X391,X392),bool)),'_dst_x0x1_2'(s(fun(X394,fun('type.pair.prod'(X391,X392),bool)),X395),s(fun(X393,X394),X396))),s(X393,X397)),X392,X391))))) | p(s(bool,happ(s(fun('type.pair.prod'(X391,X392),bool),happ(s(fun(X393,fun('type.pair.prod'(X391,X392),bool)),'_dst_x0x1_2'(s(fun(X394,fun('type.pair.prod'(X391,X392),bool)),X395),s(fun(X393,X394),X396))),s(X393,X397))),s('type.pair.prod'(X391,X392),X398))))) )),
  inference(superposition,[],[f377,f344])).
fof(f6479,plain,(
  ( ! [X383,X389,X385,X387,X388,X390,X384,X386] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X383,X384),bool),happ(s(fun(X386,fun('type.pair.prod'(X383,X384),bool)),X387),s(X386,happ(s(fun(X385,X386),X388),s(X385,X389))))),s(X383,sK13(happ(s(fun(X385,fun('type.pair.prod'(X383,X384),bool)),'_dst_x0x1_2'(s(fun(X386,fun('type.pair.prod'(X383,X384),bool)),X387),s(fun(X385,X386),X388))),s(X385,X389)),X384,X383)),s(X384,sK14(happ(s(fun(X385,fun('type.pair.prod'(X383,X384),bool)),'_dst_x0x1_2'(s(fun(X386,fun('type.pair.prod'(X383,X384),bool)),X387),s(fun(X385,X386),X388))),s(X385,X389)),X384,X383))))) | p(s(bool,happ(s(fun('type.pair.prod'(X383,X384),bool),happ(s(fun(X385,fun('type.pair.prod'(X383,X384),bool)),'_dst_x0x1_2'(s(fun(X386,fun('type.pair.prod'(X383,X384),bool)),X387),s(fun(X385,X386),X388))),s(X385,X389))),s('type.pair.prod'(X383,X384),X390))))) )),
  inference(superposition,[],[f378,f344])).
fof(f6478,plain,(
  ( ! [X377,X379,X380,X382,X376,X378,X375,X381] : (~p(s(bool,happ(s(fun('type.pair.prod'(X375,X376),bool),happ(s(fun(X378,fun('type.pair.prod'(X375,X376),bool)),X379),s(X378,happ(s(fun(X377,X378),X380),s(X377,X381))))),s('type.pair.prod'(X375,X376),X382)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X375,X376),bool),happ(s(fun(X377,fun('type.pair.prod'(X375,X376),bool)),'_dst_x0x1_2'(s(fun(X378,fun('type.pair.prod'(X375,X376),bool)),X379),s(fun(X377,X378),X380))),s(X377,X381))),s(X375,sK20(happ(s(fun(X377,fun('type.pair.prod'(X375,X376),bool)),'_dst_x0x1_2'(s(fun(X378,fun('type.pair.prod'(X375,X376),bool)),X379),s(fun(X377,X378),X380))),s(X377,X381)),X376,X375)),s(X376,sK21(happ(s(fun(X377,fun('type.pair.prod'(X375,X376),bool)),'_dst_x0x1_2'(s(fun(X378,fun('type.pair.prod'(X375,X376),bool)),X379),s(fun(X377,X378),X380))),s(X377,X381)),X376,X375)))))) )),
  inference(superposition,[],[f379,f344])).
fof(f6477,plain,(
  ( ! [X372,X374,X368,X370,X373,X369,X371] : (~p(s(bool,happ(s(fun(X368,bool),happ(s(fun(X370,fun(X368,bool)),X371),s(X370,happ(s(fun(X369,X370),X372),s(X369,X373))))),s(X368,X374)))) | p(s(bool,happ(s(fun(fun(X368,bool),bool),'?_1'),s(fun(X368,bool),happ(s(fun(X369,fun(X368,bool)),'_dst_x0x1_2'(s(fun(X370,fun(X368,bool)),X371),s(fun(X369,X370),X372))),s(X369,X373))))))) )),
  inference(superposition,[],[f303,f344])).
fof(f6476,plain,(
  ( ! [X364,X366,X362,X365,X367,X363] : (~p(s(bool,happ(s(fun(fun(X362,bool),bool),'?_1'),s(fun(X362,bool),happ(s(fun(X364,fun(X362,bool)),X365),s(X364,happ(s(fun(X363,X364),X366),s(X363,X367)))))))) | p(s(bool,happ(s(fun(X362,bool),happ(s(fun(X363,fun(X362,bool)),'_dst_x0x1_2'(s(fun(X364,fun(X362,bool)),X365),s(fun(X363,X364),X366))),s(X363,X367))),s(X362,sK9(happ(s(fun(X363,fun(X362,bool)),'_dst_x0x1_2'(s(fun(X364,fun(X362,bool)),X365),s(fun(X363,X364),X366))),s(X363,X367)),X362)))))) )),
  inference(superposition,[],[f302,f344])).
fof(f6475,plain,(
  ( ! [X356,X358,X360,X357,X359,X361] : (~p(s(bool,happ(s(fun(X356,bool),happ(s(fun(X358,fun(X356,bool)),X359),s(X358,happ(s(fun(X357,X358),X360),s(X357,X361))))),s(X356,sK8(happ(s(fun(X357,fun(X356,bool)),'_dst_x0x1_2'(s(fun(X358,fun(X356,bool)),X359),s(fun(X357,X358),X360))),s(X357,X361)),X356))))) | p(s(bool,happ(s(fun(fun(X356,bool),bool),'!_1'),s(fun(X356,bool),happ(s(fun(X357,fun(X356,bool)),'_dst_x0x1_2'(s(fun(X358,fun(X356,bool)),X359),s(fun(X357,X358),X360))),s(X357,X361))))))) )),
  inference(superposition,[],[f301,f344])).
fof(f6474,plain,(
  ( ! [X350,X352,X354,X349,X351,X353,X355] : (~p(s(bool,happ(s(fun(fun(X349,bool),bool),'!_1'),s(fun(X349,bool),happ(s(fun(X351,fun(X349,bool)),X352),s(X351,happ(s(fun(X350,X351),X353),s(X350,X354)))))))) | p(s(bool,happ(s(fun(X349,bool),happ(s(fun(X350,fun(X349,bool)),'_dst_x0x1_2'(s(fun(X351,fun(X349,bool)),X352),s(fun(X350,X351),X353))),s(X350,X354))),s(X349,X355))))) )),
  inference(superposition,[],[f300,f344])).
fof(f6473,plain,(
  ( ! [X344,X346,X343,X345,X347,X342,X348] : (p(s(bool,'|dollar|exists_1'(s(fun(X342,bool),happ(s(fun(X344,fun(X342,bool)),X345),s(X344,happ(s(fun(X343,X344),X346),s(X343,X347)))))))) | s(X342,X348) = s(X342,sK7(X348,happ(s(fun(X343,fun(X342,bool)),'_dst_x0x1_2'(s(fun(X344,fun(X342,bool)),X345),s(fun(X343,X344),X346))),s(X343,X347)),X342))) )),
  inference(superposition,[],[f292,f344])).
fof(f6472,plain,(
  ( ! [X335,X341,X337,X339,X340,X336,X338] : (~p(s(bool,happ(s(fun(X335,bool),happ(s(fun(X337,fun(X335,bool)),X338),s(X337,happ(s(fun(X336,X337),X339),s(X336,X340))))),s(X335,X341)))) | p(s(bool,'|dollar|exists_1'(s(fun(X335,bool),happ(s(fun(X336,fun(X335,bool)),'_dst_x0x1_2'(s(fun(X337,fun(X335,bool)),X338),s(fun(X336,X337),X339))),s(X336,X340))))))) )),
  inference(superposition,[],[f289,f344])).
fof(f6471,plain,(
  ( ! [X333,X329,X331,X332,X334,X330] : (p(s(bool,happ(s(fun(X329,bool),happ(s(fun(X331,fun(X329,bool)),X332),s(X331,happ(s(fun(X330,X331),X333),s(X330,X334))))),s(X329,sK6(happ(s(fun(X330,fun(X329,bool)),'_dst_x0x1_2'(s(fun(X331,fun(X329,bool)),X332),s(fun(X330,X331),X333))),s(X330,X334)),X329))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X329,bool),happ(s(fun(X330,fun(X329,bool)),'_dst_x0x1_2'(s(fun(X331,fun(X329,bool)),X332),s(fun(X330,X331),X333))),s(X330,X334))))))) )),
  inference(superposition,[],[f288,f344])).
fof(f6470,plain,(
  ( ! [X325,X327,X323,X324,X326,X328] : (~p(s(bool,'|dollar|exists_1'(s(fun(X323,bool),happ(s(fun(X325,fun(X323,bool)),X326),s(X325,happ(s(fun(X324,X325),X327),s(X324,X328)))))))) | p(s(bool,happ(s(fun(X323,bool),happ(s(fun(X324,fun(X323,bool)),'_dst_x0x1_2'(s(fun(X325,fun(X323,bool)),X326),s(fun(X324,X325),X327))),s(X324,X328))),s(X323,sK5(happ(s(fun(X324,fun(X323,bool)),'_dst_x0x1_2'(s(fun(X325,fun(X323,bool)),X326),s(fun(X324,X325),X327))),s(X324,X328)),X323))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X323,bool),happ(s(fun(X324,fun(X323,bool)),'_dst_x0x1_2'(s(fun(X325,fun(X323,bool)),X326),s(fun(X324,X325),X327))),s(X324,X328))))))) )),
  inference(superposition,[],[f286,f344])).
fof(f6469,plain,(
  ( ! [X321,X318,X320,X322,X317,X319] : (~p(s(bool,'|dollar|exists_1'(s(fun(X317,bool),happ(s(fun(X319,fun(X317,bool)),X320),s(X319,happ(s(fun(X318,X319),X321),s(X318,X322)))))))) | p(s(bool,happ(s(fun(X317,bool),happ(s(fun(X318,fun(X317,bool)),'_dst_x0x1_2'(s(fun(X319,fun(X317,bool)),X320),s(fun(X318,X319),X321))),s(X318,X322))),s(X317,sK4(happ(s(fun(X318,fun(X317,bool)),'_dst_x0x1_2'(s(fun(X319,fun(X317,bool)),X320),s(fun(X318,X319),X321))),s(X318,X322)),X317))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X317,bool),happ(s(fun(X318,fun(X317,bool)),'_dst_x0x1_2'(s(fun(X319,fun(X317,bool)),X320),s(fun(X318,X319),X321))),s(X318,X322))))))) )),
  inference(superposition,[],[f285,f344])).
fof(f6468,plain,(
  ( ! [X313,X315,X310,X316,X312,X314,X309,X311] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X309,bool),happ(s(fun(X311,fun(X309,bool)),X312),s(X311,happ(s(fun(X310,X311),X313),s(X310,X314)))))))) | ~p(s(bool,happ(s(fun(X309,bool),happ(s(fun(X310,fun(X309,bool)),'_dst_x0x1_2'(s(fun(X311,fun(X309,bool)),X312),s(fun(X310,X311),X313))),s(X310,X314))),s(X309,X315)))) | ~p(s(bool,happ(s(fun(X309,bool),happ(s(fun(X310,fun(X309,bool)),'_dst_x0x1_2'(s(fun(X311,fun(X309,bool)),X312),s(fun(X310,X311),X313))),s(X310,X314))),s(X309,X316)))) | s(X309,X315) = s(X309,X316)) )),
  inference(superposition,[],[f284,f344])).
fof(f6467,plain,(
  ( ! [X308,X304,X306,X303,X305,X307] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X303,bool),happ(s(fun(X305,fun(X303,bool)),X306),s(X305,happ(s(fun(X304,X305),X307),s(X304,X308)))))))) | p(s(bool,'|dollar|exists_1'(s(fun(X303,bool),happ(s(fun(X304,fun(X303,bool)),'_dst_x0x1_2'(s(fun(X305,fun(X303,bool)),X306),s(fun(X304,X305),X307))),s(X304,X308))))))) )),
  inference(superposition,[],[f283,f344])).
fof(f6465,plain,(
  ( ! [X292,X294,X290,X296,X293,X295,X291] : (~p(s(bool,happ(s(fun(X290,bool),happ(s(fun(X292,fun(X290,bool)),X293),s(X292,happ(s(fun(X291,X292),X294),s(X291,X295))))),s(X290,X296)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X291,X290),bool),'const.pair.UNCURRY_1'(s(fun(X291,fun(X290,bool)),'_dst_x0x1_2'(s(fun(X292,fun(X290,bool)),X293),s(fun(X291,X292),X294))))))))) )),
  inference(superposition,[],[f319,f344])).
fof(f6464,plain,(
  ( ! [X288,X285,X287,X289,X284,X286] : (p(s(bool,happ(s(fun(X285,bool),X286),s(X285,happ(s(fun(X284,X285),X287),s(X284,X288)))))) | s(bool,X289) = s(bool,happ(s(fun(X284,bool),'_dst_x0x1_2'(s(fun(X285,bool),X286),s(fun(X284,X285),X287))),s(X284,X288))) | s(bool,'T_0') = s(bool,X289)) )),
  inference(superposition,[],[f468,f344])).
fof(f6463,plain,(
  ( ! [X280,X282,X279,X281,X283,X278] : (~p(s(bool,happ(s(fun(X279,bool),X280),s(X279,happ(s(fun(X278,X279),X281),s(X278,X282)))))) | ~p(s(bool,X283)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun(X278,bool),'_dst_x0x1_2'(s(fun(X279,bool),X280),s(fun(X278,X279),X281))),s(X278,X282))))),s(bool,X283))))) )),
  inference(superposition,[],[f306,f344])).
fof(f6462,plain,(
  ( ! [X277,X273,X275,X276,X272,X274] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X277))),s(bool,happ(s(fun(X273,bool),X274),s(X273,happ(s(fun(X272,X273),X275),s(X272,X276)))))))) | p(s(bool,happ(s(fun(X272,bool),'_dst_x0x1_2'(s(fun(X273,bool),X274),s(fun(X272,X273),X275))),s(X272,X276))))) )),
  inference(superposition,[],[f305,f344])).
fof(f6459,plain,(
  ( ! [X255,X257,X259,X254,X256,X258] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun(X255,bool),X256),s(X255,happ(s(fun(X254,X255),X257),s(X254,X258))))))),s(bool,X259)))) | p(s(bool,happ(s(fun(X254,bool),'_dst_x0x1_2'(s(fun(X255,bool),X256),s(fun(X254,X255),X257))),s(X254,X258))))) )),
  inference(superposition,[],[f304,f344])).
fof(f6458,plain,(
  ( ! [X253,X251,X249,X252,X250,X248] : (~p(s(bool,happ(s(fun(X249,bool),X250),s(X249,happ(s(fun(X248,X249),X251),s(X248,X252)))))) | ~p(s(bool,X253)) | s(bool,X253) = s(bool,happ(s(fun(X248,bool),'_dst_x0x1_2'(s(fun(X249,bool),X250),s(fun(X248,X249),X251))),s(X248,X252)))) )),
  inference(superposition,[],[f299,f344])).
fof(f6456,plain,(
  ( ! [X241,X239,X237,X240,X238,X236] : (p(s(bool,happ(s(fun(X237,bool),X238),s(X237,happ(s(fun(X236,X237),X239),s(X236,X240)))))) | p(s(bool,X241)) | s(bool,X241) = s(bool,happ(s(fun(X236,bool),'_dst_x0x1_2'(s(fun(X237,bool),X238),s(fun(X236,X237),X239))),s(X236,X240)))) )),
  inference(superposition,[],[f296,f344])).
fof(f6455,plain,(
  ( ! [X231,X235,X233,X234,X232] : (p(s(bool,happ(s(fun(X232,bool),X233),s(X232,happ(s(fun(X231,X232),X234),s(X231,X235)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X231,bool),'_dst_x0x1_2'(s(fun(X232,bool),X233),s(fun(X231,X232),X234))),s(X231,X235)))) )),
  inference(superposition,[],[f270,f344])).
fof(f6454,plain,(
  ( ! [X229,X227,X230,X228,X226] : (~p(s(bool,happ(s(fun(X227,bool),X228),s(X227,happ(s(fun(X226,X227),X229),s(X226,X230)))))) | s(bool,'F_0') != s(bool,happ(s(fun(X226,bool),'_dst_x0x1_2'(s(fun(X227,bool),X228),s(fun(X226,X227),X229))),s(X226,X230)))) )),
  inference(superposition,[],[f269,f344])).
fof(f6453,plain,(
  ( ! [X222,X225,X223,X221,X224] : (~p(s(bool,happ(s(fun(X222,bool),X223),s(X222,happ(s(fun(X221,X222),X224),s(X221,X225)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X221,bool),'_dst_x0x1_2'(s(fun(X222,bool),X223),s(fun(X221,X222),X224))),s(X221,X225)))) )),
  inference(superposition,[],[f266,f344])).
fof(f6452,plain,(
  ( ! [X218,X216,X220,X219,X217] : (s(bool,'T_0') != s(bool,happ(s(fun(X217,bool),X218),s(X217,happ(s(fun(X216,X217),X219),s(X216,X220))))) | p(s(bool,happ(s(fun(X216,bool),'_dst_x0x1_2'(s(fun(X217,bool),X218),s(fun(X216,X217),X219))),s(X216,X220))))) )),
  inference(superposition,[],[f265,f344])).
fof(f6450,plain,(
  ( ! [X210,X208,X206,X209,X207,X205] : (~p(s(bool,happ(s(fun(X207,bool),X208),s(X207,happ(s(fun('type.pair.prod'(X205,X206),X207),X209),s('type.pair.prod'(X205,X206),X210)))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X205,X206),bool),'_dst_x0x1_2'(s(fun(X207,bool),X208),s(fun('type.pair.prod'(X205,X206),X207),X209))),s(X205,sK20('_dst_x0x1_2'(s(fun(X207,bool),X208),s(fun('type.pair.prod'(X205,X206),X207),X209)),X206,X205)),s(X206,sK21('_dst_x0x1_2'(s(fun(X207,bool),X208),s(fun('type.pair.prod'(X205,X206),X207),X209)),X206,X205)))))) )),
  inference(superposition,[],[f379,f344])).
fof(f6449,plain,(
  ( ! [X198,X204,X202,X200,X199,X203,X201] : (~p(s(bool,happ(s(fun(X200,bool),X201),s(X200,happ(s(fun('type.pair.prod'(X198,X199),X200),X202),s('type.pair.prod'(X198,X199),sK12('_dst_x0x1_2'(s(fun(X200,bool),X201),s(fun('type.pair.prod'(X198,X199),X200),X202)),X199,X198))))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X198,X199),bool),'_dst_x0x1_2'(s(fun(X200,bool),X201),s(fun('type.pair.prod'(X198,X199),X200),X202))),s(X198,X203),s(X199,X204))))) )),
  inference(superposition,[],[f376,f344])).
fof(f6448,plain,(
  ( ! [X196,X194,X197,X195] : (~p(s(bool,happ(s(fun(X195,bool),X196),s(X195,happ(s(fun(X194,X195),X197),s(X194,sK8('_dst_x0x1_2'(s(fun(X195,bool),X196),s(fun(X194,X195),X197)),X194))))))) | p(s(bool,happ(s(fun(fun(X194,bool),bool),'!_1'),s(fun(X194,bool),'_dst_x0x1_2'(s(fun(X195,bool),X196),s(fun(X194,X195),X197))))))) )),
  inference(superposition,[],[f301,f344])).
fof(f6447,plain,(
  ( ! [X191,X192,X190,X193] : (p(s(bool,happ(s(fun(X191,bool),X192),s(X191,happ(s(fun(X190,X191),X193),s(X190,sK6('_dst_x0x1_2'(s(fun(X191,bool),X192),s(fun(X190,X191),X193)),X190))))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X190,bool),'_dst_x0x1_2'(s(fun(X191,bool),X192),s(fun(X190,X191),X193))))))) )),
  inference(superposition,[],[f288,f344])).
fof(f6445,plain,(
  ( ! [X185,X182,X184,X183,X181] : (~p(s(bool,happ(s(fun(X182,bool),X183),s(X182,happ(s(fun(X181,X182),X184),s(X181,X185)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X181,bool),'_dst_x0x1_2'(s(fun(X182,bool),X183),s(fun(X181,X182),X184))))))) )),
  inference(superposition,[],[f289,f344])).
fof(f6444,plain,(
  ( ! [X177,X180,X178,X176,X179] : (~p(s(bool,happ(s(fun(X177,bool),X178),s(X177,happ(s(fun(X176,X177),X179),s(X176,X180)))))) | p(s(bool,happ(s(fun(fun(X176,bool),bool),'?_1'),s(fun(X176,bool),'_dst_x0x1_2'(s(fun(X177,bool),X178),s(fun(X176,X177),X179))))))) )),
  inference(superposition,[],[f303,f344])).
fof(f6439,plain,(
  ( ! [X134,X132,X136,X135,X133,X131,X137] : (p(s(bool,happ(s(fun(X131,bool),happ(s(fun(X131,fun(X131,bool)),'=_0'),s(X131,X137))),s(X131,happ(s(fun(X133,X131),X134),s(X133,happ(s(fun(X132,X133),X135),s(X132,X136)))))))) | s(X131,X137) != s(X131,happ(s(fun(X132,X131),'_dst_x0x1_2'(s(fun(X133,X131),X134),s(fun(X132,X133),X135))),s(X132,X136)))) )),
  inference(superposition,[],[f331,f344])).
fof(f6438,plain,(
  ( ! [X127,X125,X130,X128,X126,X124,X129] : (p(s(bool,happ(s(fun(X124,bool),happ(s(fun(X124,fun(X124,bool)),'=_0'),s(X124,happ(s(fun(X126,X124),X127),s(X126,happ(s(fun(X125,X126),X128),s(X125,X129))))))),s(X124,X130)))) | s(X124,X130) != s(X124,happ(s(fun(X125,X124),'_dst_x0x1_2'(s(fun(X126,X124),X127),s(fun(X125,X126),X128))),s(X125,X129)))) )),
  inference(superposition,[],[f331,f344])).
fof(f6437,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (~p(s(bool,happ(s(fun(X117,bool),happ(s(fun(X117,fun(X117,bool)),'=_0'),s(X117,X123))),s(X117,happ(s(fun(X119,X117),X120),s(X119,happ(s(fun(X118,X119),X121),s(X118,X122)))))))) | s(X117,X123) = s(X117,happ(s(fun(X118,X117),'_dst_x0x1_2'(s(fun(X119,X117),X120),s(fun(X118,X119),X121))),s(X118,X122)))) )),
  inference(superposition,[],[f330,f344])).
fof(f6436,plain,(
  ( ! [X111,X116,X114,X112,X110,X115,X113] : (~p(s(bool,happ(s(fun(X110,bool),happ(s(fun(X110,fun(X110,bool)),'=_0'),s(X110,happ(s(fun(X112,X110),X113),s(X112,happ(s(fun(X111,X112),X114),s(X111,X115))))))),s(X110,X116)))) | s(X110,X116) = s(X110,happ(s(fun(X111,X110),'_dst_x0x1_2'(s(fun(X112,X110),X113),s(fun(X111,X112),X114))),s(X111,X115)))) )),
  inference(superposition,[],[f330,f344])).
fof(f6435,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (~p(s(bool,happ(s(fun(X102,bool),happ(s(fun(X108,fun(X102,bool)),sK32(X102,X108)),s(X108,X109))),s(X102,happ(s(fun(X104,X102),X105),s(X104,happ(s(fun(X103,X104),X106),s(X103,X107)))))))) | s(X102,sK34(X102,X108)) = s(X102,happ(s(fun(X103,X102),'_dst_x0x1_2'(s(fun(X104,X102),X105),s(fun(X103,X104),X106))),s(X103,X107)))) )),
  inference(superposition,[],[f328,f344])).
fof(f6432,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X84,X82] : (~p(s(bool,happ(s(fun(X84,bool),happ(s(fun(X78,fun(X84,bool)),sK32(X84,X78)),s(X78,happ(s(fun(X80,X78),X81),s(X80,happ(s(fun(X79,X80),X82),s(X79,X83))))))),s(X84,X85)))) | s(X78,sK33(X84,X78)) = s(X78,happ(s(fun(X79,X78),'_dst_x0x1_2'(s(fun(X80,X78),X81),s(fun(X79,X80),X82))),s(X79,X83)))) )),
  inference(superposition,[],[f327,f344])).
fof(f6412,plain,(
  ( ! [X580,X582,X576,X578,X581,X577,X579] : (s(X576,sK28(happ(s(fun(X578,'type.pair.prod'(X576,X577)),X579),s(X578,happ(s(fun(X580,X578),X581),s(X580,X582)))),X577,X576)) = s(X576,'const.pair.FST_1'(s('type.pair.prod'(X576,X577),happ(s(fun(X580,'type.pair.prod'(X576,X577)),'_dst_x0x1_2'(s(fun(X578,'type.pair.prod'(X576,X577)),X579),s(fun(X580,X578),X581))),s(X580,X582)))))) )),
  inference(superposition,[],[f701,f344])).
fof(f6411,plain,(
  ( ! [X573,X575,X569,X571,X572,X574,X570] : (s(X570,sK29(happ(s(fun(X571,'type.pair.prod'(X569,X570)),X572),s(X571,happ(s(fun(X573,X571),X574),s(X573,X575)))),X570,X569)) = s(X570,'const.pair.SND_1'(s('type.pair.prod'(X569,X570),happ(s(fun(X573,'type.pair.prod'(X569,X570)),'_dst_x0x1_2'(s(fun(X571,'type.pair.prod'(X569,X570)),X572),s(fun(X573,X571),X574))),s(X573,X575)))))) )),
  inference(superposition,[],[f702,f344])).
fof(f6404,plain,(
  ( ! [X509,X511,X516,X512,X514,X508,X510,X513,X515] : (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X508,X509),bool),'const.pair.UNCURRY_1'(s(fun(X508,fun(X509,bool)),happ(s(fun(X512,fun(X508,fun(X509,bool))),'_dst_x0x1_2'(s(fun(X510,fun(X508,fun(X509,bool))),X511),s(fun(X512,X510),X513))),s(X512,X514)))))))) | p(s(bool,happ(s(fun(X509,bool),happ(s(fun(X508,fun(X509,bool)),happ(s(fun(X510,fun(X508,fun(X509,bool))),X511),s(X510,happ(s(fun(X512,X510),X513),s(X512,X514))))),s(X508,X515))),s(X509,X516))))) )),
  inference(superposition,[],[f314,f344])).
fof(f6403,plain,(
  ( ! [X499,X505,X507,X500,X502,X504,X506,X501,X503] : (~p(s(bool,happ(s(fun(X500,bool),happ(s(fun(X499,fun(X500,bool)),happ(s(fun(X503,fun(X499,fun(X500,bool))),'_dst_x0x1_2'(s(fun(X501,fun(X499,fun(X500,bool))),X502),s(fun(X503,X501),X504))),s(X503,X505))),s(X499,X506))),s(X500,X507)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X499,X500),bool),'const.pair.UNCURRY_1'(s(fun(X499,fun(X500,bool)),happ(s(fun(X501,fun(X499,fun(X500,bool))),X502),s(X501,happ(s(fun(X503,X501),X504),s(X503,X505))))))))))) )),
  inference(superposition,[],[f319,f344])).
fof(f6402,plain,(
  ( ! [X492,X494,X490,X496,X498,X493,X495,X491,X497] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X490,X491),bool),happ(s(fun(X494,fun('type.pair.prod'(X490,X491),bool)),'_dst_x0x1_2'(s(fun(X492,fun('type.pair.prod'(X490,X491),bool)),X493),s(fun(X494,X492),X495))),s(X494,X496))),s(X490,X497),s(X491,X498)))) | p(s(bool,happ(s(fun('type.pair.prod'(X490,X491),bool),happ(s(fun(X492,fun('type.pair.prod'(X490,X491),bool)),X493),s(X492,happ(s(fun(X494,X492),X495),s(X494,X496))))),s('type.pair.prod'(X490,X491),sK22(happ(s(fun(X492,fun('type.pair.prod'(X490,X491),bool)),X493),s(X492,happ(s(fun(X494,X492),X495),s(X494,X496)))),X491,X490)))))) )),
  inference(superposition,[],[f375,f344])).
fof(f6401,plain,(
  ( ! [X484,X486,X482,X488,X485,X487,X481,X483,X489] : (~p(s(bool,happ(s(fun('type.pair.prod'(X481,X482),bool),happ(s(fun(X485,fun('type.pair.prod'(X481,X482),bool)),'_dst_x0x1_2'(s(fun(X483,fun('type.pair.prod'(X481,X482),bool)),X484),s(fun(X485,X483),X486))),s(X485,X487))),s('type.pair.prod'(X481,X482),sK12(happ(s(fun(X483,fun('type.pair.prod'(X481,X482),bool)),X484),s(X483,happ(s(fun(X485,X483),X486),s(X485,X487)))),X482,X481))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X481,X482),bool),happ(s(fun(X483,fun('type.pair.prod'(X481,X482),bool)),X484),s(X483,happ(s(fun(X485,X483),X486),s(X485,X487))))),s(X481,X488),s(X482,X489))))) )),
  inference(superposition,[],[f376,f344])).
fof(f6400,plain,(
  ( ! [X476,X478,X474,X480,X475,X477,X479,X473] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X473,X474),bool),happ(s(fun(X477,fun('type.pair.prod'(X473,X474),bool)),'_dst_x0x1_2'(s(fun(X475,fun('type.pair.prod'(X473,X474),bool)),X476),s(fun(X477,X475),X478))),s(X477,X479))),s(X473,sK10(happ(s(fun(X475,fun('type.pair.prod'(X473,X474),bool)),X476),s(X475,happ(s(fun(X477,X475),X478),s(X477,X479)))),X474,X473)),s(X474,sK11(happ(s(fun(X475,fun('type.pair.prod'(X473,X474),bool)),X476),s(X475,happ(s(fun(X477,X475),X478),s(X477,X479)))),X474,X473))))) | p(s(bool,happ(s(fun('type.pair.prod'(X473,X474),bool),happ(s(fun(X475,fun('type.pair.prod'(X473,X474),bool)),X476),s(X475,happ(s(fun(X477,X475),X478),s(X477,X479))))),s('type.pair.prod'(X473,X474),X480))))) )),
  inference(superposition,[],[f377,f344])).
fof(f6399,plain,(
  ( ! [X466,X472,X469,X471,X465,X467,X468,X470] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X465,X466),bool),happ(s(fun(X469,fun('type.pair.prod'(X465,X466),bool)),'_dst_x0x1_2'(s(fun(X467,fun('type.pair.prod'(X465,X466),bool)),X468),s(fun(X469,X467),X470))),s(X469,X471))),s(X465,sK13(happ(s(fun(X467,fun('type.pair.prod'(X465,X466),bool)),X468),s(X467,happ(s(fun(X469,X467),X470),s(X469,X471)))),X466,X465)),s(X466,sK14(happ(s(fun(X467,fun('type.pair.prod'(X465,X466),bool)),X468),s(X467,happ(s(fun(X469,X467),X470),s(X469,X471)))),X466,X465))))) | p(s(bool,happ(s(fun('type.pair.prod'(X465,X466),bool),happ(s(fun(X467,fun('type.pair.prod'(X465,X466),bool)),X468),s(X467,happ(s(fun(X469,X467),X470),s(X469,X471))))),s('type.pair.prod'(X465,X466),X472))))) )),
  inference(superposition,[],[f378,f344])).
fof(f6398,plain,(
  ( ! [X461,X463,X457,X459,X460,X462,X458,X464] : (~p(s(bool,happ(s(fun('type.pair.prod'(X457,X458),bool),happ(s(fun(X461,fun('type.pair.prod'(X457,X458),bool)),'_dst_x0x1_2'(s(fun(X459,fun('type.pair.prod'(X457,X458),bool)),X460),s(fun(X461,X459),X462))),s(X461,X463))),s('type.pair.prod'(X457,X458),X464)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X457,X458),bool),happ(s(fun(X459,fun('type.pair.prod'(X457,X458),bool)),X460),s(X459,happ(s(fun(X461,X459),X462),s(X461,X463))))),s(X457,sK20(happ(s(fun(X459,fun('type.pair.prod'(X457,X458),bool)),X460),s(X459,happ(s(fun(X461,X459),X462),s(X461,X463)))),X458,X457)),s(X458,sK21(happ(s(fun(X459,fun('type.pair.prod'(X457,X458),bool)),X460),s(X459,happ(s(fun(X461,X459),X462),s(X461,X463)))),X458,X457)))))) )),
  inference(superposition,[],[f379,f344])).
fof(f6397,plain,(
  ( ! [X453,X455,X451,X452,X454,X450,X456] : (~p(s(bool,happ(s(fun(X450,bool),happ(s(fun(X453,fun(X450,bool)),'_dst_x0x1_2'(s(fun(X451,fun(X450,bool)),X452),s(fun(X453,X451),X454))),s(X453,X455))),s(X450,X456)))) | p(s(bool,happ(s(fun(fun(X450,bool),bool),'?_1'),s(fun(X450,bool),happ(s(fun(X451,fun(X450,bool)),X452),s(X451,happ(s(fun(X453,X451),X454),s(X453,X455))))))))) )),
  inference(superposition,[],[f303,f344])).
fof(f6396,plain,(
  ( ! [X445,X447,X449,X444,X446,X448] : (~p(s(bool,happ(s(fun(fun(X444,bool),bool),'?_1'),s(fun(X444,bool),happ(s(fun(X447,fun(X444,bool)),'_dst_x0x1_2'(s(fun(X445,fun(X444,bool)),X446),s(fun(X447,X445),X448))),s(X447,X449)))))) | p(s(bool,happ(s(fun(X444,bool),happ(s(fun(X445,fun(X444,bool)),X446),s(X445,happ(s(fun(X447,X445),X448),s(X447,X449))))),s(X444,sK9(happ(s(fun(X445,fun(X444,bool)),X446),s(X445,happ(s(fun(X447,X445),X448),s(X447,X449)))),X444)))))) )),
  inference(superposition,[],[f302,f344])).
fof(f6395,plain,(
  ( ! [X441,X443,X438,X440,X442,X439] : (~p(s(bool,happ(s(fun(X438,bool),happ(s(fun(X441,fun(X438,bool)),'_dst_x0x1_2'(s(fun(X439,fun(X438,bool)),X440),s(fun(X441,X439),X442))),s(X441,X443))),s(X438,sK8(happ(s(fun(X439,fun(X438,bool)),X440),s(X439,happ(s(fun(X441,X439),X442),s(X441,X443)))),X438))))) | p(s(bool,happ(s(fun(fun(X438,bool),bool),'!_1'),s(fun(X438,bool),happ(s(fun(X439,fun(X438,bool)),X440),s(X439,happ(s(fun(X441,X439),X442),s(X441,X443))))))))) )),
  inference(superposition,[],[f301,f344])).
fof(f6394,plain,(
  ( ! [X436,X432,X434,X431,X437,X433,X435] : (~p(s(bool,happ(s(fun(fun(X431,bool),bool),'!_1'),s(fun(X431,bool),happ(s(fun(X434,fun(X431,bool)),'_dst_x0x1_2'(s(fun(X432,fun(X431,bool)),X433),s(fun(X434,X432),X435))),s(X434,X436)))))) | p(s(bool,happ(s(fun(X431,bool),happ(s(fun(X432,fun(X431,bool)),X433),s(X432,happ(s(fun(X434,X432),X435),s(X434,X436))))),s(X431,X437))))) )),
  inference(superposition,[],[f300,f344])).
fof(f6393,plain,(
  ( ! [X428,X430,X424,X426,X429,X425,X427] : (p(s(bool,'|dollar|exists_1'(s(fun(X424,bool),happ(s(fun(X427,fun(X424,bool)),'_dst_x0x1_2'(s(fun(X425,fun(X424,bool)),X426),s(fun(X427,X425),X428))),s(X427,X429)))))) | s(X424,X430) = s(X424,sK7(X430,happ(s(fun(X425,fun(X424,bool)),X426),s(X425,happ(s(fun(X427,X425),X428),s(X427,X429)))),X424))) )),
  inference(superposition,[],[f292,f344])).
fof(f6392,plain,(
  ( ! [X420,X422,X418,X421,X423,X417,X419] : (~p(s(bool,happ(s(fun(X417,bool),happ(s(fun(X420,fun(X417,bool)),'_dst_x0x1_2'(s(fun(X418,fun(X417,bool)),X419),s(fun(X420,X418),X421))),s(X420,X422))),s(X417,X423)))) | p(s(bool,'|dollar|exists_1'(s(fun(X417,bool),happ(s(fun(X418,fun(X417,bool)),X419),s(X418,happ(s(fun(X420,X418),X421),s(X420,X422))))))))) )),
  inference(superposition,[],[f289,f344])).
fof(f6391,plain,(
  ( ! [X412,X414,X416,X413,X415,X411] : (p(s(bool,happ(s(fun(X411,bool),happ(s(fun(X414,fun(X411,bool)),'_dst_x0x1_2'(s(fun(X412,fun(X411,bool)),X413),s(fun(X414,X412),X415))),s(X414,X416))),s(X411,sK6(happ(s(fun(X412,fun(X411,bool)),X413),s(X412,happ(s(fun(X414,X412),X415),s(X414,X416)))),X411))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X411,bool),happ(s(fun(X412,fun(X411,bool)),X413),s(X412,happ(s(fun(X414,X412),X415),s(X414,X416))))))))) )),
  inference(superposition,[],[f288,f344])).
fof(f6390,plain,(
  ( ! [X408,X410,X405,X407,X409,X406] : (~p(s(bool,'|dollar|exists_1'(s(fun(X405,bool),happ(s(fun(X408,fun(X405,bool)),'_dst_x0x1_2'(s(fun(X406,fun(X405,bool)),X407),s(fun(X408,X406),X409))),s(X408,X410)))))) | p(s(bool,happ(s(fun(X405,bool),happ(s(fun(X406,fun(X405,bool)),X407),s(X406,happ(s(fun(X408,X406),X409),s(X408,X410))))),s(X405,sK5(happ(s(fun(X406,fun(X405,bool)),X407),s(X406,happ(s(fun(X408,X406),X409),s(X408,X410)))),X405))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X405,bool),happ(s(fun(X406,fun(X405,bool)),X407),s(X406,happ(s(fun(X408,X406),X409),s(X408,X410))))))))) )),
  inference(superposition,[],[f286,f344])).
fof(f6389,plain,(
  ( ! [X399,X401,X403,X404,X400,X402] : (~p(s(bool,'|dollar|exists_1'(s(fun(X399,bool),happ(s(fun(X402,fun(X399,bool)),'_dst_x0x1_2'(s(fun(X400,fun(X399,bool)),X401),s(fun(X402,X400),X403))),s(X402,X404)))))) | p(s(bool,happ(s(fun(X399,bool),happ(s(fun(X400,fun(X399,bool)),X401),s(X400,happ(s(fun(X402,X400),X403),s(X402,X404))))),s(X399,sK4(happ(s(fun(X400,fun(X399,bool)),X401),s(X400,happ(s(fun(X402,X400),X403),s(X402,X404)))),X399))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X399,bool),happ(s(fun(X400,fun(X399,bool)),X401),s(X400,happ(s(fun(X402,X400),X403),s(X402,X404))))))))) )),
  inference(superposition,[],[f285,f344])).
fof(f6388,plain,(
  ( ! [X391,X397,X393,X395,X396,X398,X392,X394] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X391,bool),happ(s(fun(X394,fun(X391,bool)),'_dst_x0x1_2'(s(fun(X392,fun(X391,bool)),X393),s(fun(X394,X392),X395))),s(X394,X396)))))) | ~p(s(bool,happ(s(fun(X391,bool),happ(s(fun(X392,fun(X391,bool)),X393),s(X392,happ(s(fun(X394,X392),X395),s(X394,X396))))),s(X391,X397)))) | ~p(s(bool,happ(s(fun(X391,bool),happ(s(fun(X392,fun(X391,bool)),X393),s(X392,happ(s(fun(X394,X392),X395),s(X394,X396))))),s(X391,X398)))) | s(X391,X397) = s(X391,X398)) )),
  inference(superposition,[],[f284,f344])).
fof(f6387,plain,(
  ( ! [X389,X385,X387,X388,X390,X386] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X385,bool),happ(s(fun(X388,fun(X385,bool)),'_dst_x0x1_2'(s(fun(X386,fun(X385,bool)),X387),s(fun(X388,X386),X389))),s(X388,X390)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X385,bool),happ(s(fun(X386,fun(X385,bool)),X387),s(X386,happ(s(fun(X388,X386),X389),s(X388,X390))))))))) )),
  inference(superposition,[],[f283,f344])).
fof(f6385,plain,(
  ( ! [X350,X352,X354,X349,X351,X353,X348] : (~p(s(bool,happ(s(fun(X348,bool),happ(s(fun(X351,fun(X348,bool)),'_dst_x0x1_2'(s(fun(X349,fun(X348,bool)),X350),s(fun(X351,X349),X352))),s(X351,X353))),s(X348,X354)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X349,X348),bool),'const.pair.UNCURRY_1'(s(fun(X349,fun(X348,bool)),X350))))))) )),
  inference(superposition,[],[f319,f344])).
fof(f6384,plain,(
  ( ! [X344,X346,X343,X345,X347,X342] : (~p(s(bool,happ(s(fun(X342,bool),happ(s(fun(X344,fun(X342,bool)),'_dst_x0x1_2'(s(fun(X343,fun(X342,bool)),sK32(X342,X343)),s(fun(X344,X343),X345))),s(X344,X346))),s(X342,X347)))) | s(X343,sK33(X342,X343)) = s(X343,happ(s(fun(X344,X343),X345),s(X344,X346)))) )),
  inference(superposition,[],[f327,f344])).
fof(f6383,plain,(
  ( ! [X341,X337,X339,X340,X336,X338] : (~p(s(bool,happ(s(fun(X336,bool),happ(s(fun(X338,fun(X336,bool)),'_dst_x0x1_2'(s(fun(X337,fun(X336,bool)),sK32(X336,X337)),s(fun(X338,X337),X339))),s(X338,X340))),s(X336,X341)))) | s(X336,X341) = s(X336,sK34(X336,X337))) )),
  inference(superposition,[],[f328,f344])).
fof(f6382,plain,(
  ( ! [X325,X327,X323,X324,X326,X322] : (p(s(bool,happ(s(fun(X324,bool),'_dst_x0x1_2'(s(fun(X322,bool),X323),s(fun(X324,X322),X325))),s(X324,X326)))) | s(bool,X327) = s(bool,happ(s(fun(X322,bool),X323),s(X322,happ(s(fun(X324,X322),X325),s(X324,X326))))) | s(bool,'T_0') = s(bool,X327)) )),
  inference(superposition,[],[f468,f344])).
fof(f6381,plain,(
  ( ! [X321,X316,X318,X320,X317,X319] : (~p(s(bool,happ(s(fun(X318,bool),'_dst_x0x1_2'(s(fun(X316,bool),X317),s(fun(X318,X316),X319))),s(X318,X320)))) | ~p(s(bool,X321)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun(X316,bool),X317),s(X316,happ(s(fun(X318,X316),X319),s(X318,X320))))))),s(bool,X321))))) )),
  inference(superposition,[],[f306,f344])).
fof(f6380,plain,(
  ( ! [X313,X315,X310,X312,X314,X311] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X315))),s(bool,happ(s(fun(X312,bool),'_dst_x0x1_2'(s(fun(X310,bool),X311),s(fun(X312,X310),X313))),s(X312,X314)))))) | p(s(bool,happ(s(fun(X310,bool),X311),s(X310,happ(s(fun(X312,X310),X313),s(X312,X314))))))) )),
  inference(superposition,[],[f305,f344])).
fof(f6377,plain,(
  ( ! [X292,X294,X296,X293,X295,X297] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun(X294,bool),'_dst_x0x1_2'(s(fun(X292,bool),X293),s(fun(X294,X292),X295))),s(X294,X296))))),s(bool,X297)))) | p(s(bool,happ(s(fun(X292,bool),X293),s(X292,happ(s(fun(X294,X292),X295),s(X294,X296))))))) )),
  inference(superposition,[],[f304,f344])).
fof(f6376,plain,(
  ( ! [X288,X290,X287,X289,X291,X286] : (~p(s(bool,happ(s(fun(X288,bool),'_dst_x0x1_2'(s(fun(X286,bool),X287),s(fun(X288,X286),X289))),s(X288,X290)))) | ~p(s(bool,X291)) | s(bool,X291) = s(bool,happ(s(fun(X286,bool),X287),s(X286,happ(s(fun(X288,X286),X289),s(X288,X290)))))) )),
  inference(superposition,[],[f299,f344])).
fof(f6374,plain,(
  ( ! [X277,X279,X275,X276,X278,X274] : (p(s(bool,happ(s(fun(X276,bool),'_dst_x0x1_2'(s(fun(X274,bool),X275),s(fun(X276,X274),X277))),s(X276,X278)))) | p(s(bool,X279)) | s(bool,X279) = s(bool,happ(s(fun(X274,bool),X275),s(X274,happ(s(fun(X276,X274),X277),s(X276,X278)))))) )),
  inference(superposition,[],[f296,f344])).
fof(f6373,plain,(
  ( ! [X269,X271,X273,X270,X272] : (p(s(bool,happ(s(fun(X271,bool),'_dst_x0x1_2'(s(fun(X269,bool),X270),s(fun(X271,X269),X272))),s(X271,X273)))) | s(bool,'F_0') = s(bool,happ(s(fun(X269,bool),X270),s(X269,happ(s(fun(X271,X269),X272),s(X271,X273)))))) )),
  inference(superposition,[],[f270,f344])).
fof(f6372,plain,(
  ( ! [X265,X267,X268,X264,X266] : (~p(s(bool,happ(s(fun(X266,bool),'_dst_x0x1_2'(s(fun(X264,bool),X265),s(fun(X266,X264),X267))),s(X266,X268)))) | s(bool,'F_0') != s(bool,happ(s(fun(X264,bool),X265),s(X264,happ(s(fun(X266,X264),X267),s(X266,X268)))))) )),
  inference(superposition,[],[f269,f344])).
fof(f6371,plain,(
  ( ! [X261,X263,X259,X260,X262] : (~p(s(bool,happ(s(fun(X261,bool),'_dst_x0x1_2'(s(fun(X259,bool),X260),s(fun(X261,X259),X262))),s(X261,X263)))) | s(bool,'T_0') = s(bool,happ(s(fun(X259,bool),X260),s(X259,happ(s(fun(X261,X259),X262),s(X261,X263)))))) )),
  inference(superposition,[],[f266,f344])).
fof(f6370,plain,(
  ( ! [X255,X257,X254,X256,X258] : (s(bool,'T_0') != s(bool,happ(s(fun(X256,bool),'_dst_x0x1_2'(s(fun(X254,bool),X255),s(fun(X256,X254),X257))),s(X256,X258))) | p(s(bool,happ(s(fun(X254,bool),X255),s(X254,happ(s(fun(X256,X254),X257),s(X256,X258))))))) )),
  inference(superposition,[],[f265,f344])).
fof(f6368,plain,(
  ( ! [X239,X237,X235,X240,X238,X236] : (~p(s(bool,happ(s(fun(X238,bool),'_dst_x0x1_2'(s(fun('type.pair.prod'(X235,X236),bool),X237),s(fun(X238,'type.pair.prod'(X235,X236)),X239))),s(X238,X240)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X235,X236),bool),X237),s(X235,sK20(X237,X236,X235)),s(X236,sK21(X237,X236,X235)))))) )),
  inference(superposition,[],[f379,f344])).
fof(f6367,plain,(
  ( ! [X231,X233,X234,X232] : (~p(s(bool,happ(s(fun(X232,bool),'_dst_x0x1_2'(s(fun(fun(X231,bool),bool),'?_1'),s(fun(X232,fun(X231,bool)),X233))),s(X232,X234)))) | p(s(bool,happ(s(fun(X231,bool),happ(s(fun(X232,fun(X231,bool)),X233),s(X232,X234))),s(X231,sK9(happ(s(fun(X232,fun(X231,bool)),X233),s(X232,X234)),X231)))))) )),
  inference(superposition,[],[f302,f344])).
fof(f6366,plain,(
  ( ! [X229,X227,X230,X228,X226] : (~p(s(bool,happ(s(fun(X227,bool),'_dst_x0x1_2'(s(fun(fun(X226,bool),bool),'!_1'),s(fun(X227,fun(X226,bool)),X228))),s(X227,X229)))) | p(s(bool,happ(s(fun(X226,bool),happ(s(fun(X227,fun(X226,bool)),X228),s(X227,X229))),s(X226,X230))))) )),
  inference(superposition,[],[f300,f344])).
fof(f6365,plain,(
  ( ! [X212,X216,X215,X213,X211,X217,X214] : (~p(s(bool,happ(s(fun(X215,bool),'_dst_x0x1_2'(s(fun(X211,bool),happ(s(fun(X212,fun(X211,bool)),X213),s(X212,X214))),s(fun(X215,X211),X216))),s(X215,X217)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X212,X211),bool),'const.pair.UNCURRY_1'(s(fun(X212,fun(X211,bool)),X213))))))) )),
  inference(superposition,[],[f319,f344])).
fof(f6364,plain,(
  ( ! [X210,X208,X206,X209,X207,X205] : (~p(s(bool,happ(s(fun(X208,bool),'_dst_x0x1_2'(s(fun(X205,bool),happ(s(fun(X206,fun(X205,bool)),sK32(X205,X206)),s(X206,X207))),s(fun(X208,X205),X209))),s(X208,X210)))) | s(X206,X207) = s(X206,sK33(X205,X206))) )),
  inference(superposition,[],[f327,f344])).
fof(f6363,plain,(
  ( ! [X204,X202,X200,X199,X203,X201] : (~p(s(bool,happ(s(fun(X202,bool),'_dst_x0x1_2'(s(fun(X199,bool),happ(s(fun(X200,fun(X199,bool)),sK32(X199,X200)),s(X200,X201))),s(fun(X202,X199),X203))),s(X202,X204)))) | s(X199,sK34(X199,X200)) = s(X199,happ(s(fun(X202,X199),X203),s(X202,X204)))) )),
  inference(superposition,[],[f328,f344])).
fof(f6362,plain,(
  ( ! [X198,X196,X194,X197,X195] : (~p(s(bool,happ(s(fun(X196,bool),'_dst_x0x1_2'(s(fun(X194,bool),happ(s(fun(X194,fun(X194,bool)),'=_0'),s(X194,X195))),s(fun(X196,X194),X197))),s(X196,X198)))) | s(X194,X195) = s(X194,happ(s(fun(X196,X194),X197),s(X196,X198)))) )),
  inference(superposition,[],[f330,f344])).
fof(f6361,plain,(
  ( ! [X191,X189,X192,X190,X193] : (p(s(bool,happ(s(fun(X191,bool),'_dst_x0x1_2'(s(fun(X189,bool),happ(s(fun(X189,fun(X189,bool)),'=_0'),s(X189,X190))),s(fun(X191,X189),X192))),s(X191,X193)))) | s(X189,X190) != s(X189,happ(s(fun(X191,X189),X192),s(X191,X193)))) )),
  inference(superposition,[],[f331,f344])).
fof(f6360,plain,(
  ( ! [X187,X185,X188,X186,X184] : (~p(s(bool,happ(s(fun(X186,bool),'_dst_x0x1_2'(s(fun(X184,bool),X185),s(fun(X186,X184),X187))),s(X186,X188)))) | p(s(bool,'|dollar|exists_1'(s(fun(X184,bool),X185))))) )),
  inference(superposition,[],[f289,f344])).
fof(f6359,plain,(
  ( ! [X182,X180,X183,X181,X179] : (~p(s(bool,happ(s(fun(X181,bool),'_dst_x0x1_2'(s(fun(X179,bool),X180),s(fun(X181,X179),X182))),s(X181,X183)))) | p(s(bool,happ(s(fun(fun(X179,bool),bool),'?_1'),s(fun(X179,bool),X180))))) )),
  inference(superposition,[],[f303,f344])).
fof(f6354,plain,(
  ( ! [X134,X140,X138,X136,X135,X139,X137] : (p(s(bool,happ(s(fun(X134,bool),happ(s(fun(X134,fun(X134,bool)),'=_0'),s(X134,X140))),s(X134,happ(s(fun(X137,X134),'_dst_x0x1_2'(s(fun(X135,X134),X136),s(fun(X137,X135),X138))),s(X137,X139)))))) | s(X134,X140) != s(X134,happ(s(fun(X135,X134),X136),s(X135,happ(s(fun(X137,X135),X138),s(X137,X139)))))) )),
  inference(superposition,[],[f331,f344])).
fof(f6353,plain,(
  ( ! [X127,X132,X130,X128,X133,X131,X129] : (p(s(bool,happ(s(fun(X127,bool),happ(s(fun(X127,fun(X127,bool)),'=_0'),s(X127,happ(s(fun(X130,X127),'_dst_x0x1_2'(s(fun(X128,X127),X129),s(fun(X130,X128),X131))),s(X130,X132))))),s(X127,X133)))) | s(X127,X133) != s(X127,happ(s(fun(X128,X127),X129),s(X128,happ(s(fun(X130,X128),X131),s(X130,X132)))))) )),
  inference(superposition,[],[f331,f344])).
fof(f6352,plain,(
  ( ! [X125,X123,X121,X126,X124,X122,X120] : (~p(s(bool,happ(s(fun(X120,bool),happ(s(fun(X120,fun(X120,bool)),'=_0'),s(X120,X126))),s(X120,happ(s(fun(X123,X120),'_dst_x0x1_2'(s(fun(X121,X120),X122),s(fun(X123,X121),X124))),s(X123,X125)))))) | s(X120,X126) = s(X120,happ(s(fun(X121,X120),X122),s(X121,happ(s(fun(X123,X121),X124),s(X123,X125)))))) )),
  inference(superposition,[],[f330,f344])).
fof(f6351,plain,(
  ( ! [X118,X116,X114,X119,X117,X115,X113] : (~p(s(bool,happ(s(fun(X113,bool),happ(s(fun(X113,fun(X113,bool)),'=_0'),s(X113,happ(s(fun(X116,X113),'_dst_x0x1_2'(s(fun(X114,X113),X115),s(fun(X116,X114),X117))),s(X116,X118))))),s(X113,X119)))) | s(X113,X119) = s(X113,happ(s(fun(X114,X113),X115),s(X114,happ(s(fun(X116,X114),X117),s(X116,X118)))))) )),
  inference(superposition,[],[f330,f344])).
fof(f6350,plain,(
  ( ! [X111,X109,X107,X105,X112,X110,X108,X106] : (~p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X111,fun(X105,bool)),sK32(X105,X111)),s(X111,X112))),s(X105,happ(s(fun(X108,X105),'_dst_x0x1_2'(s(fun(X106,X105),X107),s(fun(X108,X106),X109))),s(X108,X110)))))) | s(X105,sK34(X105,X111)) = s(X105,happ(s(fun(X106,X105),X107),s(X106,happ(s(fun(X108,X106),X109),s(X108,X110)))))) )),
  inference(superposition,[],[f328,f344])).
fof(f6347,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,happ(s(fun(X87,bool),happ(s(fun(X81,fun(X87,bool)),sK32(X87,X81)),s(X81,happ(s(fun(X84,X81),'_dst_x0x1_2'(s(fun(X82,X81),X83),s(fun(X84,X82),X85))),s(X84,X86))))),s(X87,X88)))) | s(X81,sK33(X87,X81)) = s(X81,happ(s(fun(X82,X81),X83),s(X82,happ(s(fun(X84,X82),X85),s(X84,X86)))))) )),
  inference(superposition,[],[f327,f344])).
fof(f6336,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,happ(s(fun(X27,bool),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),'_dst_x0x1_2'(s(fun(X27,bool),X28),s(fun(X26,X27),X29))),s(X26,X30)))) )),
  inference(superposition,[],[f344,f264])).
fof(f6335,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'|dollar|exists_1'(s(fun(X22,bool),'_dst_x0x1_2'(s(fun(X23,bool),X24),s(fun(X22,X23),X25))))) = s(bool,happ(s(fun(X23,bool),X24),s(X23,happ(s(fun(X22,X23),X25),s(X22,'const.min.@_1'(s(fun(X22,bool),'_dst_x0x1_2'(s(fun(X23,bool),X24),s(fun(X22,X23),X25)))))))))) )),
  inference(superposition,[],[f344,f280])).
fof(f6502,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),X4),s(X3,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X3),X5),s(X1,X6),s(X2,X7))))) = s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),'_dst_x0x1_2'(s(fun(X3,X0),X4),s(fun('type.pair.prod'(X1,X2),X3),X5))),s(X1,X6),s(X2,X7)))) )),
  inference(forward_demodulation,[],[f6333,f343])).
fof(f6333,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),'_dst_x0x1_2'(s(fun(X3,X0),X4),s(fun('type.pair.prod'(X1,X2),X3),X5))),s(X1,X6),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),X4),s(X3,happ(s(fun('type.pair.prod'(X1,X2),X3),X5),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X6),s(X2,X7)))))))) )),
  inference(superposition,[],[f344,f343])).
fof(f6281,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),'_dst_x0x1_2'(s(fun(X29,bool),X30),s(fun(X31,X29),X32))),s(X31,X33))) | s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),X30),s(X29,happ(s(fun(X31,X29),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f344,f264])).
fof(f6280,plain,(
  ( ! [X24,X23,X21,X19,X17,X22,X20,X18] : (s(X17,happ(s(fun(X19,X17),X20),s(X19,happ(s(fun(X18,X19),X21),s(X18,happ(s(fun(X22,X18),X23),s(X22,X24))))))) = s(X17,happ(s(fun(X22,X17),'_dst_x0x1_2'(s(fun(X18,X17),'_dst_x0x1_2'(s(fun(X19,X17),X20),s(fun(X18,X19),X21))),s(fun(X22,X18),X23))),s(X22,X24)))) )),
  inference(superposition,[],[f344,f344])).
fof(f6279,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s(X4,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,X4)),X7),s('type.pair.prod'(X5,X6),happ(s(fun(X8,'type.pair.prod'(X5,X6)),X9),s(X8,X10))))) = s(X4,happ(s(fun(X8,X4),'_dst_x0x1_2'(s(fun('type.pair.prod'(X5,X6),X4),'const.pair.UNCURRY_1'(s(fun(X5,fun(X6,X4)),X7))),s(fun(X8,'type.pair.prod'(X5,X6)),X9))),s(X8,X10)))) )),
  inference(superposition,[],[f344,f335])).
fof(f6278,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X34,X32] : (s(X33,happ(s(fun(X27,X33),X34),s(X27,happ(s(fun(X29,X27),X30),s(X29,happ(s(fun(X28,X29),X31),s(X28,X32))))))) = s(X33,happ(s(fun(X28,X33),'_dst_x0x1_2'(s(fun(X27,X33),X34),s(fun(X28,X27),'_dst_x0x1_2'(s(fun(X29,X27),X30),s(fun(X28,X29),X31))))),s(X28,X32)))) )),
  inference(superposition,[],[f344,f344])).
fof(f6277,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (s(X25,happ(s(fun(X19,X25),X26),s(X19,happ(s(fun(X22,X19),'_dst_x0x1_2'(s(fun(X20,X19),X21),s(fun(X22,X20),X23))),s(X22,X24))))) = s(X25,happ(s(fun(X20,X25),'_dst_x0x1_2'(s(fun(X19,X25),X26),s(fun(X20,X19),X21))),s(X20,happ(s(fun(X22,X20),X23),s(X22,X24)))))) )),
  inference(superposition,[],[f344,f344])).
fof(f6276,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X18,X16] : (s(X17,happ(s(fun(X11,X17),X18),s(X11,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X12,X13),X11),X14),s(X12,X15),s(X13,X16))))) = s(X17,happ(s(fun('type.pair.prod'(X12,X13),X17),'_dst_x0x1_2'(s(fun(X11,X17),X18),s(fun('type.pair.prod'(X12,X13),X11),X14))),s('type.pair.prod'(X12,X13),'const.pair.,_2'(s(X12,X15),s(X13,X16)))))) )),
  inference(superposition,[],[f344,f343])).
fof(f6275,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s(X9,happ(s(fun(X4,X9),X10),s(X4,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,X4)),X7),s('type.pair.prod'(X5,X6),X8))))) = s(X9,happ(s(fun('type.pair.prod'(X5,X6),X9),'_dst_x0x1_2'(s(fun(X4,X9),X10),s(fun('type.pair.prod'(X5,X6),X4),'const.pair.UNCURRY_1'(s(fun(X5,fun(X6,X4)),X7))))),s('type.pair.prod'(X5,X6),X8)))) )),
  inference(superposition,[],[f344,f335])).
fof(f6273,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))),s(X0,happ(s(fun(X3,X0),'_dst_x0x1_2'(s(fun(X1,X0),X2),s(fun(X3,X1),X4))),s(X3,X5))))))) )),
  inference(unit_resulting_resolution,[],[f344,f331])).
fof(f6274,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X3),s(fun(X1,X2),X4))),s(X1,X5))))),s(X0,happ(s(fun(X2,X0),X3),s(X2,happ(s(fun(X1,X2),X4),s(X1,X5))))))))) )),
  inference(unit_resulting_resolution,[],[f344,f331])).
fof(f6205,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X11,X12),bool),happ(s(fun(X13,fun('type.pair.prod'(X11,X12),bool)),X14),s(X13,X15))),s(X11,X16),s(X12,X17)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X13,'type.pair.prod'(X11,X12)),bool),'const.pair.UNCURRY_1'(s(fun(X13,fun('type.pair.prod'(X11,X12),bool)),X14))))))) )),
  inference(superposition,[],[f319,f343])).
fof(f6202,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2))))))) )),
  inference(superposition,[],[f319,f280])).
fof(f6174,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X10,X11),fun(X9,bool)),X12),s(X10,X13),s(X11,X14))),s(X9,X15)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'('type.pair.prod'(X10,X11),X9),bool),'const.pair.UNCURRY_1'(s(fun('type.pair.prod'(X10,X11),fun(X9,bool)),X12))))))) )),
  inference(superposition,[],[f319,f343])).
fof(f6173,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),'const.pair.UNCURRY_2'(s(fun(X4,fun(X5,fun(X3,bool))),X6),s('type.pair.prod'(X4,X5),X7))),s(X3,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'('type.pair.prod'(X4,X5),X3),bool),'const.pair.UNCURRY_1'(s(fun('type.pair.prod'(X4,X5),fun(X3,bool)),'const.pair.UNCURRY_1'(s(fun(X4,fun(X5,fun(X3,bool))),X6))))))))) )),
  inference(superposition,[],[f319,f335])).
fof(f6172,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(fun(X0,bool),X0),bool),'const.pair.UNCURRY_1'(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'))))))) )),
  inference(superposition,[],[f319,f279])).
fof(f6055,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,X3),s(X1,X4)))) | ~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))) )),
  inference(forward_demodulation,[],[f6042,f347])).
fof(f6042,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))),s(X0,X3))),s(X1,X4))))) )),
  inference(superposition,[],[f314,f333])).
fof(f5838,plain,(
  ( ! [X14,X12,X15,X13,X11] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X11,X12),bool),X13),s(X11,X14),s(X12,X15)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X11,X12),bool),X13),s(X11,sK20(X13,X12,X11)),s(X12,sK21(X13,X12,X11)))))) )),
  inference(superposition,[],[f379,f343])).
fof(f5848,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X7,fun(X8,bool)),X9),s(X7,sK20('const.pair.UNCURRY_1'(s(fun(X7,fun(X8,bool)),X9)),X8,X7)))),s(X8,sK21('const.pair.UNCURRY_1'(s(fun(X7,fun(X8,bool)),X9)),X8,X7))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X7,fun(X8,bool)),X9),s('type.pair.prod'(X7,X8),X10))))) )),
  inference(forward_demodulation,[],[f5837,f348])).
fof(f5837,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X7,fun(X8,bool)),X9),s('type.pair.prod'(X7,X8),X10)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X7,X8),bool),'const.pair.UNCURRY_1'(s(fun(X7,fun(X8,bool)),X9))),s(X7,sK20('const.pair.UNCURRY_1'(s(fun(X7,fun(X8,bool)),X9)),X8,X7)),s(X8,sK21('const.pair.UNCURRY_1'(s(fun(X7,fun(X8,bool)),X9)),X8,X7)))))) )),
  inference(superposition,[],[f379,f335])).
fof(f5835,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,sK20(X2,X1,X0)),s(X1,sK21(X2,X1,X0)))))) )),
  inference(superposition,[],[f379,f280])).
fof(f5847,plain,(
  ( ! [X43,X41,X44,X42,X40] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X40,X41),bool),sK28(X43,X42,fun('type.pair.prod'(X40,X41),bool))),s(X40,sK20('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43)),X41,X40)),s(X41,sK21('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43)),X41,X40))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X40,X41),bool),sK28(X43,X42,fun('type.pair.prod'(X40,X41),bool))),s('type.pair.prod'(X40,X41),X44))))) )),
  inference(forward_demodulation,[],[f5826,f701])).
fof(f5826,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,happ(s(fun('type.pair.prod'(X40,X41),bool),sK28(X43,X42,fun('type.pair.prod'(X40,X41),bool))),s('type.pair.prod'(X40,X41),X44)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X40,X41),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43))),s(X40,sK20('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43)),X41,X40)),s(X41,sK21('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43)),X41,X40)))))) )),
  inference(superposition,[],[f379,f701])).
fof(f5846,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X34,sK20('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38)))),X35,X34)),s(X35,sK21('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38)))),X35,X34))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X34,X35),bool),X37),s('type.pair.prod'(X34,X35),X39))))) )),
  inference(forward_demodulation,[],[f5825,f323])).
fof(f5825,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (~p(s(bool,happ(s(fun('type.pair.prod'(X34,X35),bool),X37),s('type.pair.prod'(X34,X35),X39)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X34,X35),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38))))),s(X34,sK20('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38)))),X35,X34)),s(X35,sK21('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38)))),X35,X34)))))) )),
  inference(superposition,[],[f379,f323])).
fof(f5845,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X29,X30),bool),sK29(X32,fun('type.pair.prod'(X29,X30),bool),X31)),s(X29,sK20('const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32)),X30,X29)),s(X30,sK21('const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32)),X30,X29))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),sK29(X32,fun('type.pair.prod'(X29,X30),bool),X31)),s('type.pair.prod'(X29,X30),X33))))) )),
  inference(forward_demodulation,[],[f5824,f702])).
fof(f5824,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),sK29(X32,fun('type.pair.prod'(X29,X30),bool),X31)),s('type.pair.prod'(X29,X30),X33)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X29,X30),bool),'const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32))),s(X29,sK20('const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32)),X30,X29)),s(X30,sK21('const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32)),X30,X29)))))) )),
  inference(superposition,[],[f379,f702])).
fof(f5844,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X23,X24),bool),X27),s(X23,sK20('const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27)))),X24,X23)),s(X24,sK21('const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27)))),X24,X23))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X23,X24),bool),X27),s('type.pair.prod'(X23,X24),X28))))) )),
  inference(forward_demodulation,[],[f5823,f322])).
fof(f5823,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(X23,X24),bool),X27),s('type.pair.prod'(X23,X24),X28)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X23,X24),bool),'const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27))))),s(X23,sK20('const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27)))),X24,X23)),s(X24,sK21('const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27)))),X24,X23)))))) )),
  inference(superposition,[],[f379,f322])).
fof(f5843,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X15,X16),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s(X17,X20),s(X18,X21))),s(X15,sK20(happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))),X16,X15)),s(X16,sK21(happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))),X16,X15))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X15,X16),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s(X17,X20),s(X18,X21))),s('type.pair.prod'(X15,X16),X22))))) )),
  inference(forward_demodulation,[],[f5822,f343])).
fof(f5822,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,happ(s(fun('type.pair.prod'(X15,X16),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s(X17,X20),s(X18,X21))),s('type.pair.prod'(X15,X16),X22)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X15,X16),bool),happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21))))),s(X15,sK20(happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))),X16,X15)),s(X16,sK21(happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))),X16,X15)))))) )),
  inference(superposition,[],[f379,f343])).
fof(f5842,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X8,X9),bool),'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12),s('type.pair.prod'(X10,X11),X13))),s(X8,sK20(happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13)),X9,X8)),s(X9,sK21(happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13)),X9,X8))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X8,X9),bool),'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12),s('type.pair.prod'(X10,X11),X13))),s('type.pair.prod'(X8,X9),X14))))) )),
  inference(forward_demodulation,[],[f5821,f335])).
fof(f5821,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X8,X9),bool),'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12),s('type.pair.prod'(X10,X11),X13))),s('type.pair.prod'(X8,X9),X14)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X8,X9),bool),happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13))),s(X8,sK20(happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13)),X9,X8)),s(X9,sK21(happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13)),X9,X8)))))) )),
  inference(superposition,[],[f379,f335])).
fof(f5841,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X4,X5),bool),X6),s(X4,sK20(happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6)),X5,X4)),s(X5,sK21(happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6)),X5,X4))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),X7))))) )),
  inference(forward_demodulation,[],[f5820,f279])).
fof(f5820,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),X7)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X4,X5),bool),happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6))),s(X4,sK20(happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6)),X5,X4)),s(X5,sK21(happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6)),X5,X4)))))) )),
  inference(superposition,[],[f379,f279])).
fof(f5840,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,sK20('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)),s(X1,sK21('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(forward_demodulation,[],[f5839,f347])).
fof(f5839,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))),s(X0,sK20('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)))),s(X1,sK21('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(forward_demodulation,[],[f5819,f348])).
fof(f5819,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))),s(X0,sK20('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)),s(X1,sK21('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)))))) )),
  inference(superposition,[],[f379,f333])).
fof(f5725,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,bool),bool),X0),s(X1,sK13(X0,bool,X1)),s(bool,'F_0')))) | p(s(bool,happ(s(fun('type.pair.prod'(X1,bool),bool),X0),s('type.pair.prod'(X1,bool),X2)))) | s(bool,'T_0') = s(bool,sK14(X0,bool,X1))) )),
  inference(superposition,[],[f378,f264])).
fof(f5724,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(bool,X1),bool),X0),s(bool,'F_0'),s(X1,sK14(X0,X1,bool))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X1),bool),X0),s('type.pair.prod'(bool,X1),X2)))) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool))) )),
  inference(superposition,[],[f378,f264])).
fof(f5735,plain,(
  ( ! [X43,X41,X44,X42,X40] : (p(s(bool,happ(s(fun('type.pair.prod'(X40,X41),bool),sK28(X43,X42,fun('type.pair.prod'(X40,X41),bool))),s('type.pair.prod'(X40,X41),X44)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X40,X41),bool),sK28(X43,X42,fun('type.pair.prod'(X40,X41),bool))),s(X40,sK13('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43)),X41,X40)),s(X41,sK14('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43)),X41,X40)))))) )),
  inference(forward_demodulation,[],[f5723,f701])).
fof(f5723,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X40,X41),bool),sK28(X43,X42,fun('type.pair.prod'(X40,X41),bool))),s(X40,sK13('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43)),X41,X40)),s(X41,sK14('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43)),X41,X40))))) | p(s(bool,happ(s(fun('type.pair.prod'(X40,X41),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43))),s('type.pair.prod'(X40,X41),X44))))) )),
  inference(superposition,[],[f378,f701])).
fof(f5734,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (p(s(bool,happ(s(fun('type.pair.prod'(X34,X35),bool),X37),s('type.pair.prod'(X34,X35),X39)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X34,sK13('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38)))),X35,X34)),s(X35,sK14('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38)))),X35,X34)))))) )),
  inference(forward_demodulation,[],[f5722,f323])).
fof(f5722,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X34,sK13('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38)))),X35,X34)),s(X35,sK14('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38)))),X35,X34))))) | p(s(bool,happ(s(fun('type.pair.prod'(X34,X35),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38))))),s('type.pair.prod'(X34,X35),X39))))) )),
  inference(superposition,[],[f378,f323])).
fof(f5733,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),sK29(X32,fun('type.pair.prod'(X29,X30),bool),X31)),s('type.pair.prod'(X29,X30),X33)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X29,X30),bool),sK29(X32,fun('type.pair.prod'(X29,X30),bool),X31)),s(X29,sK13('const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32)),X30,X29)),s(X30,sK14('const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32)),X30,X29)))))) )),
  inference(forward_demodulation,[],[f5721,f702])).
fof(f5721,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X29,X30),bool),sK29(X32,fun('type.pair.prod'(X29,X30),bool),X31)),s(X29,sK13('const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32)),X30,X29)),s(X30,sK14('const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32)),X30,X29))))) | p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),'const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32))),s('type.pair.prod'(X29,X30),X33))))) )),
  inference(superposition,[],[f378,f702])).
fof(f5732,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (p(s(bool,happ(s(fun('type.pair.prod'(X23,X24),bool),X27),s('type.pair.prod'(X23,X24),X28)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X23,X24),bool),X27),s(X23,sK13('const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27)))),X24,X23)),s(X24,sK14('const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27)))),X24,X23)))))) )),
  inference(forward_demodulation,[],[f5720,f322])).
fof(f5720,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X23,X24),bool),X27),s(X23,sK13('const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27)))),X24,X23)),s(X24,sK14('const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27)))),X24,X23))))) | p(s(bool,happ(s(fun('type.pair.prod'(X23,X24),bool),'const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27))))),s('type.pair.prod'(X23,X24),X28))))) )),
  inference(superposition,[],[f378,f322])).
fof(f5731,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (p(s(bool,happ(s(fun('type.pair.prod'(X15,X16),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s(X17,X20),s(X18,X21))),s('type.pair.prod'(X15,X16),X22)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X15,X16),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s(X17,X20),s(X18,X21))),s(X15,sK13(happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))),X16,X15)),s(X16,sK14(happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))),X16,X15)))))) )),
  inference(forward_demodulation,[],[f5719,f343])).
fof(f5719,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X15,X16),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s(X17,X20),s(X18,X21))),s(X15,sK13(happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))),X16,X15)),s(X16,sK14(happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))),X16,X15))))) | p(s(bool,happ(s(fun('type.pair.prod'(X15,X16),bool),happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21))))),s('type.pair.prod'(X15,X16),X22))))) )),
  inference(superposition,[],[f378,f343])).
fof(f5730,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun('type.pair.prod'(X8,X9),bool),'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12),s('type.pair.prod'(X10,X11),X13))),s('type.pair.prod'(X8,X9),X14)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X8,X9),bool),'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12),s('type.pair.prod'(X10,X11),X13))),s(X8,sK13(happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13)),X9,X8)),s(X9,sK14(happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13)),X9,X8)))))) )),
  inference(forward_demodulation,[],[f5718,f335])).
fof(f5718,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X8,X9),bool),'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12),s('type.pair.prod'(X10,X11),X13))),s(X8,sK13(happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13)),X9,X8)),s(X9,sK14(happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13)),X9,X8))))) | p(s(bool,happ(s(fun('type.pair.prod'(X8,X9),bool),happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13))),s('type.pair.prod'(X8,X9),X14))))) )),
  inference(superposition,[],[f378,f335])).
fof(f5729,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),X7)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X4,X5),bool),X6),s(X4,sK13(happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6)),X5,X4)),s(X5,sK14(happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6)),X5,X4)))))) )),
  inference(forward_demodulation,[],[f5717,f279])).
fof(f5717,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X4,X5),bool),X6),s(X4,sK13(happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6)),X5,X4)),s(X5,sK14(happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6)),X5,X4))))) | p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6))),s('type.pair.prod'(X4,X5),X7))))) )),
  inference(superposition,[],[f378,f279])).
fof(f5728,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,sK13('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)),s(X1,sK14('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)))))) )),
  inference(forward_demodulation,[],[f5727,f5249])).
fof(f5727,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))),s('type.pair.prod'(X0,X1),X3)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,sK13('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)),s(X1,sK14('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)))))) )),
  inference(forward_demodulation,[],[f5716,f335])).
fof(f5716,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,sK13('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)),s(X1,sK14('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(superposition,[],[f378,f333])).
fof(f5633,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,bool),bool),X0),s(X1,sK10(X0,bool,X1)),s(bool,'F_0')))) | p(s(bool,happ(s(fun('type.pair.prod'(X1,bool),bool),X0),s('type.pair.prod'(X1,bool),X2)))) | s(bool,'T_0') = s(bool,sK11(X0,bool,X1))) )),
  inference(superposition,[],[f377,f264])).
fof(f5632,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(bool,X1),bool),X0),s(bool,'F_0'),s(X1,sK11(X0,X1,bool))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X1),bool),X0),s('type.pair.prod'(bool,X1),X2)))) | s(bool,'T_0') = s(bool,sK10(X0,X1,bool))) )),
  inference(superposition,[],[f377,f264])).
fof(f5643,plain,(
  ( ! [X43,X41,X44,X42,X40] : (p(s(bool,happ(s(fun('type.pair.prod'(X40,X41),bool),sK28(X43,X42,fun('type.pair.prod'(X40,X41),bool))),s('type.pair.prod'(X40,X41),X44)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X40,X41),bool),sK28(X43,X42,fun('type.pair.prod'(X40,X41),bool))),s(X40,sK10('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43)),X41,X40)),s(X41,sK11('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43)),X41,X40)))))) )),
  inference(forward_demodulation,[],[f5631,f701])).
fof(f5631,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X40,X41),bool),sK28(X43,X42,fun('type.pair.prod'(X40,X41),bool))),s(X40,sK10('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43)),X41,X40)),s(X41,sK11('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43)),X41,X40))))) | p(s(bool,happ(s(fun('type.pair.prod'(X40,X41),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X40,X41),bool),X42),X43))),s('type.pair.prod'(X40,X41),X44))))) )),
  inference(superposition,[],[f377,f701])).
fof(f5642,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (p(s(bool,happ(s(fun('type.pair.prod'(X34,X35),bool),X37),s('type.pair.prod'(X34,X35),X39)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X34,sK10('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38)))),X35,X34)),s(X35,sK11('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38)))),X35,X34)))))) )),
  inference(forward_demodulation,[],[f5630,f323])).
fof(f5630,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X34,sK10('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38)))),X35,X34)),s(X35,sK11('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38)))),X35,X34))))) | p(s(bool,happ(s(fun('type.pair.prod'(X34,X35),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X34,X35),bool),X36),'const.pair.,_2'(s(fun('type.pair.prod'(X34,X35),bool),X37),s(X36,X38))))),s('type.pair.prod'(X34,X35),X39))))) )),
  inference(superposition,[],[f377,f323])).
fof(f5641,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),sK29(X32,fun('type.pair.prod'(X29,X30),bool),X31)),s('type.pair.prod'(X29,X30),X33)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X29,X30),bool),sK29(X32,fun('type.pair.prod'(X29,X30),bool),X31)),s(X29,sK10('const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32)),X30,X29)),s(X30,sK11('const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32)),X30,X29)))))) )),
  inference(forward_demodulation,[],[f5629,f702])).
fof(f5629,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X29,X30),bool),sK29(X32,fun('type.pair.prod'(X29,X30),bool),X31)),s(X29,sK10('const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32)),X30,X29)),s(X30,sK11('const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32)),X30,X29))))) | p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),'const.pair.SND_1'(s('type.pair.prod'(X31,fun('type.pair.prod'(X29,X30),bool)),X32))),s('type.pair.prod'(X29,X30),X33))))) )),
  inference(superposition,[],[f377,f702])).
fof(f5640,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (p(s(bool,happ(s(fun('type.pair.prod'(X23,X24),bool),X27),s('type.pair.prod'(X23,X24),X28)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X23,X24),bool),X27),s(X23,sK10('const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27)))),X24,X23)),s(X24,sK11('const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27)))),X24,X23)))))) )),
  inference(forward_demodulation,[],[f5628,f322])).
fof(f5628,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X23,X24),bool),X27),s(X23,sK10('const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27)))),X24,X23)),s(X24,sK11('const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27)))),X24,X23))))) | p(s(bool,happ(s(fun('type.pair.prod'(X23,X24),bool),'const.pair.SND_1'(s('type.pair.prod'(X25,fun('type.pair.prod'(X23,X24),bool)),'const.pair.,_2'(s(X25,X26),s(fun('type.pair.prod'(X23,X24),bool),X27))))),s('type.pair.prod'(X23,X24),X28))))) )),
  inference(superposition,[],[f377,f322])).
fof(f5639,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (p(s(bool,happ(s(fun('type.pair.prod'(X15,X16),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s(X17,X20),s(X18,X21))),s('type.pair.prod'(X15,X16),X22)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X15,X16),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s(X17,X20),s(X18,X21))),s(X15,sK10(happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))),X16,X15)),s(X16,sK11(happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))),X16,X15)))))) )),
  inference(forward_demodulation,[],[f5627,f343])).
fof(f5627,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X15,X16),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s(X17,X20),s(X18,X21))),s(X15,sK10(happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))),X16,X15)),s(X16,sK11(happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))),X16,X15))))) | p(s(bool,happ(s(fun('type.pair.prod'(X15,X16),bool),happ(s(fun('type.pair.prod'(X17,X18),fun('type.pair.prod'(X15,X16),bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21))))),s('type.pair.prod'(X15,X16),X22))))) )),
  inference(superposition,[],[f377,f343])).
fof(f5638,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun('type.pair.prod'(X8,X9),bool),'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12),s('type.pair.prod'(X10,X11),X13))),s('type.pair.prod'(X8,X9),X14)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X8,X9),bool),'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12),s('type.pair.prod'(X10,X11),X13))),s(X8,sK10(happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13)),X9,X8)),s(X9,sK11(happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13)),X9,X8)))))) )),
  inference(forward_demodulation,[],[f5626,f335])).
fof(f5626,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X8,X9),bool),'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12),s('type.pair.prod'(X10,X11),X13))),s(X8,sK10(happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13)),X9,X8)),s(X9,sK11(happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13)),X9,X8))))) | p(s(bool,happ(s(fun('type.pair.prod'(X8,X9),bool),happ(s(fun('type.pair.prod'(X10,X11),fun('type.pair.prod'(X8,X9),bool)),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,fun('type.pair.prod'(X8,X9),bool))),X12))),s('type.pair.prod'(X10,X11),X13))),s('type.pair.prod'(X8,X9),X14))))) )),
  inference(superposition,[],[f377,f335])).
fof(f5637,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),X7)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X4,X5),bool),X6),s(X4,sK10(happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6)),X5,X4)),s(X5,sK11(happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6)),X5,X4)))))) )),
  inference(forward_demodulation,[],[f5625,f279])).
fof(f5625,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X4,X5),bool),X6),s(X4,sK10(happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6)),X5,X4)),s(X5,sK11(happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6)),X5,X4))))) | p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),happ(s(fun(fun('type.pair.prod'(X4,X5),bool),fun('type.pair.prod'(X4,X5),bool)),'_i_0_0'),s(fun('type.pair.prod'(X4,X5),bool),X6))),s('type.pair.prod'(X4,X5),X7))))) )),
  inference(superposition,[],[f377,f279])).
fof(f5636,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,sK10('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)),s(X1,sK11('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)))))) )),
  inference(forward_demodulation,[],[f5635,f5249])).
fof(f5635,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))),s('type.pair.prod'(X0,X1),X3)))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,sK10('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)),s(X1,sK11('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)))))) )),
  inference(forward_demodulation,[],[f5624,f335])).
fof(f5624,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,sK10('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)),s(X1,sK11('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(superposition,[],[f377,f333])).
fof(f5583,plain,(
  ( ! [X408,X405,X407,X403,X404,X406,X402] : (s(X402,sK28('const.pair.CURRY_3'(s(fun('type.pair.prod'(X404,X405),'type.pair.prod'(X402,X403)),X406),s(X404,X407),s(X405,X408)),X403,X402)) = s(X402,sK28(happ(s(fun('type.pair.prod'(X404,X405),'type.pair.prod'(X402,X403)),X406),s('type.pair.prod'(X404,X405),'const.pair.,_2'(s(X404,X407),s(X405,X408)))),X403,X402))) )),
  inference(forward_demodulation,[],[f5535,f701])).
fof(f5535,plain,(
  ( ! [X408,X405,X407,X403,X404,X406,X402] : (s(X402,'const.pair.FST_1'(s('type.pair.prod'(X402,X403),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X404,X405),'type.pair.prod'(X402,X403)),X406),s(X404,X407),s(X405,X408))))) = s(X402,sK28(happ(s(fun('type.pair.prod'(X404,X405),'type.pair.prod'(X402,X403)),X406),s('type.pair.prod'(X404,X405),'const.pair.,_2'(s(X404,X407),s(X405,X408)))),X403,X402))) )),
  inference(superposition,[],[f701,f343])).
fof(f5582,plain,(
  ( ! [X397,X399,X395,X401,X396,X398,X400] : (s(X396,sK29('const.pair.CURRY_3'(s(fun('type.pair.prod'(X397,X398),'type.pair.prod'(X395,X396)),X399),s(X397,X400),s(X398,X401)),X396,X395)) = s(X396,sK29(happ(s(fun('type.pair.prod'(X397,X398),'type.pair.prod'(X395,X396)),X399),s('type.pair.prod'(X397,X398),'const.pair.,_2'(s(X397,X400),s(X398,X401)))),X396,X395))) )),
  inference(forward_demodulation,[],[f5534,f702])).
fof(f5534,plain,(
  ( ! [X397,X399,X395,X401,X396,X398,X400] : (s(X396,'const.pair.SND_1'(s('type.pair.prod'(X395,X396),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X397,X398),'type.pair.prod'(X395,X396)),X399),s(X397,X400),s(X398,X401))))) = s(X396,sK29(happ(s(fun('type.pair.prod'(X397,X398),'type.pair.prod'(X395,X396)),X399),s('type.pair.prod'(X397,X398),'const.pair.,_2'(s(X397,X400),s(X398,X401)))),X396,X395))) )),
  inference(superposition,[],[f702,f343])).
fof(f5581,plain,(
  ( ! [X356,X358,X352,X354,X360,X357,X359,X353,X355] : (p(s(bool,happ(s(fun('type.pair.prod'(X352,X353),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X354,X355),fun('type.pair.prod'(X352,X353),bool)),X356),s(X354,X357),s(X355,X358))),s('type.pair.prod'(X352,X353),sK22(happ(s(fun('type.pair.prod'(X354,X355),fun('type.pair.prod'(X352,X353),bool)),X356),s('type.pair.prod'(X354,X355),'const.pair.,_2'(s(X354,X357),s(X355,X358)))),X353,X352))))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X352,X353),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X354,X355),fun('type.pair.prod'(X352,X353),bool)),X356),s(X354,X357),s(X355,X358))),s(X352,X359),s(X353,X360))))) )),
  inference(forward_demodulation,[],[f5529,f343])).
fof(f5529,plain,(
  ( ! [X356,X358,X352,X354,X360,X357,X359,X353,X355] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X352,X353),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X354,X355),fun('type.pair.prod'(X352,X353),bool)),X356),s(X354,X357),s(X355,X358))),s(X352,X359),s(X353,X360)))) | p(s(bool,happ(s(fun('type.pair.prod'(X352,X353),bool),happ(s(fun('type.pair.prod'(X354,X355),fun('type.pair.prod'(X352,X353),bool)),X356),s('type.pair.prod'(X354,X355),'const.pair.,_2'(s(X354,X357),s(X355,X358))))),s('type.pair.prod'(X352,X353),sK22(happ(s(fun('type.pair.prod'(X354,X355),fun('type.pair.prod'(X352,X353),bool)),X356),s('type.pair.prod'(X354,X355),'const.pair.,_2'(s(X354,X357),s(X355,X358)))),X353,X352)))))) )),
  inference(superposition,[],[f375,f343])).
fof(f5580,plain,(
  ( ! [X350,X344,X346,X343,X349,X351,X345,X347,X348] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X343,X344),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X345,X346),fun('type.pair.prod'(X343,X344),bool)),X347),s(X345,X348),s(X346,X349))),s(X343,X350),s(X344,X351)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X343,X344),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X345,X346),fun('type.pair.prod'(X343,X344),bool)),X347),s(X345,X348),s(X346,X349))),s('type.pair.prod'(X343,X344),sK12(happ(s(fun('type.pair.prod'(X345,X346),fun('type.pair.prod'(X343,X344),bool)),X347),s('type.pair.prod'(X345,X346),'const.pair.,_2'(s(X345,X348),s(X346,X349)))),X344,X343)))))) )),
  inference(forward_demodulation,[],[f5528,f343])).
fof(f5528,plain,(
  ( ! [X350,X344,X346,X343,X349,X351,X345,X347,X348] : (~p(s(bool,happ(s(fun('type.pair.prod'(X343,X344),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X345,X346),fun('type.pair.prod'(X343,X344),bool)),X347),s(X345,X348),s(X346,X349))),s('type.pair.prod'(X343,X344),sK12(happ(s(fun('type.pair.prod'(X345,X346),fun('type.pair.prod'(X343,X344),bool)),X347),s('type.pair.prod'(X345,X346),'const.pair.,_2'(s(X345,X348),s(X346,X349)))),X344,X343))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X343,X344),bool),happ(s(fun('type.pair.prod'(X345,X346),fun('type.pair.prod'(X343,X344),bool)),X347),s('type.pair.prod'(X345,X346),'const.pair.,_2'(s(X345,X348),s(X346,X349))))),s(X343,X350),s(X344,X351))))) )),
  inference(superposition,[],[f376,f343])).
fof(f5578,plain,(
  ( ! [X333,X335,X331,X332,X334,X330] : (p(s(bool,happ(s(fun(X330,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X331,X332),fun(X330,bool)),X333),s(X331,X334),s(X332,X335))),s(X330,sK9(happ(s(fun('type.pair.prod'(X331,X332),fun(X330,bool)),X333),s('type.pair.prod'(X331,X332),'const.pair.,_2'(s(X331,X334),s(X332,X335)))),X330))))) | ~p(s(bool,happ(s(fun(fun(X330,bool),bool),'?_1'),s(fun(X330,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X331,X332),fun(X330,bool)),X333),s(X331,X334),s(X332,X335))))))) )),
  inference(forward_demodulation,[],[f5526,f343])).
fof(f5526,plain,(
  ( ! [X333,X335,X331,X332,X334,X330] : (~p(s(bool,happ(s(fun(fun(X330,bool),bool),'?_1'),s(fun(X330,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X331,X332),fun(X330,bool)),X333),s(X331,X334),s(X332,X335)))))) | p(s(bool,happ(s(fun(X330,bool),happ(s(fun('type.pair.prod'(X331,X332),fun(X330,bool)),X333),s('type.pair.prod'(X331,X332),'const.pair.,_2'(s(X331,X334),s(X332,X335))))),s(X330,sK9(happ(s(fun('type.pair.prod'(X331,X332),fun(X330,bool)),X333),s('type.pair.prod'(X331,X332),'const.pair.,_2'(s(X331,X334),s(X332,X335)))),X330)))))) )),
  inference(superposition,[],[f302,f343])).
fof(f5577,plain,(
  ( ! [X325,X327,X329,X324,X326,X328] : (p(s(bool,happ(s(fun(fun(X324,bool),bool),'!_1'),s(fun(X324,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X325,X326),fun(X324,bool)),X327),s(X325,X328),s(X326,X329)))))) | ~p(s(bool,happ(s(fun(X324,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X325,X326),fun(X324,bool)),X327),s(X325,X328),s(X326,X329))),s(X324,sK8(happ(s(fun('type.pair.prod'(X325,X326),fun(X324,bool)),X327),s('type.pair.prod'(X325,X326),'const.pair.,_2'(s(X325,X328),s(X326,X329)))),X324)))))) )),
  inference(forward_demodulation,[],[f5525,f343])).
fof(f5525,plain,(
  ( ! [X325,X327,X329,X324,X326,X328] : (~p(s(bool,happ(s(fun(X324,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X325,X326),fun(X324,bool)),X327),s(X325,X328),s(X326,X329))),s(X324,sK8(happ(s(fun('type.pair.prod'(X325,X326),fun(X324,bool)),X327),s('type.pair.prod'(X325,X326),'const.pair.,_2'(s(X325,X328),s(X326,X329)))),X324))))) | p(s(bool,happ(s(fun(fun(X324,bool),bool),'!_1'),s(fun(X324,bool),happ(s(fun('type.pair.prod'(X325,X326),fun(X324,bool)),X327),s('type.pair.prod'(X325,X326),'const.pair.,_2'(s(X325,X328),s(X326,X329))))))))) )),
  inference(superposition,[],[f301,f343])).
fof(f5523,plain,(
  ( ! [X313,X315,X310,X316,X312,X314,X311] : (p(s(bool,'|dollar|exists_1'(s(fun(X310,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X311,X312),fun(X310,bool)),X313),s(X311,X314),s(X312,X315)))))) | s(X310,X316) = s(X310,sK7(X316,happ(s(fun('type.pair.prod'(X311,X312),fun(X310,bool)),X313),s('type.pair.prod'(X311,X312),'const.pair.,_2'(s(X311,X314),s(X312,X315)))),X310))) )),
  inference(superposition,[],[f292,f343])).
fof(f5574,plain,(
  ( ! [X300,X302,X298,X301,X297,X299] : (~p(s(bool,'|dollar|exists_1'(s(fun(X297,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X298,X299),fun(X297,bool)),X300),s(X298,X301),s(X299,X302)))))) | p(s(bool,happ(s(fun(X297,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X298,X299),fun(X297,bool)),X300),s(X298,X301),s(X299,X302))),s(X297,sK6(happ(s(fun('type.pair.prod'(X298,X299),fun(X297,bool)),X300),s('type.pair.prod'(X298,X299),'const.pair.,_2'(s(X298,X301),s(X299,X302)))),X297)))))) )),
  inference(forward_demodulation,[],[f5521,f343])).
fof(f5521,plain,(
  ( ! [X300,X302,X298,X301,X297,X299] : (p(s(bool,happ(s(fun(X297,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X298,X299),fun(X297,bool)),X300),s(X298,X301),s(X299,X302))),s(X297,sK6(happ(s(fun('type.pair.prod'(X298,X299),fun(X297,bool)),X300),s('type.pair.prod'(X298,X299),'const.pair.,_2'(s(X298,X301),s(X299,X302)))),X297))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X297,bool),happ(s(fun('type.pair.prod'(X298,X299),fun(X297,bool)),X300),s('type.pair.prod'(X298,X299),'const.pair.,_2'(s(X298,X301),s(X299,X302))))))))) )),
  inference(superposition,[],[f288,f343])).
fof(f5573,plain,(
  ( ! [X292,X294,X296,X293,X295,X291] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X291,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X292,X293),fun(X291,bool)),X294),s(X292,X295),s(X293,X296)))))) | p(s(bool,happ(s(fun(X291,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X292,X293),fun(X291,bool)),X294),s(X292,X295),s(X293,X296))),s(X291,sK5(happ(s(fun('type.pair.prod'(X292,X293),fun(X291,bool)),X294),s('type.pair.prod'(X292,X293),'const.pair.,_2'(s(X292,X295),s(X293,X296)))),X291))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X291,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X292,X293),fun(X291,bool)),X294),s(X292,X295),s(X293,X296))))))) )),
  inference(forward_demodulation,[],[f5572,f343])).
fof(f5572,plain,(
  ( ! [X292,X294,X296,X293,X295,X291] : (p(s(bool,happ(s(fun(X291,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X292,X293),fun(X291,bool)),X294),s(X292,X295),s(X293,X296))),s(X291,sK5(happ(s(fun('type.pair.prod'(X292,X293),fun(X291,bool)),X294),s('type.pair.prod'(X292,X293),'const.pair.,_2'(s(X292,X295),s(X293,X296)))),X291))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X291,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X292,X293),fun(X291,bool)),X294),s(X292,X295),s(X293,X296)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X291,bool),happ(s(fun('type.pair.prod'(X292,X293),fun(X291,bool)),X294),s('type.pair.prod'(X292,X293),'const.pair.,_2'(s(X292,X295),s(X293,X296))))))))) )),
  inference(forward_demodulation,[],[f5520,f343])).
fof(f5520,plain,(
  ( ! [X292,X294,X296,X293,X295,X291] : (~p(s(bool,'|dollar|exists_1'(s(fun(X291,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X292,X293),fun(X291,bool)),X294),s(X292,X295),s(X293,X296)))))) | p(s(bool,happ(s(fun(X291,bool),happ(s(fun('type.pair.prod'(X292,X293),fun(X291,bool)),X294),s('type.pair.prod'(X292,X293),'const.pair.,_2'(s(X292,X295),s(X293,X296))))),s(X291,sK5(happ(s(fun('type.pair.prod'(X292,X293),fun(X291,bool)),X294),s('type.pair.prod'(X292,X293),'const.pair.,_2'(s(X292,X295),s(X293,X296)))),X291))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X291,bool),happ(s(fun('type.pair.prod'(X292,X293),fun(X291,bool)),X294),s('type.pair.prod'(X292,X293),'const.pair.,_2'(s(X292,X295),s(X293,X296))))))))) )),
  inference(superposition,[],[f286,f343])).
fof(f5571,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X285,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X286,X287),fun(X285,bool)),X288),s(X286,X289),s(X287,X290)))))) | p(s(bool,happ(s(fun(X285,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X286,X287),fun(X285,bool)),X288),s(X286,X289),s(X287,X290))),s(X285,sK4(happ(s(fun('type.pair.prod'(X286,X287),fun(X285,bool)),X288),s('type.pair.prod'(X286,X287),'const.pair.,_2'(s(X286,X289),s(X287,X290)))),X285))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X285,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X286,X287),fun(X285,bool)),X288),s(X286,X289),s(X287,X290))))))) )),
  inference(forward_demodulation,[],[f5570,f343])).
fof(f5570,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (p(s(bool,happ(s(fun(X285,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X286,X287),fun(X285,bool)),X288),s(X286,X289),s(X287,X290))),s(X285,sK4(happ(s(fun('type.pair.prod'(X286,X287),fun(X285,bool)),X288),s('type.pair.prod'(X286,X287),'const.pair.,_2'(s(X286,X289),s(X287,X290)))),X285))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X285,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X286,X287),fun(X285,bool)),X288),s(X286,X289),s(X287,X290)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X285,bool),happ(s(fun('type.pair.prod'(X286,X287),fun(X285,bool)),X288),s('type.pair.prod'(X286,X287),'const.pair.,_2'(s(X286,X289),s(X287,X290))))))))) )),
  inference(forward_demodulation,[],[f5519,f343])).
fof(f5519,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (~p(s(bool,'|dollar|exists_1'(s(fun(X285,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X286,X287),fun(X285,bool)),X288),s(X286,X289),s(X287,X290)))))) | p(s(bool,happ(s(fun(X285,bool),happ(s(fun('type.pair.prod'(X286,X287),fun(X285,bool)),X288),s('type.pair.prod'(X286,X287),'const.pair.,_2'(s(X286,X289),s(X287,X290))))),s(X285,sK4(happ(s(fun('type.pair.prod'(X286,X287),fun(X285,bool)),X288),s('type.pair.prod'(X286,X287),'const.pair.,_2'(s(X286,X289),s(X287,X290)))),X285))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X285,bool),happ(s(fun('type.pair.prod'(X286,X287),fun(X285,bool)),X288),s('type.pair.prod'(X286,X287),'const.pair.,_2'(s(X286,X289),s(X287,X290))))))))) )),
  inference(superposition,[],[f285,f343])).
fof(f5516,plain,(
  ( ! [X269,X271,X267,X268,X270,X272] : (~p(s(bool,happ(s(fun(X267,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X268,X269),fun(X267,bool)),sK32(X267,'type.pair.prod'(X268,X269))),s(X268,X270),s(X269,X271))),s(X267,X272)))) | s('type.pair.prod'(X268,X269),sK33(X267,'type.pair.prod'(X268,X269))) = s('type.pair.prod'(X268,X269),'const.pair.,_2'(s(X268,X270),s(X269,X271)))) )),
  inference(superposition,[],[f327,f343])).
fof(f5515,plain,(
  ( ! [X261,X263,X265,X262,X264,X266] : (~p(s(bool,happ(s(fun(X261,bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X262,X263),fun(X261,bool)),sK32(X261,'type.pair.prod'(X262,X263))),s(X262,X264),s(X263,X265))),s(X261,X266)))) | s(X261,X266) = s(X261,sK34(X261,'type.pair.prod'(X262,X263)))) )),
  inference(superposition,[],[f328,f343])).
fof(f5514,plain,(
  ( ! [X257,X259,X260,X256,X258] : (~p(s(bool,happ(s(fun('type.pair.prod'(X256,X257),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X256,X257),fun('type.pair.prod'(X256,X257),bool)),'=_0'),s(X256,X258),s(X257,X259))),s('type.pair.prod'(X256,X257),X260)))) | s('type.pair.prod'(X256,X257),X260) = s('type.pair.prod'(X256,X257),'const.pair.,_2'(s(X256,X258),s(X257,X259)))) )),
  inference(superposition,[],[f330,f343])).
fof(f5513,plain,(
  ( ! [X255,X253,X251,X254,X252] : (p(s(bool,happ(s(fun('type.pair.prod'(X251,X252),bool),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X251,X252),fun('type.pair.prod'(X251,X252),bool)),'=_0'),s(X251,X253),s(X252,X254))),s('type.pair.prod'(X251,X252),X255)))) | s('type.pair.prod'(X251,X252),X255) != s('type.pair.prod'(X251,X252),'const.pair.,_2'(s(X251,X253),s(X252,X254)))) )),
  inference(superposition,[],[f331,f343])).
fof(f5498,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X166,X167),bool),happ(s(fun(X168,fun('type.pair.prod'(X166,X167),bool)),sK32('type.pair.prod'(X166,X167),X168)),s(X168,X169))),s(X166,X170),s(X167,X171)))) | s(X168,X169) = s(X168,sK33('type.pair.prod'(X166,X167),X168))) )),
  inference(superposition,[],[f327,f343])).
fof(f5497,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X160,X161),bool),happ(s(fun(X162,fun('type.pair.prod'(X160,X161),bool)),sK32('type.pair.prod'(X160,X161),X162)),s(X162,X163))),s(X160,X164),s(X161,X165)))) | s('type.pair.prod'(X160,X161),sK34('type.pair.prod'(X160,X161),X162)) = s('type.pair.prod'(X160,X161),'const.pair.,_2'(s(X160,X164),s(X161,X165)))) )),
  inference(superposition,[],[f328,f343])).
fof(f5496,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X155,X156),bool),happ(s(fun('type.pair.prod'(X155,X156),fun('type.pair.prod'(X155,X156),bool)),'=_0'),s('type.pair.prod'(X155,X156),X157))),s(X155,X158),s(X156,X159)))) | s('type.pair.prod'(X155,X156),X157) = s('type.pair.prod'(X155,X156),'const.pair.,_2'(s(X155,X158),s(X156,X159)))) )),
  inference(superposition,[],[f330,f343])).
fof(f5495,plain,(
  ( ! [X154,X152,X151,X153,X150] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X150,X151),bool),happ(s(fun('type.pair.prod'(X150,X151),fun('type.pair.prod'(X150,X151),bool)),'=_0'),s('type.pair.prod'(X150,X151),X152))),s(X150,X153),s(X151,X154)))) | s('type.pair.prod'(X150,X151),X152) != s('type.pair.prod'(X150,X151),'const.pair.,_2'(s(X150,X153),s(X151,X154)))) )),
  inference(superposition,[],[f331,f343])).
fof(f5494,plain,(
  ( ! [X146,X149,X147,X145,X148] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X145,X146),bool),X147),s(X145,X148),s(X146,X149)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X145,X146),bool),X147))))) )),
  inference(superposition,[],[f289,f343])).
fof(f5493,plain,(
  ( ! [X144,X142,X140,X143,X141] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X140,X141),bool),X142),s(X140,X143),s(X141,X144)))) | p(s(bool,happ(s(fun(fun('type.pair.prod'(X140,X141),bool),bool),'?_1'),s(fun('type.pair.prod'(X140,X141),bool),X142))))) )),
  inference(superposition,[],[f303,f343])).
fof(f5473,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'F_0') = s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X16,X17),bool),X18),s(X16,X19),s(X17,X20))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X16,X17),bool),X18),s('type.pair.prod'(X16,X17),'const.pair.,_2'(s(X16,X19),s(X17,X20)))))) )),
  inference(superposition,[],[f343,f264])).
fof(f5472,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.pair.prod'(X12,X13),'const.pair.,_2'(s(X12,X14),s(X13,X15))) = s('type.pair.prod'(X12,X13),'const.pair.CURRY_3'(s(fun('type.pair.prod'(X12,X13),'type.pair.prod'(X12,X13)),'_i_0_0'),s(X12,X14),s(X13,X15)))) )),
  inference(superposition,[],[f343,f279])).
fof(f5471,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))) = s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s(X1,X4),s(X2,X5)))) )),
  inference(superposition,[],[f343,f335])).
fof(f5470,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X3,happ(s(fun('type.pair.prod'(X0,X1),X3),X4),s('type.pair.prod'(X0,X1),X2))) = s(X3,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),X3),X4),s(X0,sK28(X2,X1,X0)),s(X1,sK29(X2,X1,X0))))) )),
  inference(superposition,[],[f343,f321])).
fof(f5363,plain,(
  ( ! [X300,X296,X298,X295,X297,X299] : (s(X295,sK28('const.pair.UNCURRY_2'(s(fun(X297,fun(X298,'type.pair.prod'(X295,X296))),X299),s('type.pair.prod'(X297,X298),X300)),X296,X295)) = s(X295,sK28(happ(s(fun('type.pair.prod'(X297,X298),'type.pair.prod'(X295,X296)),'const.pair.UNCURRY_1'(s(fun(X297,fun(X298,'type.pair.prod'(X295,X296))),X299))),s('type.pair.prod'(X297,X298),X300)),X296,X295))) )),
  inference(forward_demodulation,[],[f5312,f701])).
fof(f5312,plain,(
  ( ! [X300,X296,X298,X295,X297,X299] : (s(X295,'const.pair.FST_1'(s('type.pair.prod'(X295,X296),'const.pair.UNCURRY_2'(s(fun(X297,fun(X298,'type.pair.prod'(X295,X296))),X299),s('type.pair.prod'(X297,X298),X300))))) = s(X295,sK28(happ(s(fun('type.pair.prod'(X297,X298),'type.pair.prod'(X295,X296)),'const.pair.UNCURRY_1'(s(fun(X297,fun(X298,'type.pair.prod'(X295,X296))),X299))),s('type.pair.prod'(X297,X298),X300)),X296,X295))) )),
  inference(superposition,[],[f701,f335])).
fof(f5362,plain,(
  ( ! [X292,X294,X290,X293,X289,X291] : (s(X290,sK29('const.pair.UNCURRY_2'(s(fun(X291,fun(X292,'type.pair.prod'(X289,X290))),X293),s('type.pair.prod'(X291,X292),X294)),X290,X289)) = s(X290,sK29(happ(s(fun('type.pair.prod'(X291,X292),'type.pair.prod'(X289,X290)),'const.pair.UNCURRY_1'(s(fun(X291,fun(X292,'type.pair.prod'(X289,X290))),X293))),s('type.pair.prod'(X291,X292),X294)),X290,X289))) )),
  inference(forward_demodulation,[],[f5311,f702])).
fof(f5311,plain,(
  ( ! [X292,X294,X290,X293,X289,X291] : (s(X290,'const.pair.SND_1'(s('type.pair.prod'(X289,X290),'const.pair.UNCURRY_2'(s(fun(X291,fun(X292,'type.pair.prod'(X289,X290))),X293),s('type.pair.prod'(X291,X292),X294))))) = s(X290,sK29(happ(s(fun('type.pair.prod'(X291,X292),'type.pair.prod'(X289,X290)),'const.pair.UNCURRY_1'(s(fun(X291,fun(X292,'type.pair.prod'(X289,X290))),X293))),s('type.pair.prod'(X291,X292),X294)),X290,X289))) )),
  inference(superposition,[],[f702,f335])).
fof(f5361,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272,X274] : (p(s(bool,happ(s(fun('type.pair.prod'(X267,X268),bool),'const.pair.UNCURRY_2'(s(fun(X269,fun(X270,fun('type.pair.prod'(X267,X268),bool))),X271),s('type.pair.prod'(X269,X270),X272))),s('type.pair.prod'(X267,X268),sK22(happ(s(fun('type.pair.prod'(X269,X270),fun('type.pair.prod'(X267,X268),bool)),'const.pair.UNCURRY_1'(s(fun(X269,fun(X270,fun('type.pair.prod'(X267,X268),bool))),X271))),s('type.pair.prod'(X269,X270),X272)),X268,X267))))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X267,X268),bool),'const.pair.UNCURRY_2'(s(fun(X269,fun(X270,fun('type.pair.prod'(X267,X268),bool))),X271),s('type.pair.prod'(X269,X270),X272))),s(X267,X273),s(X268,X274))))) )),
  inference(forward_demodulation,[],[f5308,f335])).
fof(f5308,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272,X274] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X267,X268),bool),'const.pair.UNCURRY_2'(s(fun(X269,fun(X270,fun('type.pair.prod'(X267,X268),bool))),X271),s('type.pair.prod'(X269,X270),X272))),s(X267,X273),s(X268,X274)))) | p(s(bool,happ(s(fun('type.pair.prod'(X267,X268),bool),happ(s(fun('type.pair.prod'(X269,X270),fun('type.pair.prod'(X267,X268),bool)),'const.pair.UNCURRY_1'(s(fun(X269,fun(X270,fun('type.pair.prod'(X267,X268),bool))),X271))),s('type.pair.prod'(X269,X270),X272))),s('type.pair.prod'(X267,X268),sK22(happ(s(fun('type.pair.prod'(X269,X270),fun('type.pair.prod'(X267,X268),bool)),'const.pair.UNCURRY_1'(s(fun(X269,fun(X270,fun('type.pair.prod'(X267,X268),bool))),X271))),s('type.pair.prod'(X269,X270),X272)),X268,X267)))))) )),
  inference(superposition,[],[f375,f335])).
fof(f5360,plain,(
  ( ! [X261,X263,X259,X265,X260,X262,X264,X266] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X259,X260),bool),'const.pair.UNCURRY_2'(s(fun(X261,fun(X262,fun('type.pair.prod'(X259,X260),bool))),X263),s('type.pair.prod'(X261,X262),X264))),s(X259,X265),s(X260,X266)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X259,X260),bool),'const.pair.UNCURRY_2'(s(fun(X261,fun(X262,fun('type.pair.prod'(X259,X260),bool))),X263),s('type.pair.prod'(X261,X262),X264))),s('type.pair.prod'(X259,X260),sK12(happ(s(fun('type.pair.prod'(X261,X262),fun('type.pair.prod'(X259,X260),bool)),'const.pair.UNCURRY_1'(s(fun(X261,fun(X262,fun('type.pair.prod'(X259,X260),bool))),X263))),s('type.pair.prod'(X261,X262),X264)),X260,X259)))))) )),
  inference(forward_demodulation,[],[f5307,f335])).
fof(f5307,plain,(
  ( ! [X261,X263,X259,X265,X260,X262,X264,X266] : (~p(s(bool,happ(s(fun('type.pair.prod'(X259,X260),bool),'const.pair.UNCURRY_2'(s(fun(X261,fun(X262,fun('type.pair.prod'(X259,X260),bool))),X263),s('type.pair.prod'(X261,X262),X264))),s('type.pair.prod'(X259,X260),sK12(happ(s(fun('type.pair.prod'(X261,X262),fun('type.pair.prod'(X259,X260),bool)),'const.pair.UNCURRY_1'(s(fun(X261,fun(X262,fun('type.pair.prod'(X259,X260),bool))),X263))),s('type.pair.prod'(X261,X262),X264)),X260,X259))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X259,X260),bool),happ(s(fun('type.pair.prod'(X261,X262),fun('type.pair.prod'(X259,X260),bool)),'const.pair.UNCURRY_1'(s(fun(X261,fun(X262,fun('type.pair.prod'(X259,X260),bool))),X263))),s('type.pair.prod'(X261,X262),X264))),s(X259,X265),s(X260,X266))))) )),
  inference(superposition,[],[f376,f335])).
fof(f5358,plain,(
  ( ! [X251,X249,X252,X250,X248] : (p(s(bool,happ(s(fun(X248,bool),'const.pair.UNCURRY_2'(s(fun(X249,fun(X250,fun(X248,bool))),X251),s('type.pair.prod'(X249,X250),X252))),s(X248,sK9(happ(s(fun('type.pair.prod'(X249,X250),fun(X248,bool)),'const.pair.UNCURRY_1'(s(fun(X249,fun(X250,fun(X248,bool))),X251))),s('type.pair.prod'(X249,X250),X252)),X248))))) | ~p(s(bool,happ(s(fun(fun(X248,bool),bool),'?_1'),s(fun(X248,bool),'const.pair.UNCURRY_2'(s(fun(X249,fun(X250,fun(X248,bool))),X251),s('type.pair.prod'(X249,X250),X252))))))) )),
  inference(forward_demodulation,[],[f5305,f335])).
fof(f5305,plain,(
  ( ! [X251,X249,X252,X250,X248] : (~p(s(bool,happ(s(fun(fun(X248,bool),bool),'?_1'),s(fun(X248,bool),'const.pair.UNCURRY_2'(s(fun(X249,fun(X250,fun(X248,bool))),X251),s('type.pair.prod'(X249,X250),X252)))))) | p(s(bool,happ(s(fun(X248,bool),happ(s(fun('type.pair.prod'(X249,X250),fun(X248,bool)),'const.pair.UNCURRY_1'(s(fun(X249,fun(X250,fun(X248,bool))),X251))),s('type.pair.prod'(X249,X250),X252))),s(X248,sK9(happ(s(fun('type.pair.prod'(X249,X250),fun(X248,bool)),'const.pair.UNCURRY_1'(s(fun(X249,fun(X250,fun(X248,bool))),X251))),s('type.pair.prod'(X249,X250),X252)),X248)))))) )),
  inference(superposition,[],[f302,f335])).
fof(f5357,plain,(
  ( ! [X243,X246,X244,X247,X245] : (p(s(bool,happ(s(fun(fun(X243,bool),bool),'!_1'),s(fun(X243,bool),'const.pair.UNCURRY_2'(s(fun(X244,fun(X245,fun(X243,bool))),X246),s('type.pair.prod'(X244,X245),X247)))))) | ~p(s(bool,happ(s(fun(X243,bool),'const.pair.UNCURRY_2'(s(fun(X244,fun(X245,fun(X243,bool))),X246),s('type.pair.prod'(X244,X245),X247))),s(X243,sK8(happ(s(fun('type.pair.prod'(X244,X245),fun(X243,bool)),'const.pair.UNCURRY_1'(s(fun(X244,fun(X245,fun(X243,bool))),X246))),s('type.pair.prod'(X244,X245),X247)),X243)))))) )),
  inference(forward_demodulation,[],[f5304,f335])).
fof(f5304,plain,(
  ( ! [X243,X246,X244,X247,X245] : (~p(s(bool,happ(s(fun(X243,bool),'const.pair.UNCURRY_2'(s(fun(X244,fun(X245,fun(X243,bool))),X246),s('type.pair.prod'(X244,X245),X247))),s(X243,sK8(happ(s(fun('type.pair.prod'(X244,X245),fun(X243,bool)),'const.pair.UNCURRY_1'(s(fun(X244,fun(X245,fun(X243,bool))),X246))),s('type.pair.prod'(X244,X245),X247)),X243))))) | p(s(bool,happ(s(fun(fun(X243,bool),bool),'!_1'),s(fun(X243,bool),happ(s(fun('type.pair.prod'(X244,X245),fun(X243,bool)),'const.pair.UNCURRY_1'(s(fun(X244,fun(X245,fun(X243,bool))),X246))),s('type.pair.prod'(X244,X245),X247))))))) )),
  inference(superposition,[],[f301,f335])).
fof(f5302,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,'|dollar|exists_1'(s(fun(X231,bool),'const.pair.UNCURRY_2'(s(fun(X232,fun(X233,fun(X231,bool))),X234),s('type.pair.prod'(X232,X233),X235)))))) | s(X231,X236) = s(X231,sK7(X236,happ(s(fun('type.pair.prod'(X232,X233),fun(X231,bool)),'const.pair.UNCURRY_1'(s(fun(X232,fun(X233,fun(X231,bool))),X234))),s('type.pair.prod'(X232,X233),X235)),X231))) )),
  inference(superposition,[],[f292,f335])).
fof(f5354,plain,(
  ( ! [X222,X220,X223,X221,X224] : (~p(s(bool,'|dollar|exists_1'(s(fun(X220,bool),'const.pair.UNCURRY_2'(s(fun(X221,fun(X222,fun(X220,bool))),X223),s('type.pair.prod'(X221,X222),X224)))))) | p(s(bool,happ(s(fun(X220,bool),'const.pair.UNCURRY_2'(s(fun(X221,fun(X222,fun(X220,bool))),X223),s('type.pair.prod'(X221,X222),X224))),s(X220,sK6(happ(s(fun('type.pair.prod'(X221,X222),fun(X220,bool)),'const.pair.UNCURRY_1'(s(fun(X221,fun(X222,fun(X220,bool))),X223))),s('type.pair.prod'(X221,X222),X224)),X220)))))) )),
  inference(forward_demodulation,[],[f5300,f335])).
fof(f5300,plain,(
  ( ! [X222,X220,X223,X221,X224] : (p(s(bool,happ(s(fun(X220,bool),'const.pair.UNCURRY_2'(s(fun(X221,fun(X222,fun(X220,bool))),X223),s('type.pair.prod'(X221,X222),X224))),s(X220,sK6(happ(s(fun('type.pair.prod'(X221,X222),fun(X220,bool)),'const.pair.UNCURRY_1'(s(fun(X221,fun(X222,fun(X220,bool))),X223))),s('type.pair.prod'(X221,X222),X224)),X220))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X220,bool),happ(s(fun('type.pair.prod'(X221,X222),fun(X220,bool)),'const.pair.UNCURRY_1'(s(fun(X221,fun(X222,fun(X220,bool))),X223))),s('type.pair.prod'(X221,X222),X224))))))) )),
  inference(superposition,[],[f288,f335])).
fof(f5353,plain,(
  ( ! [X218,X216,X215,X219,X217] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X215,bool),'const.pair.UNCURRY_2'(s(fun(X216,fun(X217,fun(X215,bool))),X218),s('type.pair.prod'(X216,X217),X219)))))) | p(s(bool,happ(s(fun(X215,bool),'const.pair.UNCURRY_2'(s(fun(X216,fun(X217,fun(X215,bool))),X218),s('type.pair.prod'(X216,X217),X219))),s(X215,sK5(happ(s(fun('type.pair.prod'(X216,X217),fun(X215,bool)),'const.pair.UNCURRY_1'(s(fun(X216,fun(X217,fun(X215,bool))),X218))),s('type.pair.prod'(X216,X217),X219)),X215))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X215,bool),'const.pair.UNCURRY_2'(s(fun(X216,fun(X217,fun(X215,bool))),X218),s('type.pair.prod'(X216,X217),X219))))))) )),
  inference(forward_demodulation,[],[f5352,f335])).
fof(f5352,plain,(
  ( ! [X218,X216,X215,X219,X217] : (p(s(bool,happ(s(fun(X215,bool),'const.pair.UNCURRY_2'(s(fun(X216,fun(X217,fun(X215,bool))),X218),s('type.pair.prod'(X216,X217),X219))),s(X215,sK5(happ(s(fun('type.pair.prod'(X216,X217),fun(X215,bool)),'const.pair.UNCURRY_1'(s(fun(X216,fun(X217,fun(X215,bool))),X218))),s('type.pair.prod'(X216,X217),X219)),X215))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X215,bool),'const.pair.UNCURRY_2'(s(fun(X216,fun(X217,fun(X215,bool))),X218),s('type.pair.prod'(X216,X217),X219)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X215,bool),happ(s(fun('type.pair.prod'(X216,X217),fun(X215,bool)),'const.pair.UNCURRY_1'(s(fun(X216,fun(X217,fun(X215,bool))),X218))),s('type.pair.prod'(X216,X217),X219))))))) )),
  inference(forward_demodulation,[],[f5299,f335])).
fof(f5299,plain,(
  ( ! [X218,X216,X215,X219,X217] : (~p(s(bool,'|dollar|exists_1'(s(fun(X215,bool),'const.pair.UNCURRY_2'(s(fun(X216,fun(X217,fun(X215,bool))),X218),s('type.pair.prod'(X216,X217),X219)))))) | p(s(bool,happ(s(fun(X215,bool),happ(s(fun('type.pair.prod'(X216,X217),fun(X215,bool)),'const.pair.UNCURRY_1'(s(fun(X216,fun(X217,fun(X215,bool))),X218))),s('type.pair.prod'(X216,X217),X219))),s(X215,sK5(happ(s(fun('type.pair.prod'(X216,X217),fun(X215,bool)),'const.pair.UNCURRY_1'(s(fun(X216,fun(X217,fun(X215,bool))),X218))),s('type.pair.prod'(X216,X217),X219)),X215))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X215,bool),happ(s(fun('type.pair.prod'(X216,X217),fun(X215,bool)),'const.pair.UNCURRY_1'(s(fun(X216,fun(X217,fun(X215,bool))),X218))),s('type.pair.prod'(X216,X217),X219))))))) )),
  inference(superposition,[],[f286,f335])).
fof(f5351,plain,(
  ( ! [X212,X210,X213,X211,X214] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X210,bool),'const.pair.UNCURRY_2'(s(fun(X211,fun(X212,fun(X210,bool))),X213),s('type.pair.prod'(X211,X212),X214)))))) | p(s(bool,happ(s(fun(X210,bool),'const.pair.UNCURRY_2'(s(fun(X211,fun(X212,fun(X210,bool))),X213),s('type.pair.prod'(X211,X212),X214))),s(X210,sK4(happ(s(fun('type.pair.prod'(X211,X212),fun(X210,bool)),'const.pair.UNCURRY_1'(s(fun(X211,fun(X212,fun(X210,bool))),X213))),s('type.pair.prod'(X211,X212),X214)),X210))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X210,bool),'const.pair.UNCURRY_2'(s(fun(X211,fun(X212,fun(X210,bool))),X213),s('type.pair.prod'(X211,X212),X214))))))) )),
  inference(forward_demodulation,[],[f5350,f335])).
fof(f5350,plain,(
  ( ! [X212,X210,X213,X211,X214] : (p(s(bool,happ(s(fun(X210,bool),'const.pair.UNCURRY_2'(s(fun(X211,fun(X212,fun(X210,bool))),X213),s('type.pair.prod'(X211,X212),X214))),s(X210,sK4(happ(s(fun('type.pair.prod'(X211,X212),fun(X210,bool)),'const.pair.UNCURRY_1'(s(fun(X211,fun(X212,fun(X210,bool))),X213))),s('type.pair.prod'(X211,X212),X214)),X210))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X210,bool),'const.pair.UNCURRY_2'(s(fun(X211,fun(X212,fun(X210,bool))),X213),s('type.pair.prod'(X211,X212),X214)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X210,bool),happ(s(fun('type.pair.prod'(X211,X212),fun(X210,bool)),'const.pair.UNCURRY_1'(s(fun(X211,fun(X212,fun(X210,bool))),X213))),s('type.pair.prod'(X211,X212),X214))))))) )),
  inference(forward_demodulation,[],[f5298,f335])).
fof(f5298,plain,(
  ( ! [X212,X210,X213,X211,X214] : (~p(s(bool,'|dollar|exists_1'(s(fun(X210,bool),'const.pair.UNCURRY_2'(s(fun(X211,fun(X212,fun(X210,bool))),X213),s('type.pair.prod'(X211,X212),X214)))))) | p(s(bool,happ(s(fun(X210,bool),happ(s(fun('type.pair.prod'(X211,X212),fun(X210,bool)),'const.pair.UNCURRY_1'(s(fun(X211,fun(X212,fun(X210,bool))),X213))),s('type.pair.prod'(X211,X212),X214))),s(X210,sK4(happ(s(fun('type.pair.prod'(X211,X212),fun(X210,bool)),'const.pair.UNCURRY_1'(s(fun(X211,fun(X212,fun(X210,bool))),X213))),s('type.pair.prod'(X211,X212),X214)),X210))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X210,bool),happ(s(fun('type.pair.prod'(X211,X212),fun(X210,bool)),'const.pair.UNCURRY_1'(s(fun(X211,fun(X212,fun(X210,bool))),X213))),s('type.pair.prod'(X211,X212),X214))))))) )),
  inference(superposition,[],[f285,f335])).
fof(f5281,plain,(
  ( ! [X134,X132,X133] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X132,fun(X133,bool)),X134),s('type.pair.prod'(X132,X133),sK8('const.pair.UNCURRY_1'(s(fun(X132,fun(X133,bool)),X134)),'type.pair.prod'(X132,X133)))))) | p(s(bool,happ(s(fun(fun('type.pair.prod'(X132,X133),bool),bool),'!_1'),s(fun('type.pair.prod'(X132,X133),bool),'const.pair.UNCURRY_1'(s(fun(X132,fun(X133,bool)),X134))))))) )),
  inference(superposition,[],[f301,f335])).
fof(f5280,plain,(
  ( ! [X130,X131,X129] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X129,fun(X130,bool)),X131),s('type.pair.prod'(X129,X130),sK6('const.pair.UNCURRY_1'(s(fun(X129,fun(X130,bool)),X131)),'type.pair.prod'(X129,X130)))))) | ~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X129,X130),bool),'const.pair.UNCURRY_1'(s(fun(X129,fun(X130,bool)),X131))))))) )),
  inference(superposition,[],[f288,f335])).
fof(f5278,plain,(
  ( ! [X125,X123,X124,X122] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X122,fun(X123,bool)),X124),s('type.pair.prod'(X122,X123),X125)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X122,X123),bool),'const.pair.UNCURRY_1'(s(fun(X122,fun(X123,bool)),X124))))))) )),
  inference(superposition,[],[f289,f335])).
fof(f5277,plain,(
  ( ! [X121,X118,X120,X119] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X118,fun(X119,bool)),X120),s('type.pair.prod'(X118,X119),X121)))) | p(s(bool,happ(s(fun(fun('type.pair.prod'(X118,X119),bool),bool),'?_1'),s(fun('type.pair.prod'(X118,X119),bool),'const.pair.UNCURRY_1'(s(fun(X118,fun(X119,bool)),X120))))))) )),
  inference(superposition,[],[f303,f335])).
fof(f5337,plain,(
  ( ! [X116,X114,X117,X115,X113] : (p(s(bool,happ(s(fun(X114,bool),happ(s(fun(X113,fun(X114,bool)),X115),s(X113,X116))),s(X114,X117)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X113,fun(X114,bool)),X115),s('type.pair.prod'(X113,X114),sK12('const.pair.UNCURRY_1'(s(fun(X113,fun(X114,bool)),X115)),X114,X113)))))) )),
  inference(forward_demodulation,[],[f5276,f348])).
fof(f5276,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X113,fun(X114,bool)),X115),s('type.pair.prod'(X113,X114),sK12('const.pair.UNCURRY_1'(s(fun(X113,fun(X114,bool)),X115)),X114,X113))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X113,X114),bool),'const.pair.UNCURRY_1'(s(fun(X113,fun(X114,bool)),X115))),s(X113,X116),s(X114,X117))))) )),
  inference(superposition,[],[f376,f335])).
fof(f5258,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X8,fun(X9,bool)),X10),s('type.pair.prod'(X8,X9),X11))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X8,X9),bool),'const.pair.UNCURRY_1'(s(fun(X8,fun(X9,bool)),X10))),s('type.pair.prod'(X8,X9),X11)))) )),
  inference(superposition,[],[f335,f264])).
fof(f5257,plain,(
  ( ! [X6,X7,X5] : (s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X5,X6),bool),'const.pair.UNCURRY_1'(s(fun(X5,fun(X6,bool)),X7))))) = s(bool,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,bool)),X7),s('type.pair.prod'(X5,X6),'const.min.@_1'(s(fun('type.pair.prod'(X5,X6),bool),'const.pair.UNCURRY_1'(s(fun(X5,fun(X6,bool)),X7)))))))) )),
  inference(superposition,[],[f335,f280])).
fof(f5330,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK22('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0))))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,X3),s(X1,X4))))) )),
  inference(backward_demodulation,[],[f5249,f5026])).
fof(f5329,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK9('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1)))))) | ~p(s(bool,happ(s(fun(fun('type.pair.prod'(X0,X1),bool),bool),'?_1'),s(fun('type.pair.prod'(X0,X1),bool),X2))))) )),
  inference(backward_demodulation,[],[f5249,f3638])).
fof(f5328,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK5('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))) | ~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))) )),
  inference(backward_demodulation,[],[f5249,f4416])).
fof(f5327,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK4('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))) | ~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))) )),
  inference(backward_demodulation,[],[f5249,f4230])).
fof(f5249,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X2,happ(s(fun('type.pair.prod'(X0,X1),X2),X3),s('type.pair.prod'(X0,X1),X4))) = s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),X4)))) )),
  inference(superposition,[],[f335,f333])).
fof(f5133,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X30,X31),bool),sK28(X33,X32,fun('type.pair.prod'(X30,X31),bool))),s(X30,X34),s(X31,X35)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X30,X31),bool),sK28(X33,X32,fun('type.pair.prod'(X30,X31),bool))),s('type.pair.prod'(X30,X31),sK12('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X30,X31),bool),X32),X33)),X31,X30)))))) )),
  inference(forward_demodulation,[],[f5125,f701])).
fof(f5125,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(X30,X31),bool),sK28(X33,X32,fun('type.pair.prod'(X30,X31),bool))),s('type.pair.prod'(X30,X31),sK12('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X30,X31),bool),X32),X33)),X31,X30))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X30,X31),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X30,X31),bool),X32),X33))),s(X30,X34),s(X31,X35))))) )),
  inference(superposition,[],[f376,f701])).
fof(f5132,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X23,X24),bool),X26),s(X23,X28),s(X24,X29)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X23,X24),bool),X26),s('type.pair.prod'(X23,X24),sK12('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X23,X24),bool),X25),'const.pair.,_2'(s(fun('type.pair.prod'(X23,X24),bool),X26),s(X25,X27)))),X24,X23)))))) )),
  inference(forward_demodulation,[],[f5124,f323])).
fof(f5124,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(X23,X24),bool),X26),s('type.pair.prod'(X23,X24),sK12('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X23,X24),bool),X25),'const.pair.,_2'(s(fun('type.pair.prod'(X23,X24),bool),X26),s(X25,X27)))),X24,X23))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X23,X24),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X23,X24),bool),X25),'const.pair.,_2'(s(fun('type.pair.prod'(X23,X24),bool),X26),s(X25,X27))))),s(X23,X28),s(X24,X29))))) )),
  inference(superposition,[],[f376,f323])).
fof(f5131,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),bool),sK29(X20,fun('type.pair.prod'(X17,X18),bool),X19)),s(X17,X21),s(X18,X22)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X17,X18),bool),sK29(X20,fun('type.pair.prod'(X17,X18),bool),X19)),s('type.pair.prod'(X17,X18),sK12('const.pair.SND_1'(s('type.pair.prod'(X19,fun('type.pair.prod'(X17,X18),bool)),X20)),X18,X17)))))) )),
  inference(forward_demodulation,[],[f5123,f702])).
fof(f5123,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (~p(s(bool,happ(s(fun('type.pair.prod'(X17,X18),bool),sK29(X20,fun('type.pair.prod'(X17,X18),bool),X19)),s('type.pair.prod'(X17,X18),sK12('const.pair.SND_1'(s('type.pair.prod'(X19,fun('type.pair.prod'(X17,X18),bool)),X20)),X18,X17))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),bool),'const.pair.SND_1'(s('type.pair.prod'(X19,fun('type.pair.prod'(X17,X18),bool)),X20))),s(X17,X21),s(X18,X22))))) )),
  inference(superposition,[],[f376,f702])).
fof(f5130,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X16] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X10,X11),bool),X14),s(X10,X15),s(X11,X16)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),X14),s('type.pair.prod'(X10,X11),sK12('const.pair.SND_1'(s('type.pair.prod'(X12,fun('type.pair.prod'(X10,X11),bool)),'const.pair.,_2'(s(X12,X13),s(fun('type.pair.prod'(X10,X11),bool),X14)))),X11,X10)))))) )),
  inference(forward_demodulation,[],[f5122,f322])).
fof(f5122,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),X14),s('type.pair.prod'(X10,X11),sK12('const.pair.SND_1'(s('type.pair.prod'(X12,fun('type.pair.prod'(X10,X11),bool)),'const.pair.,_2'(s(X12,X13),s(fun('type.pair.prod'(X10,X11),bool),X14)))),X11,X10))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X10,X11),bool),'const.pair.SND_1'(s('type.pair.prod'(X12,fun('type.pair.prod'(X10,X11),bool)),'const.pair.,_2'(s(X12,X13),s(fun('type.pair.prod'(X10,X11),bool),X14))))),s(X10,X15),s(X11,X16))))) )),
  inference(superposition,[],[f376,f322])).
fof(f5129,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X5,X6),bool),X7),s(X5,X8),s(X6,X9)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X5,X6),bool),X7),s('type.pair.prod'(X5,X6),sK12(happ(s(fun(fun('type.pair.prod'(X5,X6),bool),fun('type.pair.prod'(X5,X6),bool)),'_i_0_0'),s(fun('type.pair.prod'(X5,X6),bool),X7)),X6,X5)))))) )),
  inference(forward_demodulation,[],[f5121,f279])).
fof(f5121,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X5,X6),bool),X7),s('type.pair.prod'(X5,X6),sK12(happ(s(fun(fun('type.pair.prod'(X5,X6),bool),fun('type.pair.prod'(X5,X6),bool)),'_i_0_0'),s(fun('type.pair.prod'(X5,X6),bool),X7)),X6,X5))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X5,X6),bool),happ(s(fun(fun('type.pair.prod'(X5,X6),bool),fun('type.pair.prod'(X5,X6),bool)),'_i_0_0'),s(fun('type.pair.prod'(X5,X6),bool),X7))),s(X5,X8),s(X6,X9))))) )),
  inference(superposition,[],[f376,f279])).
fof(f5128,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,X3),s(X1,X4)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK12('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)))))) )),
  inference(forward_demodulation,[],[f5127,f347])).
fof(f5127,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))),s(X0,X3))),s(X1,X4)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK12('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)))))) )),
  inference(forward_demodulation,[],[f5120,f348])).
fof(f5120,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK12('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))),s(X0,X3),s(X1,X4))))) )),
  inference(superposition,[],[f376,f333])).
fof(f5031,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (p(s(bool,happ(s(fun('type.pair.prod'(X30,X31),bool),sK28(X33,X32,fun('type.pair.prod'(X30,X31),bool))),s('type.pair.prod'(X30,X31),sK22('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X30,X31),bool),X32),X33)),X31,X30))))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X30,X31),bool),sK28(X33,X32,fun('type.pair.prod'(X30,X31),bool))),s(X30,X34),s(X31,X35))))) )),
  inference(forward_demodulation,[],[f4984,f701])).
fof(f4984,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X30,X31),bool),sK28(X33,X32,fun('type.pair.prod'(X30,X31),bool))),s(X30,X34),s(X31,X35)))) | p(s(bool,happ(s(fun('type.pair.prod'(X30,X31),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X30,X31),bool),X32),X33))),s('type.pair.prod'(X30,X31),sK22('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X30,X31),bool),X32),X33)),X31,X30)))))) )),
  inference(superposition,[],[f375,f701])).
fof(f5030,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,happ(s(fun('type.pair.prod'(X23,X24),bool),X26),s('type.pair.prod'(X23,X24),sK22('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X23,X24),bool),X25),'const.pair.,_2'(s(fun('type.pair.prod'(X23,X24),bool),X26),s(X25,X27)))),X24,X23))))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X23,X24),bool),X26),s(X23,X28),s(X24,X29))))) )),
  inference(forward_demodulation,[],[f4983,f323])).
fof(f4983,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X23,X24),bool),X26),s(X23,X28),s(X24,X29)))) | p(s(bool,happ(s(fun('type.pair.prod'(X23,X24),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X23,X24),bool),X25),'const.pair.,_2'(s(fun('type.pair.prod'(X23,X24),bool),X26),s(X25,X27))))),s('type.pair.prod'(X23,X24),sK22('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X23,X24),bool),X25),'const.pair.,_2'(s(fun('type.pair.prod'(X23,X24),bool),X26),s(X25,X27)))),X24,X23)))))) )),
  inference(superposition,[],[f375,f323])).
fof(f5029,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (p(s(bool,happ(s(fun('type.pair.prod'(X17,X18),bool),sK29(X20,fun('type.pair.prod'(X17,X18),bool),X19)),s('type.pair.prod'(X17,X18),sK22('const.pair.SND_1'(s('type.pair.prod'(X19,fun('type.pair.prod'(X17,X18),bool)),X20)),X18,X17))))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),bool),sK29(X20,fun('type.pair.prod'(X17,X18),bool),X19)),s(X17,X21),s(X18,X22))))) )),
  inference(forward_demodulation,[],[f4982,f702])).
fof(f4982,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X17,X18),bool),sK29(X20,fun('type.pair.prod'(X17,X18),bool),X19)),s(X17,X21),s(X18,X22)))) | p(s(bool,happ(s(fun('type.pair.prod'(X17,X18),bool),'const.pair.SND_1'(s('type.pair.prod'(X19,fun('type.pair.prod'(X17,X18),bool)),X20))),s('type.pair.prod'(X17,X18),sK22('const.pair.SND_1'(s('type.pair.prod'(X19,fun('type.pair.prod'(X17,X18),bool)),X20)),X18,X17)))))) )),
  inference(superposition,[],[f375,f702])).
fof(f5028,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X16] : (p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),X14),s('type.pair.prod'(X10,X11),sK22('const.pair.SND_1'(s('type.pair.prod'(X12,fun('type.pair.prod'(X10,X11),bool)),'const.pair.,_2'(s(X12,X13),s(fun('type.pair.prod'(X10,X11),bool),X14)))),X11,X10))))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X10,X11),bool),X14),s(X10,X15),s(X11,X16))))) )),
  inference(forward_demodulation,[],[f4981,f322])).
fof(f4981,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X16] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X10,X11),bool),X14),s(X10,X15),s(X11,X16)))) | p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),'const.pair.SND_1'(s('type.pair.prod'(X12,fun('type.pair.prod'(X10,X11),bool)),'const.pair.,_2'(s(X12,X13),s(fun('type.pair.prod'(X10,X11),bool),X14))))),s('type.pair.prod'(X10,X11),sK22('const.pair.SND_1'(s('type.pair.prod'(X12,fun('type.pair.prod'(X10,X11),bool)),'const.pair.,_2'(s(X12,X13),s(fun('type.pair.prod'(X10,X11),bool),X14)))),X11,X10)))))) )),
  inference(superposition,[],[f375,f322])).
fof(f5027,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun('type.pair.prod'(X5,X6),bool),X7),s('type.pair.prod'(X5,X6),sK22(happ(s(fun(fun('type.pair.prod'(X5,X6),bool),fun('type.pair.prod'(X5,X6),bool)),'_i_0_0'),s(fun('type.pair.prod'(X5,X6),bool),X7)),X6,X5))))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X5,X6),bool),X7),s(X5,X8),s(X6,X9))))) )),
  inference(forward_demodulation,[],[f4980,f279])).
fof(f4980,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X5,X6),bool),X7),s(X5,X8),s(X6,X9)))) | p(s(bool,happ(s(fun('type.pair.prod'(X5,X6),bool),happ(s(fun(fun('type.pair.prod'(X5,X6),bool),fun('type.pair.prod'(X5,X6),bool)),'_i_0_0'),s(fun('type.pair.prod'(X5,X6),bool),X7))),s('type.pair.prod'(X5,X6),sK22(happ(s(fun(fun('type.pair.prod'(X5,X6),bool),fun('type.pair.prod'(X5,X6),bool)),'_i_0_0'),s(fun('type.pair.prod'(X5,X6),bool),X7)),X6,X5)))))) )),
  inference(superposition,[],[f375,f279])).
fof(f5026,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))),s('type.pair.prod'(X0,X1),sK22('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0))))) | ~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,X3),s(X1,X4))))) )),
  inference(forward_demodulation,[],[f4979,f335])).
fof(f4979,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,X3),s(X1,X4)))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))),s('type.pair.prod'(X0,X1),sK22('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),X1,X0)))))) )),
  inference(superposition,[],[f375,f333])).
fof(f4851,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),sK32(X0,X1)),s(X1,X2)))))) | s(X0,sK34(X0,X1)) = s(X0,'const.min.@_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),sK32(X0,X1)),s(X1,X2)))))) )),
  inference(superposition,[],[f328,f280])).
fof(f4839,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X30,fun(bool,bool)),sK32(bool,X30)),s(X30,X31))),s(bool,X29)))) | s(bool,'F_0') = s(bool,sK34(bool,X30)) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f328,f264])).
fof(f4838,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X27,fun(bool,bool)),sK32(bool,X27)),s(X27,X28))),s(bool,'F_0')))) | s(bool,X26) = s(bool,sK34(bool,X27)) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f328,f264])).
fof(f4753,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),sK32(X0,X1)),s(X1,X2)))))) | s(X1,X2) = s(X1,sK33(X0,X1))) )),
  inference(superposition,[],[f327,f280])).
fof(f4723,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,happ(s(fun(X30,bool),happ(s(fun(bool,fun(X30,bool)),sK32(X30,bool)),s(bool,X29))),s(X30,X31)))) | s(bool,'F_0') = s(bool,sK33(X30,bool)) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f327,f264])).
fof(f4722,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,happ(s(fun(X27,bool),happ(s(fun(bool,fun(X27,bool)),sK32(X27,bool)),s(bool,'F_0'))),s(X27,X28)))) | s(bool,X26) = s(bool,sK33(X27,bool)) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f327,f264])).
fof(f4620,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK30(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK30(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f325,f264])).
fof(f4619,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK30(X0,X1,X2,bool))) )),
  inference(superposition,[],[f325,f264])).
fof(f4554,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))))) | $spl28),
  inference(forward_demodulation,[],[f4549,f346])).
fof(f4549,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))))),s(bool,'T_0')))) | $spl28),
  inference(unit_resulting_resolution,[],[f261,f4497,f306])).
fof(f4551,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))))) | $spl28),
  inference(unit_resulting_resolution,[],[f261,f4497,f306])).
fof(f4552,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))))),s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))))) | $spl28),
  inference(unit_resulting_resolution,[],[f4497,f4497,f306])).
fof(f4426,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool)))))) | p(s(bool,happ(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool))),s(X16,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18)),X16))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool))))))) )),
  inference(forward_demodulation,[],[f4425,f701])).
fof(f4425,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool))),s(X16,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18)),X16))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18))))))) )),
  inference(forward_demodulation,[],[f4412,f701])).
fof(f4412,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool)))))) | p(s(bool,happ(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18))),s(X16,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18)),X16))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18))))))) )),
  inference(superposition,[],[f286,f701])).
fof(f4424,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X12,bool),X14)))) | p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14))))) )),
  inference(forward_demodulation,[],[f4423,f323])).
fof(f4423,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X12,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15))))))))) )),
  inference(forward_demodulation,[],[f4411,f323])).
fof(f4411,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | p(s(bool,happ(s(fun(X12,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15))))),s(X12,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X12,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15))))))))) )),
  inference(superposition,[],[f286,f323])).
fof(f4422,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10))))) | p(s(bool,happ(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10)),s(X9,sK5('const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11)),X9))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10)))))) )),
  inference(forward_demodulation,[],[f4421,f702])).
fof(f4421,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10)),s(X9,sK5('const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11)),X9))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X9,bool),'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11))))))) )),
  inference(forward_demodulation,[],[f4410,f702])).
fof(f4410,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10))))) | p(s(bool,happ(s(fun(X9,bool),'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11))),s(X9,sK5('const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11)),X9))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X9,bool),'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11))))))) )),
  inference(superposition,[],[f286,f702])).
fof(f4420,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X5,bool),X8)))) | p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK5('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X8))))) )),
  inference(forward_demodulation,[],[f4419,f322])).
fof(f4419,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK5('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X8)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X5,bool),'const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8))))))))) )),
  inference(forward_demodulation,[],[f4409,f322])).
fof(f4409,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X8)))) | p(s(bool,happ(s(fun(X5,bool),'const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8))))),s(X5,sK5('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X5,bool),'const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8))))))))) )),
  inference(superposition,[],[f286,f322])).
fof(f4418,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X3,bool),X4)))) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK5(happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4)),X3))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4))))) )),
  inference(forward_demodulation,[],[f4417,f279])).
fof(f4417,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK5(happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4)),X3))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X3,bool),happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4))))))) )),
  inference(forward_demodulation,[],[f4408,f279])).
fof(f4408,plain,(
  ( ! [X4,X3] : (~p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | p(s(bool,happ(s(fun(X3,bool),happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4))),s(X3,sK5(happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4)),X3))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X3,bool),happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4))))))) )),
  inference(superposition,[],[f286,f279])).
fof(f4416,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))),s('type.pair.prod'(X0,X1),sK5('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1)))))) | ~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))) )),
  inference(forward_demodulation,[],[f4415,f333])).
fof(f4415,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))),s('type.pair.prod'(X0,X1),sK5('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1)))))) | ~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))) | p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))))))) )),
  inference(forward_demodulation,[],[f4407,f335])).
fof(f4407,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))),s('type.pair.prod'(X0,X1),sK5('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))))))) )),
  inference(superposition,[],[f286,f333])).
fof(f811,plain,(
  ( ! [X31,X32] : (s(bool,'F_0') = s(bool,sK28(X32,X31,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),X32)))) )),
  inference(superposition,[],[f264,f701])).
fof(f849,plain,(
  ( ! [X1] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1)))) )),
  inference(superposition,[],[f280,f264])).
fof(f854,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))))) )),
  inference(superposition,[],[f264,f280])).
fof(f1244,plain,(
  ( ! [X35,X36] : (s(bool,'F_0') = s(bool,sK29(X36,bool,X35)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X35,bool),X36)))) )),
  inference(superposition,[],[f264,f702])).
fof(f1944,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,sK4(X1,bool)) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X1)))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK5(X1,bool))) )),
  inference(superposition,[],[f287,f264])).
fof(f2356,plain,(
  ( ! [X28,X29] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X29)))) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f330,f264])).
fof(f2383,plain,(
  ( ! [X28,X29] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X29))),s(bool,'F_0')))) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f330,f264])).
fof(f2548,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X29)))) | s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f331,f264])).
fof(f2575,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X29))),s(bool,'F_0')))) | s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f331,f264])).
fof(f3517,plain,(
  ( ! [X8] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X8)))) | s(bool,'T_0') = s(bool,sK8(X8,bool))) )),
  inference(superposition,[],[f301,f264])).
fof(f4240,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool)))))) | p(s(bool,happ(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool))),s(X16,sK4('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18)),X16))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool))))))) )),
  inference(forward_demodulation,[],[f4239,f701])).
fof(f4239,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool))),s(X16,sK4('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18)),X16))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18))))))) )),
  inference(forward_demodulation,[],[f4219,f701])).
fof(f4219,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool)))))) | p(s(bool,happ(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18))),s(X16,sK4('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18)),X16))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18))))))) )),
  inference(superposition,[],[f285,f701])).
fof(f4238,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X12,bool),X14)))) | p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4('const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14))))) )),
  inference(forward_demodulation,[],[f4237,f323])).
fof(f4237,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4('const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X12,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15))))))))) )),
  inference(forward_demodulation,[],[f4218,f323])).
fof(f4218,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | p(s(bool,happ(s(fun(X12,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15))))),s(X12,sK4('const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X12,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15))))))))) )),
  inference(superposition,[],[f285,f323])).
fof(f4236,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10))))) | p(s(bool,happ(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10)),s(X9,sK4('const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11)),X9))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10)))))) )),
  inference(forward_demodulation,[],[f4235,f702])).
fof(f4235,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10)),s(X9,sK4('const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11)),X9))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X9,bool),'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11))))))) )),
  inference(forward_demodulation,[],[f4217,f702])).
fof(f4217,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10))))) | p(s(bool,happ(s(fun(X9,bool),'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11))),s(X9,sK4('const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11)),X9))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X9,bool),'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11))))))) )),
  inference(superposition,[],[f285,f702])).
fof(f4234,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X5,bool),X8)))) | p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK4('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X8))))) )),
  inference(forward_demodulation,[],[f4233,f322])).
fof(f4233,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK4('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X8)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X5,bool),'const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8))))))))) )),
  inference(forward_demodulation,[],[f4216,f322])).
fof(f4216,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X8)))) | p(s(bool,happ(s(fun(X5,bool),'const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8))))),s(X5,sK4('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X5,bool),'const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8))))))))) )),
  inference(superposition,[],[f285,f322])).
fof(f4232,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X3,bool),X4)))) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK4(happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4)),X3))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4))))) )),
  inference(forward_demodulation,[],[f4231,f279])).
fof(f4231,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK4(happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4)),X3))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X3,bool),happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4))))))) )),
  inference(forward_demodulation,[],[f4215,f279])).
fof(f4215,plain,(
  ( ! [X4,X3] : (~p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | p(s(bool,happ(s(fun(X3,bool),happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4))),s(X3,sK4(happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4)),X3))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X3,bool),happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4))))))) )),
  inference(superposition,[],[f285,f279])).
fof(f4230,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))),s('type.pair.prod'(X0,X1),sK4('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1)))))) | ~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))) )),
  inference(forward_demodulation,[],[f4229,f333])).
fof(f4229,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))),s('type.pair.prod'(X0,X1),sK4('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1)))))) | ~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))) | p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))))))) )),
  inference(forward_demodulation,[],[f4214,f335])).
fof(f4214,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))),s('type.pair.prod'(X0,X1),sK4('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))))))) )),
  inference(superposition,[],[f285,f333])).
fof(f4149,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) ) | $spl1),
  inference(subsumption_resolution,[],[f4103,f263])).
fof(f4103,plain,(
  ( ! [X8] : (p(s(bool,'F_0')) | s(bool,X8) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) ) | $spl1),
  inference(superposition,[],[f468,f2503])).
fof(f4080,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f2503,f264])).
fof(f4141,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))))),s(bool,'F_0')))) | $spl1),
  inference(forward_demodulation,[],[f4078,f352])).
fof(f4078,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))))) | $spl1),
  inference(unit_resulting_resolution,[],[f2503,f331])).
fof(f4065,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,sK29(X30,bool,X29)))),s(bool,X31)))) | ~p(s(bool,X31)) | s(bool,'F_0') = s(bool,sK29(X30,bool,X29))) ) | $spl1),
  inference(forward_demodulation,[],[f4064,f702])).
fof(f4064,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X31)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'const.pair.SND_1'(s('type.pair.prod'(X29,bool),X30))))),s(bool,X31)))) | s(bool,'F_0') = s(bool,sK29(X30,bool,X29))) ) | $spl1),
  inference(subsumption_resolution,[],[f4028,f261])).
fof(f4028,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'const.pair.SND_1'(s('type.pair.prod'(X29,bool),X30))))),s(bool,X31)))) | s(bool,'F_0') = s(bool,sK29(X30,bool,X29))) ) | $spl1),
  inference(superposition,[],[f306,f1279])).
fof(f4063,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,sK28(X27,X26,bool)))),s(bool,X28)))) | ~p(s(bool,X28)) | s(bool,'F_0') = s(bool,sK28(X27,X26,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f4062,f701])).
fof(f4062,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,X28)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X26),X27))))),s(bool,X28)))) | s(bool,'F_0') = s(bool,sK28(X27,X26,bool))) ) | $spl1),
  inference(subsumption_resolution,[],[f4027,f261])).
fof(f4027,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X28)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X26),X27))))),s(bool,X28)))) | s(bool,'F_0') = s(bool,sK28(X27,X26,bool))) ) | $spl1),
  inference(superposition,[],[f306,f831])).
fof(f4014,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))))),
  inference(unit_resulting_resolution,[],[f261,f261,f306])).
fof(f3633,plain,(
  ( ! [X0] : (~p(s(bool,'|dollar|exists_1'(s(fun(fun(X0,bool),bool),'?_1')))) | p(s(bool,happ(s(fun(X0,bool),'const.min.@_1'(s(fun(fun(X0,bool),bool),'?_1'))),s(X0,sK9('const.min.@_1'(s(fun(fun(X0,bool),bool),'?_1')),X0)))))) )),
  inference(superposition,[],[f302,f280])).
fof(f3643,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool))),s(X16,sK9('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18)),X16))))) | ~p(s(bool,happ(s(fun(fun(X16,bool),bool),'?_1'),s(fun(X16,bool),sK28(X18,X17,fun(X16,bool))))))) )),
  inference(forward_demodulation,[],[f3632,f701])).
fof(f3632,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(fun(X16,bool),bool),'?_1'),s(fun(X16,bool),sK28(X18,X17,fun(X16,bool)))))) | p(s(bool,happ(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18))),s(X16,sK9('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18)),X16)))))) )),
  inference(superposition,[],[f302,f701])).
fof(f3642,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK9('const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))))) | ~p(s(bool,happ(s(fun(fun(X12,bool),bool),'?_1'),s(fun(X12,bool),X14))))) )),
  inference(forward_demodulation,[],[f3631,f323])).
fof(f3631,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(fun(X12,bool),bool),'?_1'),s(fun(X12,bool),X14)))) | p(s(bool,happ(s(fun(X12,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15))))),s(X12,sK9('const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12)))))) )),
  inference(superposition,[],[f302,f323])).
fof(f3641,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10)),s(X9,sK9('const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11)),X9))))) | ~p(s(bool,happ(s(fun(fun(X9,bool),bool),'?_1'),s(fun(X9,bool),sK29(X11,fun(X9,bool),X10)))))) )),
  inference(forward_demodulation,[],[f3630,f702])).
fof(f3630,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(fun(X9,bool),bool),'?_1'),s(fun(X9,bool),sK29(X11,fun(X9,bool),X10))))) | p(s(bool,happ(s(fun(X9,bool),'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11))),s(X9,sK9('const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11)),X9)))))) )),
  inference(superposition,[],[f302,f702])).
fof(f3640,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK9('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5))))) | ~p(s(bool,happ(s(fun(fun(X5,bool),bool),'?_1'),s(fun(X5,bool),X8))))) )),
  inference(forward_demodulation,[],[f3629,f322])).
fof(f3629,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,happ(s(fun(fun(X5,bool),bool),'?_1'),s(fun(X5,bool),X8)))) | p(s(bool,happ(s(fun(X5,bool),'const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8))))),s(X5,sK9('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5)))))) )),
  inference(superposition,[],[f302,f322])).
fof(f3639,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK9(happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4)),X3))))) | ~p(s(bool,happ(s(fun(fun(X3,bool),bool),'?_1'),s(fun(X3,bool),X4))))) )),
  inference(forward_demodulation,[],[f3628,f279])).
fof(f3628,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(fun(X3,bool),bool),'?_1'),s(fun(X3,bool),X4)))) | p(s(bool,happ(s(fun(X3,bool),happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4))),s(X3,sK9(happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4)),X3)))))) )),
  inference(superposition,[],[f302,f279])).
fof(f3638,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))),s('type.pair.prod'(X0,X1),sK9('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1)))))) | ~p(s(bool,happ(s(fun(fun('type.pair.prod'(X0,X1),bool),bool),'?_1'),s(fun('type.pair.prod'(X0,X1),bool),X2))))) )),
  inference(forward_demodulation,[],[f3627,f335])).
fof(f3627,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(fun('type.pair.prod'(X0,X1),bool),bool),'?_1'),s(fun('type.pair.prod'(X0,X1),bool),X2)))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))),s('type.pair.prod'(X0,X1),sK9('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1))))))) )),
  inference(superposition,[],[f302,f333])).
fof(f3530,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,sK8(X3,bool)) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(forward_demodulation,[],[f3515,f803])).
fof(f3515,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,sK28('const.pair.,_2'(s(bool,X4),s(X5,X6)),X5,bool))))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,sK8(X3,bool)) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(superposition,[],[f301,f1051])).
fof(f3529,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK8(X0,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f3514,f803])).
fof(f3514,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,sK28('const.pair.,_2'(s(bool,'F_0'),s(X1,X2)),X1,bool))))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK8(X0,bool))) ) | $spl1),
  inference(superposition,[],[f301,f2283])).
fof(f3528,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun(fun(X16,bool),bool),'!_1'),s(fun(X16,bool),sK28(X18,X17,fun(X16,bool)))))) | ~p(s(bool,happ(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool))),s(X16,sK8('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18)),X16)))))) )),
  inference(forward_demodulation,[],[f3513,f701])).
fof(f3513,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),sK28(X18,X17,fun(X16,bool))),s(X16,sK8('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18)),X16))))) | p(s(bool,happ(s(fun(fun(X16,bool),bool),'!_1'),s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),X18))))))) )),
  inference(superposition,[],[f301,f701])).
fof(f3527,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(fun(X12,bool),bool),'!_1'),s(fun(X12,bool),X14)))) | ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK8('const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f3512,f323])).
fof(f3512,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK8('const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))))) | p(s(bool,happ(s(fun(fun(X12,bool),bool),'!_1'),s(fun(X12,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15))))))))) )),
  inference(superposition,[],[f301,f323])).
fof(f3526,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(fun(X9,bool),bool),'!_1'),s(fun(X9,bool),sK29(X11,fun(X9,bool),X10))))) | ~p(s(bool,happ(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10)),s(X9,sK8('const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11)),X9)))))) )),
  inference(forward_demodulation,[],[f3511,f702])).
fof(f3511,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),sK29(X11,fun(X9,bool),X10)),s(X9,sK8('const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11)),X9))))) | p(s(bool,happ(s(fun(fun(X9,bool),bool),'!_1'),s(fun(X9,bool),'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),X11))))))) )),
  inference(superposition,[],[f301,f702])).
fof(f3525,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(fun(X5,bool),bool),'!_1'),s(fun(X5,bool),X8)))) | ~p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK8('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5)))))) )),
  inference(forward_demodulation,[],[f3510,f322])).
fof(f3510,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK8('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5))))) | p(s(bool,happ(s(fun(fun(X5,bool),bool),'!_1'),s(fun(X5,bool),'const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8))))))))) )),
  inference(superposition,[],[f301,f322])).
fof(f3524,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),X4)))) | ~p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK8(happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4)),X3)))))) )),
  inference(forward_demodulation,[],[f3509,f279])).
fof(f3509,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK8(happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4)),X3))))) | p(s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),happ(s(fun(fun(X3,bool),fun(X3,bool)),'_i_0_0'),s(fun(X3,bool),X4))))))) )),
  inference(superposition,[],[f301,f279])).
fof(f3523,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(fun('type.pair.prod'(X0,X1),bool),bool),'!_1'),s(fun('type.pair.prod'(X0,X1),bool),X2)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK8('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1))))))) )),
  inference(forward_demodulation,[],[f3508,f333])).
fof(f3508,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK8('const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2)))),'type.pair.prod'(X0,X1)))))) | p(s(bool,happ(s(fun(fun('type.pair.prod'(X0,X1),bool),bool),'!_1'),s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),bool),X2))))))))) )),
  inference(superposition,[],[f301,f333])).
fof(f3460,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f2488,f264])).
fof(f3471,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))))),s(bool,'F_0')))) | $spl1),
  inference(forward_demodulation,[],[f3456,f352])).
fof(f3456,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))))) | $spl1),
  inference(unit_resulting_resolution,[],[f2488,f330])).
fof(f3418,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl1),
  inference(subsumption_resolution,[],[f3373,f263])).
fof(f3373,plain,(
  ( ! [X7] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl1),
  inference(superposition,[],[f468,f2479])).
fof(f3352,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f2479,f264])).
fof(f3410,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))))),s(bool,'T_0')))) | $spl1),
  inference(forward_demodulation,[],[f3350,f352])).
fof(f3350,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))))) | $spl1),
  inference(unit_resulting_resolution,[],[f2479,f331])).
fof(f3341,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X21,fun(X22,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X21,X22),bool),X23))),s('type.pair.prod'(X21,X22),X24)))) | ~p(s(bool,happ(s(fun(fun('type.pair.prod'(X21,X22),bool),bool),'!_1'),s(fun('type.pair.prod'(X21,X22),bool),X23))))) )),
  inference(forward_demodulation,[],[f3314,f335])).
fof(f3314,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,happ(s(fun(fun('type.pair.prod'(X21,X22),bool),bool),'!_1'),s(fun('type.pair.prod'(X21,X22),bool),X23)))) | p(s(bool,happ(s(fun('type.pair.prod'(X21,X22),bool),'const.pair.UNCURRY_1'(s(fun(X21,fun(X22,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X21,X22),bool),X23))))),s('type.pair.prod'(X21,X22),X24))))) )),
  inference(superposition,[],[f300,f333])).
fof(f3313,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X17,X18),bool),X19)))) | s('type.pair.prod'(X17,X18),X20) = s('type.pair.prod'(X17,X18),sK7(X20,'const.pair.UNCURRY_1'(s(fun(X17,fun(X18,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X17,X18),bool),X19)))),'type.pair.prod'(X17,X18)))) )),
  inference(superposition,[],[f292,f333])).
fof(f3339,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X10,X11),bool),X12)))) | p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),X12),s('type.pair.prod'(X10,X11),sK6('const.pair.UNCURRY_1'(s(fun(X10,fun(X11,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X10,X11),bool),X12)))),'type.pair.prod'(X10,X11))))))) )),
  inference(forward_demodulation,[],[f3311,f333])).
fof(f3311,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),X12),s('type.pair.prod'(X10,X11),sK6('const.pair.UNCURRY_1'(s(fun(X10,fun(X11,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X10,X11),bool),X12)))),'type.pair.prod'(X10,X11)))))) | ~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X10,X11),bool),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,bool)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X10,X11),bool),X12))))))))) )),
  inference(superposition,[],[f288,f333])).
fof(f3260,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f2464,f264])).
fof(f3271,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))))),s(bool,'T_0')))) | $spl1),
  inference(forward_demodulation,[],[f3256,f352])).
fof(f3256,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))))) | $spl1),
  inference(unit_resulting_resolution,[],[f2464,f330])).
fof(f3160,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,'T_0') != s(bool,sK28('const.pair.,_2'(s(bool,X86),s(X88,X89)),X88,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X86),s(X85,X87)))))) ) | $spl1),
  inference(superposition,[],[f2308,f323])).
fof(f3158,plain,(
  ( ! [X80,X78,X76,X79,X77] : (s(bool,'T_0') != s(bool,sK28('const.pair.,_2'(s(bool,X78),s(X79,X80)),X79,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X76,bool),'const.pair.,_2'(s(X76,X77),s(bool,X78)))))) ) | $spl1),
  inference(superposition,[],[f2308,f322])).
fof(f3146,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'T_0') != s(bool,sK28('const.pair.,_2'(s(bool,'|dollar|exists_1'(s(fun(X29,bool),X30))),s(X31,X32)),X31,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),X30),s(X29,'const.min.@_1'(s(fun(X29,bool),X30)))))) ) | $spl1),
  inference(superposition,[],[f2308,f280])).
fof(f2956,plain,(
  ( ! [X140,X139] : (s(bool,'T_0') != s(bool,sK28('const.pair.,_2'(s(bool,'F_0'),s(X139,X140)),X139,bool))) ) | $spl1),
  inference(subsumption_resolution,[],[f2806,f382])).
fof(f2806,plain,(
  ( ! [X140,X139] : (s(bool,'T_0') != s(bool,sK28('const.pair.,_2'(s(bool,'F_0'),s(X139,X140)),X139,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) ) | $spl1),
  inference(superposition,[],[f382,f2283])).
fof(f2736,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X86) = s(bool,sK28('const.pair.,_2'(s(bool,'F_0'),s(X88,X89)),X88,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X86),s(X85,X87)))))) ) | $spl1),
  inference(superposition,[],[f2283,f323])).
fof(f2734,plain,(
  ( ! [X80,X78,X76,X79,X77] : (s(bool,X78) = s(bool,sK28('const.pair.,_2'(s(bool,'F_0'),s(X79,X80)),X79,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X76,bool),'const.pair.,_2'(s(X76,X77),s(bool,X78)))))) ) | $spl1),
  inference(superposition,[],[f2283,f322])).
fof(f2733,plain,(
  ( ! [X74,X75,X73] : (s(bool,X73) = s(bool,sK28('const.pair.,_2'(s(bool,'F_0'),s(X74,X75)),X74,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X73)))) ) | $spl1),
  inference(superposition,[],[f2283,f279])).
fof(f2713,plain,(
  ( ! [X10,X8,X9] : (s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10)))) ) | $spl1),
  inference(subsumption_resolution,[],[f2712,f263])).
fof(f2712,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,'F_0')) | s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10)))) ) | $spl1),
  inference(forward_demodulation,[],[f2597,f391])).
fof(f2597,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10)))) ) | $spl1),
  inference(superposition,[],[f331,f595])).
fof(f2595,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1)))))) | s(X0,X1) != s(X0,'const.min.@_1'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1)))))) )),
  inference(superposition,[],[f331,f280])).
fof(f2696,plain,(
  ( ! [X45,X43,X44] : (s(bool,X45) != s(bool,sK29(X44,bool,X43)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X45))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK29(X44,bool,X43))) ) | $spl1),
  inference(forward_demodulation,[],[f2583,f702])).
fof(f2583,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X45))),s(bool,'T_0')))) | s(bool,X45) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X43,bool),X44))) | s(bool,'F_0') = s(bool,sK29(X44,bool,X43))) ) | $spl1),
  inference(superposition,[],[f331,f1279])).
fof(f2695,plain,(
  ( ! [X41,X42,X40] : (s(bool,X42) != s(bool,sK28(X41,X40,bool)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK28(X41,X40,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f2582,f701])).
fof(f2582,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,'T_0')))) | s(bool,X42) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X40),X41))) | s(bool,'F_0') = s(bool,sK28(X41,X40,bool))) ) | $spl1),
  inference(superposition,[],[f331,f831])).
fof(f2576,plain,(
  ( ! [X30,X31] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X31))),s(bool,X30)))) | s(bool,'F_0') != s(bool,X31) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f331,f264])).
fof(f2689,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X27))),s(bool,'F_0')))) | s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,X26)) ) | $spl1),
  inference(forward_demodulation,[],[f2574,f391])).
fof(f2574,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X27))),s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))))) | s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,X26)) ) | $spl1),
  inference(superposition,[],[f331,f595])).
fof(f2688,plain,(
  ( ! [X21,X25,X22] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X25))),s(bool,X22)))) | s(bool,X21) != s(bool,X25) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22)) ) | $spl1),
  inference(forward_demodulation,[],[f2573,f803])).
fof(f2573,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X25))),s(bool,sK28('const.pair.,_2'(s(bool,X22),s(X23,X24)),X23,bool))))) | s(bool,X21) != s(bool,X25) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22)) ) | $spl1),
  inference(superposition,[],[f331,f1051])).
fof(f2671,plain,(
  ( ! [X45,X43,X44] : (s(bool,X45) != s(bool,sK29(X44,bool,X43)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X45)))) | s(bool,'F_0') = s(bool,sK29(X44,bool,X43))) ) | $spl1),
  inference(forward_demodulation,[],[f2556,f702])).
fof(f2556,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X45)))) | s(bool,X45) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X43,bool),X44))) | s(bool,'F_0') = s(bool,sK29(X44,bool,X43))) ) | $spl1),
  inference(superposition,[],[f331,f1279])).
fof(f2670,plain,(
  ( ! [X41,X42,X40] : (s(bool,X42) != s(bool,sK28(X41,X40,bool)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X42)))) | s(bool,'F_0') = s(bool,sK28(X41,X40,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f2555,f701])).
fof(f2555,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X42)))) | s(bool,X42) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X40),X41))) | s(bool,'F_0') = s(bool,sK28(X41,X40,bool))) ) | $spl1),
  inference(superposition,[],[f331,f831])).
fof(f2549,plain,(
  ( ! [X30,X31] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X30))),s(bool,X31)))) | s(bool,'F_0') != s(bool,X31) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f331,f264])).
fof(f2664,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X27)))) | s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,X26)) ) | $spl1),
  inference(forward_demodulation,[],[f2547,f391])).
fof(f2547,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))))),s(bool,X27)))) | s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,X26)) ) | $spl1),
  inference(superposition,[],[f331,f595])).
fof(f2663,plain,(
  ( ! [X21,X25,X22] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X22))),s(bool,X25)))) | s(bool,X21) != s(bool,X25) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22)) ) | $spl1),
  inference(forward_demodulation,[],[f2546,f803])).
fof(f2546,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,sK28('const.pair.,_2'(s(bool,X22),s(X23,X24)),X23,bool)))),s(bool,X25)))) | s(bool,X21) != s(bool,X25) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22)) ) | $spl1),
  inference(superposition,[],[f331,f1051])).
fof(f2654,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1),s(fun(X0,X0),'_i_0_0'))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f2514,f352])).
fof(f2514,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f279,f331])).
fof(f2648,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun('type.pair.prod'(X0,X1),fun('type.pair.prod'(X0,X1),bool)),'=_0'),s('type.pair.prod'(X0,X1),X2))),s(X0,sK28(X2,X1,X0)),s(X1,sK29(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f2520,f343])).
fof(f2520,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun('type.pair.prod'(X0,X1),fun('type.pair.prod'(X0,X1),bool)),'=_0'),s('type.pair.prod'(X0,X1),X2))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK28(X2,X1,X0)),s(X1,sK29(X2,X1,X0)))))))) )),
  inference(unit_resulting_resolution,[],[f321,f331])).
fof(f2621,plain,(
  ( ! [X0,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3))),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f2620,f346])).
fof(f2620,plain,(
  ( ! [X0,X3] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3))),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f2619,f322])).
fof(f2619,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3))),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3))))))))) )),
  inference(forward_demodulation,[],[f2533,f346])).
fof(f2533,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3))))))),s(X0,X3))))) )),
  inference(unit_resulting_resolution,[],[f322,f331])).
fof(f2618,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,sK29(X2,X0,X1)))),s(X0,sK29(X2,X0,X1)))))) )),
  inference(forward_demodulation,[],[f2617,f346])).
fof(f2617,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,sK29(X2,X0,X1)))),s(X0,sK29(X2,X0,X1)))))) )),
  inference(forward_demodulation,[],[f2616,f702])).
fof(f2616,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,sK29(X2,X0,X1)))),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X1,X0),X2))))))) )),
  inference(forward_demodulation,[],[f2534,f346])).
fof(f2534,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X1,X0),X2))))),s(X0,sK29(X2,X0,X1)))))) )),
  inference(unit_resulting_resolution,[],[f702,f331])).
fof(f2615,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x1x0_2'(s(fun('type.pair.prod'(X0,X1),fun('type.pair.prod'(X0,X1),bool)),'=_0'),s('type.pair.prod'(X0,X1),X2))),s(X0,sK28(X2,X1,X0)),s(X1,sK29(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f2614,f343])).
fof(f2614,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x1x0_2'(s(fun('type.pair.prod'(X0,X1),fun('type.pair.prod'(X0,X1),bool)),'=_0'),s('type.pair.prod'(X0,X1),X2))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK28(X2,X1,X0)),s(X1,sK29(X2,X1,X0)))))))) )),
  inference(forward_demodulation,[],[f2535,f346])).
fof(f2535,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun('type.pair.prod'(X0,X1),fun('type.pair.prod'(X0,X1),bool)),'=_0'),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK28(X2,X1,X0)),s(X1,sK29(X2,X1,X0)))))),s('type.pair.prod'(X0,X1),X2))))) )),
  inference(unit_resulting_resolution,[],[f321,f331])).
fof(f2607,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),happ(s(fun('type.pair.prod'(sK1,sK0),fun('type.pair.prod'(sK1,sK0),bool)),'=_0'),s('type.pair.prod'(sK1,sK0),X0))),s('type.pair.prod'(sK1,sK0),X0))))) ) | $spl1),
  inference(forward_demodulation,[],[f2606,f346])).
fof(f2606,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'_dst_x1x0_2'(s(fun('type.pair.prod'(sK1,sK0),fun('type.pair.prod'(sK1,sK0),bool)),'=_0'),s('type.pair.prod'(sK1,sK0),X0))),s('type.pair.prod'(sK1,sK0),X0))))) ) | $spl1),
  inference(forward_demodulation,[],[f2605,f575])).
fof(f2605,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'_dst_x1x0_2'(s(fun('type.pair.prod'(sK1,sK0),fun('type.pair.prod'(sK1,sK0),bool)),'=_0'),s('type.pair.prod'(sK1,sK0),X0))),s('type.pair.prod'(sK1,sK0),sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))))))) ) | $spl1),
  inference(forward_demodulation,[],[f2538,f346])).
fof(f2538,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),happ(s(fun('type.pair.prod'(sK1,sK0),fun('type.pair.prod'(sK1,sK0),bool)),'=_0'),s('type.pair.prod'(sK1,sK0),sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))))),s('type.pair.prod'(sK1,sK0),X0))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f575,f331])).
fof(f2604,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),'=_0'),s(sK1,sK28(X0,sK0,sK1)))),s(sK1,sK28(X0,sK0,sK1)))))) ) | $spl1),
  inference(forward_demodulation,[],[f2603,f346])).
fof(f2603,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK1,bool)),'=_0'),s(sK1,sK28(X0,sK0,sK1)))),s(sK1,sK28(X0,sK0,sK1)))))) ) | $spl1),
  inference(forward_demodulation,[],[f2602,f825])).
fof(f2602,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK1,bool)),'=_0'),s(sK1,sK28(X0,sK0,sK1)))),s(sK1,sK28(sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)),sK0,sK1)))))) ) | $spl1),
  inference(forward_demodulation,[],[f2539,f346])).
fof(f2539,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),'=_0'),s(sK1,sK28(sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)),sK0,sK1)))),s(sK1,sK28(X0,sK0,sK1)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f825,f331])).
fof(f2601,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK29(X0,sK0,sK1)))),s(sK0,sK29(X0,sK0,sK1)))))) ) | $spl1),
  inference(forward_demodulation,[],[f2600,f346])).
fof(f2600,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK29(X0,sK0,sK1)))),s(sK0,sK29(X0,sK0,sK1)))))) ) | $spl1),
  inference(forward_demodulation,[],[f2599,f1270])).
fof(f2599,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK29(X0,sK0,sK1)))),s(sK0,sK29(sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)),sK0,sK1)))))) ) | $spl1),
  inference(forward_demodulation,[],[f2540,f346])).
fof(f2540,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK29(sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)),sK0,sK1)))),s(sK0,sK29(X0,sK0,sK1)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1270,f331])).
fof(f2487,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0')))))) | $spl1),
  inference(unit_resulting_resolution,[],[f2410,f300])).
fof(f2488,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))) | $spl1),
  inference(unit_resulting_resolution,[],[f2410,f265])).
fof(f2505,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))))),s(bool,'F_0'))))) ) | $spl1),
  inference(forward_demodulation,[],[f2495,f352])).
fof(f2495,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f2410,f305])).
fof(f2503,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))) | $spl1),
  inference(forward_demodulation,[],[f2497,f391])).
fof(f2497,plain,(
  s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))) | $spl1),
  inference(unit_resulting_resolution,[],[f382,f2410,f468])).
fof(f2463,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0')))))) | $spl1),
  inference(unit_resulting_resolution,[],[f2409,f300])).
fof(f2464,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))) | $spl1),
  inference(unit_resulting_resolution,[],[f2409,f265])).
fof(f2481,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))))),s(bool,'T_0'))))) ) | $spl1),
  inference(forward_demodulation,[],[f2471,f352])).
fof(f2471,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f2409,f305])).
fof(f2479,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))) | $spl1),
  inference(forward_demodulation,[],[f2473,f391])).
fof(f2473,plain,(
  s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))) | $spl1),
  inference(unit_resulting_resolution,[],[f382,f2409,f468])).
fof(f2403,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1)))))) | s(X0,X1) = s(X0,'const.min.@_1'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1)))))) )),
  inference(superposition,[],[f330,f280])).
fof(f2449,plain,(
  ( ! [X45,X43,X44] : (s(bool,X45) = s(bool,sK29(X44,bool,X43)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X45))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK29(X44,bool,X43))) ) | $spl1),
  inference(forward_demodulation,[],[f2391,f702])).
fof(f2391,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X45))),s(bool,'T_0')))) | s(bool,X45) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X43,bool),X44))) | s(bool,'F_0') = s(bool,sK29(X44,bool,X43))) ) | $spl1),
  inference(superposition,[],[f330,f1279])).
fof(f2448,plain,(
  ( ! [X41,X42,X40] : (s(bool,X42) = s(bool,sK28(X41,X40,bool)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK28(X41,X40,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f2390,f701])).
fof(f2390,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,'T_0')))) | s(bool,X42) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X40),X41))) | s(bool,'F_0') = s(bool,sK28(X41,X40,bool))) ) | $spl1),
  inference(superposition,[],[f330,f831])).
fof(f2384,plain,(
  ( ! [X30,X31] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X31))),s(bool,X30)))) | s(bool,'F_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f330,f264])).
fof(f2442,plain,(
  ( ! [X26,X27] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X27))),s(bool,'F_0')))) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X26)) ) | $spl1),
  inference(forward_demodulation,[],[f2382,f391])).
fof(f2382,plain,(
  ( ! [X26,X27] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X27))),s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))))) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X26)) ) | $spl1),
  inference(superposition,[],[f330,f595])).
fof(f2441,plain,(
  ( ! [X21,X25,X22] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X25))),s(bool,X22)))) | s(bool,X21) = s(bool,X25) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22)) ) | $spl1),
  inference(forward_demodulation,[],[f2381,f803])).
fof(f2381,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X25))),s(bool,sK28('const.pair.,_2'(s(bool,X22),s(X23,X24)),X23,bool))))) | s(bool,X21) = s(bool,X25) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22)) ) | $spl1),
  inference(superposition,[],[f330,f1051])).
fof(f2424,plain,(
  ( ! [X45,X43,X44] : (s(bool,X45) = s(bool,sK29(X44,bool,X43)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X45)))) | s(bool,'F_0') = s(bool,sK29(X44,bool,X43))) ) | $spl1),
  inference(forward_demodulation,[],[f2364,f702])).
fof(f2364,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X45)))) | s(bool,X45) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X43,bool),X44))) | s(bool,'F_0') = s(bool,sK29(X44,bool,X43))) ) | $spl1),
  inference(superposition,[],[f330,f1279])).
fof(f2423,plain,(
  ( ! [X41,X42,X40] : (s(bool,X42) = s(bool,sK28(X41,X40,bool)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X42)))) | s(bool,'F_0') = s(bool,sK28(X41,X40,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f2363,f701])).
fof(f2363,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X42)))) | s(bool,X42) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X40),X41))) | s(bool,'F_0') = s(bool,sK28(X41,X40,bool))) ) | $spl1),
  inference(superposition,[],[f330,f831])).
fof(f2357,plain,(
  ( ! [X30,X31] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X30))),s(bool,X31)))) | s(bool,'F_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f330,f264])).
fof(f2417,plain,(
  ( ! [X26,X27] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X27)))) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X26)) ) | $spl1),
  inference(forward_demodulation,[],[f2355,f391])).
fof(f2355,plain,(
  ( ! [X26,X27] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))))),s(bool,X27)))) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X26)) ) | $spl1),
  inference(superposition,[],[f330,f595])).
fof(f2416,plain,(
  ( ! [X21,X25,X22] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X22))),s(bool,X25)))) | s(bool,X21) = s(bool,X25) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22)) ) | $spl1),
  inference(forward_demodulation,[],[f2354,f803])).
fof(f2354,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,sK28('const.pair.,_2'(s(bool,X22),s(X23,X24)),X23,bool)))),s(bool,X25)))) | s(bool,X21) = s(bool,X25) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22)) ) | $spl1),
  inference(superposition,[],[f330,f1051])).
fof(f2410,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))) | $spl1),
  inference(forward_demodulation,[],[f2346,f391])).
fof(f2346,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f382,f330])).
fof(f2409,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))) | $spl1),
  inference(forward_demodulation,[],[f2408,f346])).
fof(f2408,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))) | $spl1),
  inference(forward_demodulation,[],[f2407,f391])).
fof(f2407,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))))) | $spl1),
  inference(forward_demodulation,[],[f2348,f346])).
fof(f2348,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))))),s(bool,'T_0')))) | $spl1),
  inference(unit_resulting_resolution,[],[f382,f330])).
fof(f2344,plain,(
  ( ! [X125,X123,X126,X124] : (s('type.pair.prod'(bool,X124),'const.pair.,_2'(s(bool,X123),s(X124,X125))) = s('type.pair.prod'(bool,X124),'const.pair.,_2'(s(bool,X126),s(X124,X125))) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X123)) ) | $spl1),
  inference(forward_demodulation,[],[f2200,f1233])).
fof(f2200,plain,(
  ( ! [X125,X123,X126,X124] : (s('type.pair.prod'(bool,X124),'const.pair.,_2'(s(bool,X123),s(X124,X125))) = s('type.pair.prod'(bool,X124),'const.pair.,_2'(s(bool,X126),s(X124,sK29('const.pair.,_2'(s(bool,X123),s(X124,X125)),X124,bool)))) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X123)) ) | $spl1),
  inference(superposition,[],[f321,f1051])).
fof(f2329,plain,(
  ( ! [X358,X357,X359] : (s('type.pair.prod'(X358,bool),X357) = s('type.pair.prod'(X358,bool),'const.pair.,_2'(s(X358,sK28(X357,bool,X358)),s(bool,X359))) | s(bool,'T_0') = s(bool,sK29(X357,bool,X358)) | s(bool,'T_0') = s(bool,X359)) ) | $spl1),
  inference(forward_demodulation,[],[f2173,f803])).
fof(f2173,plain,(
  ( ! [X358,X360,X357,X359,X361] : (s('type.pair.prod'(X358,bool),X357) = s('type.pair.prod'(X358,bool),'const.pair.,_2'(s(X358,sK28(X357,bool,X358)),s(bool,sK28('const.pair.,_2'(s(bool,X359),s(X360,X361)),X360,bool)))) | s(bool,'T_0') = s(bool,sK29(X357,bool,X358)) | s(bool,'T_0') = s(bool,X359)) ) | $spl1),
  inference(superposition,[],[f321,f1051])).
fof(f2328,plain,(
  ( ! [X352,X354,X353] : (s('type.pair.prod'(bool,X353),X352) = s('type.pair.prod'(bool,X353),'const.pair.,_2'(s(bool,X354),s(X353,sK29(X352,X353,bool)))) | s(bool,'T_0') = s(bool,sK28(X352,X353,bool)) | s(bool,'T_0') = s(bool,X354)) ) | $spl1),
  inference(forward_demodulation,[],[f2172,f803])).
fof(f2172,plain,(
  ( ! [X356,X352,X354,X353,X355] : (s('type.pair.prod'(bool,X353),X352) = s('type.pair.prod'(bool,X353),'const.pair.,_2'(s(bool,sK28('const.pair.,_2'(s(bool,X354),s(X355,X356)),X355,bool)),s(X353,sK29(X352,X353,bool)))) | s(bool,'T_0') = s(bool,sK28(X352,X353,bool)) | s(bool,'T_0') = s(bool,X354)) ) | $spl1),
  inference(superposition,[],[f321,f1051])).
fof(f2327,plain,(
  ( ! [X349,X348] : (p(s(bool,happ(s(fun(bool,bool),X348),s(bool,X349)))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X348)))) | s(bool,'T_0') = s(bool,sK6(X348,bool)) | s(bool,'T_0') = s(bool,X349)) ) | $spl1),
  inference(forward_demodulation,[],[f2171,f803])).
fof(f2171,plain,(
  ( ! [X350,X349,X351,X348] : (p(s(bool,happ(s(fun(bool,bool),X348),s(bool,sK28('const.pair.,_2'(s(bool,X349),s(X350,X351)),X350,bool))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X348)))) | s(bool,'T_0') = s(bool,sK6(X348,bool)) | s(bool,'T_0') = s(bool,X349)) ) | $spl1),
  inference(superposition,[],[f288,f1051])).
fof(f2326,plain,(
  ( ! [X344,X345] : (s(bool,X345) != s(bool,sK4(X344,bool)) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X344)))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X344)))) | s(bool,'T_0') = s(bool,sK5(X344,bool)) | s(bool,'T_0') = s(bool,X345)) ) | $spl1),
  inference(forward_demodulation,[],[f2170,f803])).
fof(f2170,plain,(
  ( ! [X344,X346,X345,X347] : (s(bool,sK4(X344,bool)) != s(bool,sK28('const.pair.,_2'(s(bool,X345),s(X346,X347)),X346,bool)) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X344)))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X344)))) | s(bool,'T_0') = s(bool,sK5(X344,bool)) | s(bool,'T_0') = s(bool,X345)) ) | $spl1),
  inference(superposition,[],[f287,f1051])).
fof(f2325,plain,(
  ( ! [X337,X336] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X336))) = s(bool,happ(s(fun(bool,bool),X336),s(bool,X337))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X336))) | s(bool,'T_0') = s(bool,X337)) ) | $spl1),
  inference(forward_demodulation,[],[f2169,f803])).
fof(f2169,plain,(
  ( ! [X337,X339,X336,X338] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X336))) = s(bool,happ(s(fun(bool,bool),X336),s(bool,sK28('const.pair.,_2'(s(bool,X337),s(X338,X339)),X338,bool)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X336))) | s(bool,'T_0') = s(bool,X337)) ) | $spl1),
  inference(superposition,[],[f280,f1051])).
fof(f2308,plain,(
  ( ! [X167,X166,X168] : (s(bool,'T_0') != s(bool,sK28('const.pair.,_2'(s(bool,X166),s(X167,X168)),X167,bool)) | s(bool,'T_0') = s(bool,X166)) ) | $spl1),
  inference(subsumption_resolution,[],[f2137,f382])).
fof(f2137,plain,(
  ( ! [X167,X166,X168] : (s(bool,'T_0') != s(bool,sK28('const.pair.,_2'(s(bool,X166),s(X167,X168)),X167,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | s(bool,'T_0') = s(bool,X166)) ) | $spl1),
  inference(superposition,[],[f382,f1051])).
fof(f2075,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (s(bool,X99) = s(bool,sK28('const.pair.,_2'(s(bool,X97),s(X100,X101)),X100,bool)) | s(bool,'T_0') = s(bool,X99) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X96),'const.pair.,_2'(s(bool,X97),s(X96,X98)))))) ) | $spl1),
  inference(superposition,[],[f1051,f323])).
fof(f2073,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (s(bool,X88) = s(bool,sK28('const.pair.,_2'(s(bool,X87),s(X89,X90)),X89,bool)) | s(bool,'T_0') = s(bool,X88) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X85,bool),'const.pair.,_2'(s(X85,X86),s(bool,X87)))))) ) | $spl1),
  inference(superposition,[],[f1051,f322])).
fof(f2062,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,X34) = s(bool,sK28('const.pair.,_2'(s(bool,'|dollar|exists_1'(s(fun(X32,bool),X33))),s(X35,X36)),X35,bool)) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,happ(s(fun(X32,bool),X33),s(X32,'const.min.@_1'(s(fun(X32,bool),X33)))))) ) | $spl1),
  inference(superposition,[],[f1051,f280])).
fof(f2283,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,sK28('const.pair.,_2'(s(bool,'F_0'),s(X1,X2)),X1,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(forward_demodulation,[],[f2270,f391])).
fof(f2270,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,sK28('const.pair.,_2'(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))),s(X1,X2)),X1,bool))) ) | $spl1),
  inference(condensation,[],[f2269])).
fof(f2269,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X8) = s(bool,sK28('const.pair.,_2'(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))),s(X9,X10)),X9,bool)) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) ) | $spl1),
  inference(duplicate_literal_removal,[],[f2055])).
fof(f2055,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X8) = s(bool,sK28('const.pair.,_2'(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))),s(X9,X10)),X9,bool)) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl1),
  inference(superposition,[],[f1051,f595])).
fof(f2053,plain,(
  ( ! [X103,X105,X102,X106,X104] : (s(bool,sK28(X103,X102,bool)) = s(bool,sK28('const.pair.,_2'(s(bool,X104),s(X105,X106)),X105,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X102),X103))) | s(bool,'T_0') = s(bool,X104)) ) | $spl1),
  inference(superposition,[],[f1051,f701])).
fof(f2052,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (s(bool,X97) = s(bool,sK28('const.pair.,_2'(s(bool,X99),s(X100,X101)),X100,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X96),'const.pair.,_2'(s(bool,X97),s(X96,X98))))) | s(bool,'T_0') = s(bool,X99)) ) | $spl1),
  inference(superposition,[],[f1051,f323])).
fof(f2051,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,sK29(X92,bool,X91)) = s(bool,sK28('const.pair.,_2'(s(bool,X93),s(X94,X95)),X94,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X91,bool),X92))) | s(bool,'T_0') = s(bool,X93)) ) | $spl1),
  inference(superposition,[],[f1051,f702])).
fof(f2050,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (s(bool,X87) = s(bool,sK28('const.pair.,_2'(s(bool,X88),s(X89,X90)),X89,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X85,bool),'const.pair.,_2'(s(X85,X86),s(bool,X87))))) | s(bool,'T_0') = s(bool,X88)) ) | $spl1),
  inference(superposition,[],[f1051,f322])).
fof(f2049,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,X81) = s(bool,sK28('const.pair.,_2'(s(bool,X82),s(X83,X84)),X83,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X81))) | s(bool,'T_0') = s(bool,X82)) ) | $spl1),
  inference(superposition,[],[f1051,f279])).
fof(f2041,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,'|dollar|exists_1'(s(fun(X32,bool),X33))) = s(bool,sK28('const.pair.,_2'(s(bool,X34),s(X35,X36)),X35,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(X32,bool),X33),s(X32,'const.min.@_1'(s(fun(X32,bool),X33))))) | s(bool,'T_0') = s(bool,X34)) ) | $spl1),
  inference(superposition,[],[f1051,f280])).
fof(f2278,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,sK28('const.pair.,_2'(s(bool,X0),s(X4,X5)),X4,bool)) = s(bool,sK28('const.pair.,_2'(s(bool,X1),s(X2,X3)),X2,bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(condensation,[],[f2277])).
fof(f2277,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,sK28('const.pair.,_2'(s(bool,X1),s(X2,X3)),X2,bool)) = s(bool,sK28('const.pair.,_2'(s(bool,X4),s(X5,X6)),X5,bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(duplicate_literal_removal,[],[f2035])).
fof(f2035,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,sK28('const.pair.,_2'(s(bool,X1),s(X2,X3)),X2,bool)) = s(bool,sK28('const.pair.,_2'(s(bool,X4),s(X5,X6)),X5,bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(superposition,[],[f1051,f1051])).
fof(f1945,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,sK4(X0,bool)) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0)))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK5(X0,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f1943,f391])).
fof(f1943,plain,(
  ( ! [X0] : (s(bool,sK4(X0,bool)) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0)))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK5(X0,bool))) ) | $spl1),
  inference(superposition,[],[f287,f595])).
fof(f1942,plain,(
  ( ! [X62,X63] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X62),s(bool,X63))) = s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X63),s(bool,X63))) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X63)) ) | $spl1),
  inference(forward_demodulation,[],[f1885,f1233])).
fof(f1885,plain,(
  ( ! [X62,X63] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X62),s(bool,X63))) = s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X63),s(bool,sK29('const.pair.,_2'(s(bool,X62),s(bool,X63)),bool,bool)))) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X63)) ) | $spl1),
  inference(superposition,[],[f321,f1003])).
fof(f1859,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X43) = s(bool,sK28('const.pair.,_2'(s(bool,X45),s(bool,X43)),bool,bool)) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X42),'const.pair.,_2'(s(bool,X43),s(X42,X44)))))) ) | $spl1),
  inference(superposition,[],[f1003,f323])).
fof(f1857,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) = s(bool,sK28('const.pair.,_2'(s(bool,X38),s(bool,X37)),bool,bool)) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X35,bool),'const.pair.,_2'(s(X35,X36),s(bool,X37)))))) ) | $spl1),
  inference(superposition,[],[f1003,f322])).
fof(f1848,plain,(
  ( ! [X12,X13,X11] : (s(bool,'|dollar|exists_1'(s(fun(X11,bool),X12))) = s(bool,sK28('const.pair.,_2'(s(bool,X13),s(bool,'|dollar|exists_1'(s(fun(X11,bool),X12)))),bool,bool)) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X12),s(X11,'const.min.@_1'(s(fun(X11,bool),X12)))))) ) | $spl1),
  inference(superposition,[],[f1003,f280])).
fof(f1839,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X45) = s(bool,sK28('const.pair.,_2'(s(bool,X43),s(bool,X45)),bool,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X42),'const.pair.,_2'(s(bool,X43),s(X42,X44))))) | s(bool,'T_0') = s(bool,X45)) ) | $spl1),
  inference(superposition,[],[f1003,f323])).
fof(f1837,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X38) = s(bool,sK28('const.pair.,_2'(s(bool,X37),s(bool,X38)),bool,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X35,bool),'const.pair.,_2'(s(X35,X36),s(bool,X37))))) | s(bool,'T_0') = s(bool,X38)) ) | $spl1),
  inference(superposition,[],[f1003,f322])).
fof(f1828,plain,(
  ( ! [X12,X13,X11] : (s(bool,X13) = s(bool,sK28('const.pair.,_2'(s(bool,'|dollar|exists_1'(s(fun(X11,bool),X12))),s(bool,X13)),bool,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X12),s(X11,'const.min.@_1'(s(fun(X11,bool),X12))))) | s(bool,'T_0') = s(bool,X13)) ) | $spl1),
  inference(superposition,[],[f1003,f280])).
fof(f1920,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK28('const.pair.,_2'(s(bool,'F_0'),s(bool,X0)),bool,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(forward_demodulation,[],[f1917,f391])).
fof(f1917,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,sK28('const.pair.,_2'(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))),s(bool,X0)),bool,bool))) ) | $spl1),
  inference(condensation,[],[f1916])).
fof(f1916,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK28('const.pair.,_2'(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))),s(bool,X1)),bool,bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(duplicate_literal_removal,[],[f1821])).
fof(f1821,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK28('const.pair.,_2'(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))),s(bool,X1)),bool,bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f1003,f595])).
fof(f1738,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,sK29(X28,bool,X29)) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X26,bool),X27))) | s(bool,'F_0') = s(bool,sK29(X28,bool,X29)) | s(bool,'F_0') = s(bool,sK29(X27,bool,X26))) ) | $spl1),
  inference(subsumption_resolution,[],[f1720,f261])).
fof(f1720,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'T_0')) | s(bool,sK29(X28,bool,X29)) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X26,bool),X27))) | s(bool,'F_0') = s(bool,sK29(X28,bool,X29)) | s(bool,'F_0') = s(bool,sK29(X27,bool,X26))) ) | $spl1),
  inference(superposition,[],[f1488,f1279])).
fof(f1737,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,sK29(X24,bool,X25)) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),X23))) | s(bool,'F_0') = s(bool,sK29(X24,bool,X25)) | s(bool,'F_0') = s(bool,sK28(X23,X22,bool))) ) | $spl1),
  inference(subsumption_resolution,[],[f1719,f261])).
fof(f1719,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'T_0')) | s(bool,sK29(X24,bool,X25)) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),X23))) | s(bool,'F_0') = s(bool,sK29(X24,bool,X25)) | s(bool,'F_0') = s(bool,sK28(X23,X22,bool))) ) | $spl1),
  inference(superposition,[],[f1488,f831])).
fof(f1669,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,sK29('const.pair.,_2'(s(X2,X3),s(bool,X1)),bool,X2)) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f1512,f264])).
fof(f1607,plain,(
  ( ! [X50,X51,X49] : (s('type.pair.prod'(X49,bool),'const.pair.,_2'(s(X49,X50),s(bool,'F_0'))) = s('type.pair.prod'(X49,bool),'const.pair.,_2'(s(X49,X50),s(bool,X51))) | s(bool,'T_0') = s(bool,X51)) ) | $spl1),
  inference(forward_demodulation,[],[f1578,f803])).
fof(f1578,plain,(
  ( ! [X50,X51,X49] : (s('type.pair.prod'(X49,bool),'const.pair.,_2'(s(X49,X50),s(bool,X51))) = s('type.pair.prod'(X49,bool),'const.pair.,_2'(s(X49,sK28('const.pair.,_2'(s(X49,X50),s(bool,X51)),bool,X49)),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X51)) ) | $spl1),
  inference(superposition,[],[f321,f1453])).
fof(f1605,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X30) = s(bool,sK29('const.pair.,_2'(s(X27,X28),s(bool,X29)),bool,X27)) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X29)) ) | $spl1),
  inference(subsumption_resolution,[],[f1573,f263])).
fof(f1573,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,'F_0')) | s(bool,X30) = s(bool,sK29('const.pair.,_2'(s(X27,X28),s(bool,X29)),bool,X27)) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X29)) ) | $spl1),
  inference(superposition,[],[f468,f1453])).
fof(f1563,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,'F_0') = s(bool,sK29('const.pair.,_2'(s(X59,X60),s(bool,X57)),bool,X59)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X56),'const.pair.,_2'(s(bool,X57),s(X56,X58)))))) ) | $spl1),
  inference(superposition,[],[f1453,f323])).
fof(f1561,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,sK29('const.pair.,_2'(s(X50,X51),s(bool,X49)),bool,X50)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X47,bool),'const.pair.,_2'(s(X47,X48),s(bool,X49)))))) ) | $spl1),
  inference(superposition,[],[f1453,f322])).
fof(f1553,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'F_0') = s(bool,sK29('const.pair.,_2'(s(X20,X21),s(bool,'|dollar|exists_1'(s(fun(X18,bool),X19)))),bool,X20)) | s(bool,'T_0') = s(bool,happ(s(fun(X18,bool),X19),s(X18,'const.min.@_1'(s(fun(X18,bool),X19)))))) ) | $spl1),
  inference(superposition,[],[f1453,f280])).
fof(f1531,plain,(
  ( ! [X26,X25] : (s(bool,X25) = s(bool,sK29('const.pair.,_2'(s(bool,X25),s(bool,X26)),bool,bool)) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) ) | $spl1),
  inference(superposition,[],[f1453,f264])).
fof(f1512,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X8) | s(bool,'F_0') = s(bool,sK29('const.pair.,_2'(s(X6,X7),s(bool,X8)),bool,X6))) ) | $spl1),
  inference(superposition,[],[f1486,f322])).
fof(f1521,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,sK29(sK28(X17,X16,'type.pair.prod'(X15,bool)),bool,X15)) | s(bool,'F_0') = s(bool,sK29('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X15,bool),X16),X17)),bool,X15))) ) | $spl1),
  inference(forward_demodulation,[],[f1508,f702])).
fof(f1508,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X15,bool),sK28(X17,X16,'type.pair.prod'(X15,bool))))) | s(bool,'F_0') = s(bool,sK29('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X15,bool),X16),X17)),bool,X15))) ) | $spl1),
  inference(superposition,[],[f1486,f701])).
fof(f1520,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'F_0') != s(bool,sK29(X13,bool,X11)) | s(bool,'F_0') = s(bool,sK29('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X11,bool),X12),'const.pair.,_2'(s('type.pair.prod'(X11,bool),X13),s(X12,X14)))),bool,X11))) ) | $spl1),
  inference(forward_demodulation,[],[f1507,f702])).
fof(f1507,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X11,bool),X13))) | s(bool,'F_0') = s(bool,sK29('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X11,bool),X12),'const.pair.,_2'(s('type.pair.prod'(X11,bool),X13),s(X12,X14)))),bool,X11))) ) | $spl1),
  inference(superposition,[],[f1486,f323])).
fof(f1519,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') != s(bool,sK29(sK29(X10,'type.pair.prod'(X8,bool),X9),bool,X8)) | s(bool,'F_0') = s(bool,sK29('const.pair.SND_1'(s('type.pair.prod'(X9,'type.pair.prod'(X8,bool)),X10)),bool,X8))) ) | $spl1),
  inference(forward_demodulation,[],[f1506,f702])).
fof(f1506,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X8,bool),sK29(X10,'type.pair.prod'(X8,bool),X9)))) | s(bool,'F_0') = s(bool,sK29('const.pair.SND_1'(s('type.pair.prod'(X9,'type.pair.prod'(X8,bool)),X10)),bool,X8))) ) | $spl1),
  inference(superposition,[],[f1486,f702])).
fof(f1518,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,sK29(X7,bool,X4)) | s(bool,'F_0') = s(bool,sK29('const.pair.SND_1'(s('type.pair.prod'(X5,'type.pair.prod'(X4,bool)),'const.pair.,_2'(s(X5,X6),s('type.pair.prod'(X4,bool),X7)))),bool,X4))) ) | $spl1),
  inference(forward_demodulation,[],[f1505,f702])).
fof(f1505,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),X7))) | s(bool,'F_0') = s(bool,sK29('const.pair.SND_1'(s('type.pair.prod'(X5,'type.pair.prod'(X4,bool)),'const.pair.,_2'(s(X5,X6),s('type.pair.prod'(X4,bool),X7)))),bool,X4))) ) | $spl1),
  inference(superposition,[],[f1486,f322])).
fof(f1517,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,sK29(X3,bool,X2)) | s(bool,'F_0') = s(bool,sK29(happ(s(fun('type.pair.prod'(X2,bool),'type.pair.prod'(X2,bool)),'_i_0_0'),s('type.pair.prod'(X2,bool),X3)),bool,X2))) ) | $spl1),
  inference(forward_demodulation,[],[f1504,f702])).
fof(f1504,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),X3))) | s(bool,'F_0') = s(bool,sK29(happ(s(fun('type.pair.prod'(X2,bool),'type.pair.prod'(X2,bool)),'_i_0_0'),s('type.pair.prod'(X2,bool),X3)),bool,X2))) ) | $spl1),
  inference(superposition,[],[f1486,f279])).
fof(f1516,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,sK29(X1,bool,X0)) | s(bool,'F_0') = s(bool,sK29('const.pair.,_2'(s(X0,sK28(X1,bool,X0)),s(bool,sK29(X1,bool,X0))),bool,X0))) ) | $spl1),
  inference(forward_demodulation,[],[f1503,f702])).
fof(f1503,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),X1))) | s(bool,'F_0') = s(bool,sK29('const.pair.,_2'(s(X0,sK28(X1,bool,X0)),s(bool,sK29(X1,bool,X0))),bool,X0))) ) | $spl1),
  inference(superposition,[],[f1486,f321])).
fof(f1489,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,sK28(X30,X31,bool)) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X28,bool),X29))) | s(bool,'F_0') = s(bool,sK28(X30,X31,bool)) | s(bool,'F_0') = s(bool,sK29(X29,bool,X28))) ) | $spl1),
  inference(subsumption_resolution,[],[f1466,f261])).
fof(f1466,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'T_0')) | s(bool,sK28(X30,X31,bool)) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X28,bool),X29))) | s(bool,'F_0') = s(bool,sK28(X30,X31,bool)) | s(bool,'F_0') = s(bool,sK29(X29,bool,X28))) ) | $spl1),
  inference(superposition,[],[f920,f1279])).
fof(f1488,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,sK29(X17,bool,X16)) | s(bool,'F_0') = s(bool,sK29(X17,bool,X16))) ) | $spl1),
  inference(forward_demodulation,[],[f1487,f702])).
fof(f1487,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X16,bool),X17))) | s(bool,'F_0') = s(bool,sK29(X17,bool,X16))) ) | $spl1),
  inference(subsumption_resolution,[],[f1462,f261])).
fof(f1462,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | s(bool,X18) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X16,bool),X17))) | s(bool,'F_0') = s(bool,sK29(X17,bool,X16))) ) | $spl1),
  inference(superposition,[],[f299,f1279])).
fof(f1486,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X6,bool),X7))) | s(bool,'F_0') = s(bool,sK29(X7,bool,X6))) ) | $spl1),
  inference(subsumption_resolution,[],[f1458,f261])).
fof(f1458,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X6,bool),X7))) | s(bool,'F_0') = s(bool,sK29(X7,bool,X6))) ) | $spl1),
  inference(superposition,[],[f269,f1279])).
fof(f1453,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,sK29('const.pair.,_2'(s(X6,X7),s(bool,X8)),bool,X6)) | s(bool,'T_0') = s(bool,X8)) ) | $spl1),
  inference(superposition,[],[f1279,f322])).
fof(f1484,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,sK29(sK28(X17,X16,'type.pair.prod'(X15,bool)),bool,X15)) | s(bool,'F_0') = s(bool,sK29('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X15,bool),X16),X17)),bool,X15))) ) | $spl1),
  inference(forward_demodulation,[],[f1450,f702])).
fof(f1450,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X15,bool),sK28(X17,X16,'type.pair.prod'(X15,bool))))) | s(bool,'F_0') = s(bool,sK29('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X15,bool),X16),X17)),bool,X15))) ) | $spl1),
  inference(superposition,[],[f1279,f701])).
fof(f1483,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'T_0') = s(bool,sK29(X13,bool,X11)) | s(bool,'F_0') = s(bool,sK29('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X11,bool),X12),'const.pair.,_2'(s('type.pair.prod'(X11,bool),X13),s(X12,X14)))),bool,X11))) ) | $spl1),
  inference(forward_demodulation,[],[f1449,f702])).
fof(f1449,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X11,bool),X13))) | s(bool,'F_0') = s(bool,sK29('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X11,bool),X12),'const.pair.,_2'(s('type.pair.prod'(X11,bool),X13),s(X12,X14)))),bool,X11))) ) | $spl1),
  inference(superposition,[],[f1279,f323])).
fof(f1482,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,sK29(sK29(X10,'type.pair.prod'(X8,bool),X9),bool,X8)) | s(bool,'F_0') = s(bool,sK29('const.pair.SND_1'(s('type.pair.prod'(X9,'type.pair.prod'(X8,bool)),X10)),bool,X8))) ) | $spl1),
  inference(forward_demodulation,[],[f1448,f702])).
fof(f1448,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X8,bool),sK29(X10,'type.pair.prod'(X8,bool),X9)))) | s(bool,'F_0') = s(bool,sK29('const.pair.SND_1'(s('type.pair.prod'(X9,'type.pair.prod'(X8,bool)),X10)),bool,X8))) ) | $spl1),
  inference(superposition,[],[f1279,f702])).
fof(f1481,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,sK29(X7,bool,X4)) | s(bool,'F_0') = s(bool,sK29('const.pair.SND_1'(s('type.pair.prod'(X5,'type.pair.prod'(X4,bool)),'const.pair.,_2'(s(X5,X6),s('type.pair.prod'(X4,bool),X7)))),bool,X4))) ) | $spl1),
  inference(forward_demodulation,[],[f1447,f702])).
fof(f1447,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),X7))) | s(bool,'F_0') = s(bool,sK29('const.pair.SND_1'(s('type.pair.prod'(X5,'type.pair.prod'(X4,bool)),'const.pair.,_2'(s(X5,X6),s('type.pair.prod'(X4,bool),X7)))),bool,X4))) ) | $spl1),
  inference(superposition,[],[f1279,f322])).
fof(f1480,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,sK29(X3,bool,X2)) | s(bool,'F_0') = s(bool,sK29(happ(s(fun('type.pair.prod'(X2,bool),'type.pair.prod'(X2,bool)),'_i_0_0'),s('type.pair.prod'(X2,bool),X3)),bool,X2))) ) | $spl1),
  inference(forward_demodulation,[],[f1446,f702])).
fof(f1446,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),X3))) | s(bool,'F_0') = s(bool,sK29(happ(s(fun('type.pair.prod'(X2,bool),'type.pair.prod'(X2,bool)),'_i_0_0'),s('type.pair.prod'(X2,bool),X3)),bool,X2))) ) | $spl1),
  inference(superposition,[],[f1279,f279])).
fof(f1479,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK29(X1,bool,X0)) | s(bool,'F_0') = s(bool,sK29('const.pair.,_2'(s(X0,sK28(X1,bool,X0)),s(bool,sK29(X1,bool,X0))),bool,X0))) ) | $spl1),
  inference(forward_demodulation,[],[f1445,f702])).
fof(f1445,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),X1))) | s(bool,'F_0') = s(bool,sK29('const.pair.,_2'(s(X0,sK28(X1,bool,X0)),s(bool,sK29(X1,bool,X0))),bool,X0))) ) | $spl1),
  inference(superposition,[],[f1279,f321])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK7(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK7(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X7] : ! [X8,X10] : (! [X9] : (s(X7,X9) = s(X7,X10) => p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9))))) => p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X8)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.bool.JRH_INDUCT_UTIL')).
fof(f569,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK7(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f292,f279])).
fof(f571,plain,(
  ( ! [X12,X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X10)))) | s(X8,X12) = s(X8,sK7(X12,'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),'const.pair.,_2'(s(fun(X8,bool),X10),s(X9,X11)))),X8))) )),
  inference(superposition,[],[f292,f323])).
fof(f823,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,'|dollar|exists_1'(s(fun(X62,bool),sK28(X64,X63,fun(X62,bool)))))) | s(X62,X65) = s(X62,sK7(X65,'const.pair.FST_1'(s('type.pair.prod'(fun(X62,bool),X63),X64)),X62))) )),
  inference(superposition,[],[f292,f701])).
fof(f570,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6)))) | s(X3,X7) = s(X3,sK7(X7,'const.pair.SND_1'(s('type.pair.prod'(X4,fun(X3,bool)),'const.pair.,_2'(s(X4,X5),s(fun(X3,bool),X6)))),X3))) )),
  inference(superposition,[],[f292,f322])).
fof(f1265,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,'|dollar|exists_1'(s(fun(X97,bool),sK29(X99,fun(X97,bool),X98))))) | s(X97,X100) = s(X97,sK7(X100,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X97,bool)),X99)),X97))) )),
  inference(superposition,[],[f292,f702])).
fof(f1122,plain,(
  ( ! [X0] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0)))))) | p(s(bool,'const.min.@_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))))))) )),
  inference(superposition,[],[f305,f280])).
fof(f968,plain,(
  ( ! [X0] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0)))))) | p(s(bool,X0))) )),
  inference(superposition,[],[f304,f280])).
fof(f658,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f652,f279])).
fof(f652,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f288,f279])).
fof(f660,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X8)))) | p(s(bool,happ(s(fun(X6,bool),X8),s(X6,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),'const.pair.,_2'(s(fun(X6,bool),X8),s(X7,X9)))),X6)))))) )),
  inference(forward_demodulation,[],[f654,f323])).
fof(f654,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),'const.pair.,_2'(s(fun(X6,bool),X8),s(X7,X9)))),X6))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),'const.pair.,_2'(s(fun(X6,bool),X8),s(X7,X9))))))))) )),
  inference(superposition,[],[f288,f323])).
fof(f659,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,'|dollar|exists_1'(s(fun(X2,bool),X5)))) | p(s(bool,happ(s(fun(X2,bool),X5),s(X2,sK6('const.pair.SND_1'(s('type.pair.prod'(X3,fun(X2,bool)),'const.pair.,_2'(s(X3,X4),s(fun(X2,bool),X5)))),X2)))))) )),
  inference(forward_demodulation,[],[f653,f322])).
fof(f653,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun(X2,bool),X5),s(X2,sK6('const.pair.SND_1'(s('type.pair.prod'(X3,fun(X2,bool)),'const.pair.,_2'(s(X3,X4),s(fun(X2,bool),X5)))),X2))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X2,bool),'const.pair.SND_1'(s('type.pair.prod'(X3,fun(X2,bool)),'const.pair.,_2'(s(X3,X4),s(fun(X2,bool),X5))))))))) )),
  inference(superposition,[],[f288,f322])).
fof(f374,plain,(
  ( ! [X4,X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),sK25(X2,X1,X0)))))) )),
  inference(backward_demodulation,[],[f337,f318])).
fof(f373,plain,(
  ( ! [X6,X2,X0,X7,X1] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),sK15(X2,X1,X0))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7))))) )),
  inference(backward_demodulation,[],[f337,f311])).
fof(f314,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK18(X2,X1,X0)))),s(X1,sK19(X2,X1,X0))))) | p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))))) & (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))) | ! [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19])],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((? [X5,X6] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))))) & (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))) | ! [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))))),
  inference(rectify,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((? [X3,X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))))) & (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))) | ! [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))))),
  inference(nnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : (! [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) <=> p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ! [X2] : (! [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) <=> p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X7,X6] : ! [X8] : (! [X9,X14] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,X9))),s(X6,X14)))) <=> p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X7,X6),bool),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,bool)),X8)))))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.PFORALL_THM')).
fof(f372,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),X3)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK16(X2,X1,X0)))),s(X1,sK17(X2,X1,X0)))))) )),
  inference(backward_demodulation,[],[f337,f312])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK18(X2,X1,X0)))),s(X1,sK19(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f371,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),X8)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK23(X2,X1,X0)))),s(X1,sK24(X2,X1,X0)))))) )),
  inference(backward_demodulation,[],[f337,f317])).
fof(f320,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK26(X2,X1,X0)))),s(X1,sK27(X2,X1,X0))))) | ~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f327,plain,(
  ( ! [X6,X0,X5,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),sK32(X1,X0)),s(X0,X5))),s(X1,X6)))) | s(X0,X5) = s(X0,sK33(X1,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1] : ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),sK32(X1,X0)),s(X0,X5))),s(X1,X6)))) | (s(X0,X5) = s(X0,sK33(X1,X0)) & s(X1,X6) = s(X1,sK34(X1,X0)))) & (s(X0,X5) != s(X0,sK33(X1,X0)) | s(X1,X6) != s(X1,sK34(X1,X0)) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),sK32(X1,X0)),s(X0,X5))),s(X1,X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK32,sK33,sK34])],[f253])).
fof(f253,plain,(
  ! [X0,X1] : ? [X2,X3,X4] : ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | (s(X0,X3) = s(X0,X5) & s(X1,X4) = s(X1,X6))) & (s(X0,X3) != s(X0,X5) | s(X1,X4) != s(X1,X6) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6))))))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0,X1] : ? [X2,X3,X4] : ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | (s(X0,X3) = s(X0,X5) & s(X1,X4) = s(X1,X6))) & ((s(X0,X3) != s(X0,X5) | s(X1,X4) != s(X1,X6)) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6))))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : ? [X2,X3,X4] : ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) <=> (s(X0,X3) = s(X0,X5) & s(X1,X4) = s(X1,X6)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X7,X6] : ? [X16,X9,X14] : ! [X5,X45] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X16),s(X7,X5))),s(X6,X45)))) <=> (s(X7,X5) = s(X7,X9) & s(X6,X14) = s(X6,X45)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.pair_|48|_prove_ltb')).
fof(f328,plain,(
  ( ! [X6,X0,X5,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),sK32(X1,X0)),s(X0,X5))),s(X1,X6)))) | s(X1,X6) = s(X1,sK34(X1,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f329,plain,(
  ( ! [X6,X0,X5,X1] : (s(X1,X6) != s(X1,sK34(X1,X0)) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),sK32(X1,X0)),s(X0,X5))),s(X1,X6)))) | s(X0,X5) != s(X0,sK33(X1,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X39,X40,X41] : (p(s(bool,happ(s(fun(X39,bool),happ(s(fun(X39,fun(X39,bool)),'=_0'),s(X39,X40))),s(X39,X41)))) <=> s(X39,X40) = s(X39,X41))),
  file('thm.pair.ELIM_PEXISTS_EVAL',arity_0_HL_EQ)).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f255])).
fof(f1209,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(fun(X0,bool),bool),'!_1')))) | p(s(bool,happ(s(fun(X0,bool),'const.min.@_1'(s(fun(fun(X0,bool),bool),'!_1'))),s(X0,X1))))) )),
  inference(superposition,[],[f300,f280])).
fof(f841,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(X55,bool),sK28(X57,X56,fun(X55,bool)))))) | p(s(bool,happ(s(fun(X55,bool),sK28(X57,X56,fun(X55,bool))),s(X55,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X55,bool),X56),X57)),X55)))))) )),
  inference(forward_demodulation,[],[f821,f701])).
fof(f821,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,happ(s(fun(X55,bool),sK28(X57,X56,fun(X55,bool))),s(X55,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X55,bool),X56),X57)),X55))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X55,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X55,bool),X56),X57))))))) )),
  inference(superposition,[],[f288,f701])).
fof(f1295,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'|dollar|exists_1'(s(fun(X90,bool),sK29(X92,fun(X90,bool),X91))))) | p(s(bool,happ(s(fun(X90,bool),sK29(X92,fun(X90,bool),X91)),s(X90,sK6('const.pair.SND_1'(s('type.pair.prod'(X91,fun(X90,bool)),X92)),X90)))))) )),
  inference(forward_demodulation,[],[f1263,f702])).
fof(f1263,plain,(
  ( ! [X92,X90,X91] : (p(s(bool,happ(s(fun(X90,bool),sK29(X92,fun(X90,bool),X91)),s(X90,sK6('const.pair.SND_1'(s('type.pair.prod'(X91,fun(X90,bool)),X92)),X90))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X90,bool),'const.pair.SND_1'(s('type.pair.prod'(X91,fun(X90,bool)),X92))))))) )),
  inference(superposition,[],[f288,f702])).
fof(f1127,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X5))),s(bool,X4)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1109,f263])).
fof(f1109,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X5))),s(bool,X4)))) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f305,f264])).
fof(f973,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X4))),s(bool,X5)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f942,f263])).
fof(f942,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X4))),s(bool,X5)))) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f304,f264])).
fof(f304,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) & p(s(bool,X1)))) & (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) & p(s(bool,X1)))) & ((~p(s(bool,X0)) | ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) & p(s(bool,X1))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X40,X41] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X40))),s(bool,X41)))) <=> (p(s(bool,X40)) & p(s(bool,X41))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','|slash||bslash|_HL_0')).
fof(f305,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f228])).
fof(f306,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f939,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f263,f304])).
fof(f1091,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,'F_0'))))) )),
  inference(unit_resulting_resolution,[],[f263,f305])).
fof(f375,plain,(
  ( ! [X4,X2,X0,X5,X1] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,X4),s(X1,X5)))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK22(X2,X1,X0)))))) )),
  inference(backward_demodulation,[],[f343,f316])).
fof(f376,plain,(
  ( ! [X6,X2,X0,X7,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK12(X2,X1,X0))))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,X6),s(X1,X7))))) )),
  inference(backward_demodulation,[],[f343,f309])).
fof(f377,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,sK10(X2,X1,X0)),s(X1,sK11(X2,X1,X0))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(backward_demodulation,[],[f343,f308])).
fof(f378,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,sK13(X2,X1,X0)),s(X1,sK14(X2,X1,X0))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(backward_demodulation,[],[f343,f310])).
fof(f379,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X8)))) | p(s(bool,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X0,X1),bool),X2),s(X0,sK20(X2,X1,X0)),s(X1,sK21(X2,X1,X0)))))) )),
  inference(backward_demodulation,[],[f343,f315])).
fof(f326,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(fun(X0,fun(X1,bool)),bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'?_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,bool)),'_dst_x00x11_2'(s(fun(fun(X1,bool),bool),'?_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,bool))),'_dst_x000x111_2'(s(fun(fun(X0,bool),bool),'!_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,fun(X0,bool)))),'_dst_x0000x1111_2'(s(fun(fun(X1,bool),bool),'!_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,fun(X0,fun(X1,bool))))),'_dst_x00000x10011x01111_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,bool))),'_i_0_0'),s(fun(X0,fun(X1,fun(X0,fun(X1,bool)))),'_dst_x0000x1010x0101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,fun(X0,bool)),'_dst_x01x10_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'))),s(fun(X1,fun(X1,bool)),'_dst_x01x10_2'(s(fun(X1,fun(X1,bool)),'=_0'),s(fun(X1,X1),'_i_0_0'))))))))))))))),s(fun('type.pair.prod'(X0,X1),fun(X0,fun(X1,bool))),sK31(X1,X0)))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(fun(X0,fun(X1,bool)),bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'?_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,bool)),'_dst_x00x11_2'(s(fun(fun(X1,bool),bool),'?_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,bool))),'_dst_x000x111_2'(s(fun(fun(X0,bool),bool),'!_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,fun(X0,bool)))),'_dst_x0000x1111_2'(s(fun(fun(X1,bool),bool),'!_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,fun(X0,fun(X1,bool))))),'_dst_x00000x10011x01111_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,bool))),'_i_0_0'),s(fun(X0,fun(X1,fun(X0,fun(X1,bool)))),'_dst_x0000x1010x0101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,fun(X0,bool)),'_dst_x01x10_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'))),s(fun(X1,fun(X1,bool)),'_dst_x01x10_2'(s(fun(X1,fun(X1,bool)),'=_0'),s(fun(X1,X1),'_i_0_0'))))))))))))))),s(fun('type.pair.prod'(X0,X1),fun(X0,fun(X1,bool))),sK31(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31])],[f127])).
fof(f127,plain,(
  ! [X0,X1] : ? [X2] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(fun(X0,fun(X1,bool)),bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'?_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,bool)),'_dst_x00x11_2'(s(fun(fun(X1,bool),bool),'?_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,bool))),'_dst_x000x111_2'(s(fun(fun(X0,bool),bool),'!_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,fun(X0,bool)))),'_dst_x0000x1111_2'(s(fun(fun(X1,bool),bool),'!_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,fun(X0,fun(X1,bool))))),'_dst_x00000x10011x01111_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(X0,fun(X1,bool)),fun(X0,fun(X1,bool))),'_i_0_0'),s(fun(X0,fun(X1,fun(X0,fun(X1,bool)))),'_dst_x0000x1010x0101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,fun(X0,bool)),'_dst_x01x10_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'))),s(fun(X1,fun(X1,bool)),'_dst_x01x10_2'(s(fun(X1,fun(X1,bool)),'=_0'),s(fun(X1,X1),'_i_0_0'))))))))))))))),s(fun('type.pair.prod'(X0,X1),fun(X0,fun(X1,bool))),X2))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X7,X6] : ? [X28] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun(fun(X7,fun(X6,bool)),bool),'_dst_x0x1_2'(s(fun(fun(X7,bool),bool),'?_0'),s(fun(fun(X7,fun(X6,bool)),fun(X7,bool)),'_dst_x00x11_2'(s(fun(fun(X6,bool),bool),'?_0'),s(fun(fun(X7,fun(X6,bool)),fun(X7,fun(X6,bool))),'_dst_x000x111_2'(s(fun(fun(X7,bool),bool),'!_0'),s(fun(fun(X7,fun(X6,bool)),fun(X7,fun(X6,fun(X7,bool)))),'_dst_x0000x1111_2'(s(fun(fun(X6,bool),bool),'!_0'),s(fun(fun(X7,fun(X6,bool)),fun(X7,fun(X6,fun(X7,fun(X6,bool))))),'_dst_x00000x10011x01111_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(X7,fun(X6,bool)),fun(X7,fun(X6,bool))),'_i_0_0'),s(fun(X7,fun(X6,fun(X7,fun(X6,bool)))),'_dst_x0000x1010x0101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X7,fun(X7,bool)),'_dst_x01x10_2'(s(fun(X7,fun(X7,bool)),'=_0'),s(fun(X7,X7),'_i_0_0'))),s(fun(X6,fun(X6,bool)),'_dst_x01x10_2'(s(fun(X6,fun(X6,bool)),'=_0'),s(fun(X6,X6),'_i_0_0'))))))))))))))),s(fun('type.pair.prod'(X7,X6),fun(X7,fun(X6,bool))),X28))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.prod_TY_DEF')).
fof(f1326,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(superposition,[],[f303,f280])).
fof(f1300,plain,(
  ( ! [X111,X112,X110,X113] : (s(X110,sK28(sK29(X113,'type.pair.prod'(X110,X111),X112),X111,X110)) = s(X110,sK28('const.pair.SND_1'(s('type.pair.prod'(X112,'type.pair.prod'(X110,X111)),X113)),X111,X110))) )),
  inference(forward_demodulation,[],[f1269,f701])).
fof(f1269,plain,(
  ( ! [X111,X112,X110,X113] : (s(X110,'const.pair.FST_1'(s('type.pair.prod'(X110,X111),sK29(X113,'type.pair.prod'(X110,X111),X112)))) = s(X110,sK28('const.pair.SND_1'(s('type.pair.prod'(X112,'type.pair.prod'(X110,X111)),X113)),X111,X110))) )),
  inference(superposition,[],[f701,f702])).
fof(f1299,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') = s(bool,sK28(sK29(X109,'type.pair.prod'(bool,X107),X108),X107,bool)) | s(bool,'F_0') = s(bool,sK28('const.pair.SND_1'(s('type.pair.prod'(X108,'type.pair.prod'(bool,X107)),X109)),X107,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f1268,f701])).
fof(f1268,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X107),sK29(X109,'type.pair.prod'(bool,X107),X108)))) | s(bool,'F_0') = s(bool,sK28('const.pair.SND_1'(s('type.pair.prod'(X108,'type.pair.prod'(bool,X107)),X109)),X107,bool))) ) | $spl1),
  inference(superposition,[],[f831,f702])).
fof(f1298,plain,(
  ( ! [X105,X106,X104] : (s(bool,'F_0') != s(bool,sK28(sK29(X106,'type.pair.prod'(bool,X104),X105),X104,bool)) | s(bool,'F_0') = s(bool,sK28('const.pair.SND_1'(s('type.pair.prod'(X105,'type.pair.prod'(bool,X104)),X106)),X104,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f1267,f701])).
fof(f1267,plain,(
  ( ! [X105,X106,X104] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X104),sK29(X106,'type.pair.prod'(bool,X104),X105)))) | s(bool,'F_0') = s(bool,sK28('const.pair.SND_1'(s('type.pair.prod'(X105,'type.pair.prod'(bool,X104)),X106)),X104,bool))) ) | $spl1),
  inference(superposition,[],[f918,f702])).
fof(f1279,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X7,bool),X8))) | s(bool,'F_0') = s(bool,sK29(X8,bool,X7))) ) | $spl1),
  inference(forward_demodulation,[],[f1234,f391])).
fof(f1234,plain,(
  ( ! [X8,X7] : (s(bool,sK29(X8,bool,X7)) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X7,bool),X8)))) ) | $spl1),
  inference(superposition,[],[f702,f595])).
fof(f1278,plain,(
  ( ! [X43,X44,X42] : (s('type.pair.prod'(bool,X43),'const.pair.,_2'(s(bool,'F_0'),s(X43,X44))) = s('type.pair.prod'(bool,X43),'const.pair.,_2'(s(bool,X42),s(X43,X44))) | s(bool,'T_0') = s(bool,X42)) ) | $spl1),
  inference(backward_demodulation,[],[f1233,f1030])).
fof(f1233,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,sK29('const.pair.,_2'(s(X1,X2),s(X0,X3)),X0,X1))) )),
  inference(superposition,[],[f702,f322])).
fof(f1277,plain,(
  ( ! [X24,X23,X21,X22] : (s(X22,sK29(sK28(X24,X23,'type.pair.prod'(X21,X22)),X22,X21)) = s(X22,sK29('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X21,X22),X23),X24)),X22,X21))) )),
  inference(forward_demodulation,[],[f1232,f702])).
fof(f1232,plain,(
  ( ! [X24,X23,X21,X22] : (s(X22,'const.pair.SND_1'(s('type.pair.prod'(X21,X22),sK28(X24,X23,'type.pair.prod'(X21,X22))))) = s(X22,sK29('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X21,X22),X23),X24)),X22,X21))) )),
  inference(superposition,[],[f702,f701])).
fof(f1276,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(X17,sK29(X19,X17,X16)) = s(X17,sK29('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X16,X17),X18),'const.pair.,_2'(s('type.pair.prod'(X16,X17),X19),s(X18,X20)))),X17,X16))) )),
  inference(forward_demodulation,[],[f1231,f702])).
fof(f1231,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(X17,'const.pair.SND_1'(s('type.pair.prod'(X16,X17),X19))) = s(X17,sK29('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X16,X17),X18),'const.pair.,_2'(s('type.pair.prod'(X16,X17),X19),s(X18,X20)))),X17,X16))) )),
  inference(superposition,[],[f702,f323])).
fof(f1275,plain,(
  ( ! [X14,X12,X15,X13] : (s(X13,sK29(sK29(X15,'type.pair.prod'(X12,X13),X14),X13,X12)) = s(X13,sK29('const.pair.SND_1'(s('type.pair.prod'(X14,'type.pair.prod'(X12,X13)),X15)),X13,X12))) )),
  inference(forward_demodulation,[],[f1230,f702])).
fof(f1230,plain,(
  ( ! [X14,X12,X15,X13] : (s(X13,'const.pair.SND_1'(s('type.pair.prod'(X12,X13),sK29(X15,'type.pair.prod'(X12,X13),X14)))) = s(X13,sK29('const.pair.SND_1'(s('type.pair.prod'(X14,'type.pair.prod'(X12,X13)),X15)),X13,X12))) )),
  inference(superposition,[],[f702,f702])).
fof(f1274,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(X8,sK29(X11,X8,X7)) = s(X8,sK29('const.pair.SND_1'(s('type.pair.prod'(X9,'type.pair.prod'(X7,X8)),'const.pair.,_2'(s(X9,X10),s('type.pair.prod'(X7,X8),X11)))),X8,X7))) )),
  inference(forward_demodulation,[],[f1229,f702])).
fof(f1229,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(X8,'const.pair.SND_1'(s('type.pair.prod'(X7,X8),X11))) = s(X8,sK29('const.pair.SND_1'(s('type.pair.prod'(X9,'type.pair.prod'(X7,X8)),'const.pair.,_2'(s(X9,X10),s('type.pair.prod'(X7,X8),X11)))),X8,X7))) )),
  inference(superposition,[],[f702,f322])).
fof(f1273,plain,(
  ( ! [X6,X4,X5] : (s(X5,sK29(X6,X5,X4)) = s(X5,sK29(happ(s(fun('type.pair.prod'(X4,X5),'type.pair.prod'(X4,X5)),'_i_0_0'),s('type.pair.prod'(X4,X5),X6)),X5,X4))) )),
  inference(forward_demodulation,[],[f1228,f702])).
fof(f1228,plain,(
  ( ! [X6,X4,X5] : (s(X5,'const.pair.SND_1'(s('type.pair.prod'(X4,X5),X6))) = s(X5,sK29(happ(s(fun('type.pair.prod'(X4,X5),'type.pair.prod'(X4,X5)),'_i_0_0'),s('type.pair.prod'(X4,X5),X6)),X5,X4))) )),
  inference(superposition,[],[f702,f279])).
fof(f1272,plain,(
  ( ! [X2,X3,X1] : (s(X2,sK29(X3,X2,X1)) = s(X2,sK29('const.pair.,_2'(s(X1,sK28(X3,X2,X1)),s(X2,sK29(X3,X2,X1))),X2,X1))) )),
  inference(forward_demodulation,[],[f1227,f702])).
fof(f1227,plain,(
  ( ! [X2,X3,X1] : (s(X2,'const.pair.SND_1'(s('type.pair.prod'(X1,X2),X3))) = s(X2,sK29('const.pair.,_2'(s(X1,sK28(X3,X2,X1)),s(X2,sK29(X3,X2,X1))),X2,X1))) )),
  inference(superposition,[],[f702,f321])).
fof(f1270,plain,(
  ( ! [X0] : (s(sK0,sK29(X0,sK0,sK1)) = s(sK0,sK29(sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)),sK0,sK1))) ) | $spl1),
  inference(forward_demodulation,[],[f1226,f702])).
fof(f1226,plain,(
  ( ! [X0] : (s(sK0,'const.pair.SND_1'(s('type.pair.prod'(sK1,sK0),X0))) = s(sK0,sK29(sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)),sK0,sK1))) ) | $spl1),
  inference(superposition,[],[f702,f575])).
fof(f1225,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK1,sK0),X0) = s('type.pair.prod'(sK1,sK0),'const.pair.,_2'(s(sK1,sK28(X0,sK0,sK1)),s(sK0,sK29(sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)),sK0,sK1))))) ) | $spl1),
  inference(forward_demodulation,[],[f1217,f575])).
fof(f1217,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK1,sK0),sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))) = s('type.pair.prod'(sK1,sK0),'const.pair.,_2'(s(sK1,sK28(X0,sK0,sK1)),s(sK0,sK29(sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)),sK0,sK1))))) ) | $spl1),
  inference(superposition,[],[f321,f825])).
fof(f1196,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,sK28(X24,X25,bool)) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),X23))) | s(bool,'F_0') = s(bool,sK28(X24,X25,bool)) | s(bool,'F_0') = s(bool,sK28(X23,X22,bool))) ) | $spl1),
  inference(subsumption_resolution,[],[f1181,f261])).
fof(f1181,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'T_0')) | s(bool,sK28(X24,X25,bool)) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),X23))) | s(bool,'F_0') = s(bool,sK28(X24,X25,bool)) | s(bool,'F_0') = s(bool,sK28(X23,X22,bool))) ) | $spl1),
  inference(superposition,[],[f920,f831])).
fof(f1056,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,sK28('const.pair.,_2'(s(bool,X1),s(X2,X3)),X2,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f930,f264])).
fof(f1030,plain,(
  ( ! [X43,X44,X42] : (s('type.pair.prod'(bool,X43),'const.pair.,_2'(s(bool,X42),s(X43,X44))) = s('type.pair.prod'(bool,X43),'const.pair.,_2'(s(bool,'F_0'),s(X43,sK29('const.pair.,_2'(s(bool,X42),s(X43,X44)),X43,bool)))) | s(bool,'T_0') = s(bool,X42)) ) | $spl1),
  inference(superposition,[],[f321,f891])).
fof(f1051,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X38) = s(bool,sK28('const.pair.,_2'(s(bool,X35),s(X36,X37)),X36,bool)) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X35)) ) | $spl1),
  inference(subsumption_resolution,[],[f1028,f263])).
fof(f1028,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,'F_0')) | s(bool,X38) = s(bool,sK28('const.pair.,_2'(s(bool,X35),s(X36,X37)),X36,bool)) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X35)) ) | $spl1),
  inference(superposition,[],[f468,f891])).
fof(f1003,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,sK28('const.pair.,_2'(s(bool,X5),s(bool,X4)),bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(superposition,[],[f891,f264])).
fof(f999,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'F_0') = s(bool,sK28('const.pair.,_2'(s(bool,X42),s(X44,X45)),X44,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),'const.pair.,_2'(s(bool,X42),s(X41,X43)))))) ) | $spl1),
  inference(superposition,[],[f891,f323])).
fof(f998,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') = s(bool,sK28('const.pair.,_2'(s(bool,X38),s(X39,X40)),X39,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,X38)))))) ) | $spl1),
  inference(superposition,[],[f891,f322])).
fof(f993,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'F_0') = s(bool,sK28('const.pair.,_2'(s(bool,'|dollar|exists_1'(s(fun(X18,bool),X19))),s(X20,X21)),X20,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(X18,bool),X19),s(X18,'const.min.@_1'(s(fun(X18,bool),X19)))))) ) | $spl1),
  inference(superposition,[],[f891,f280])).
fof(f930,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'F_0') = s(bool,sK28('const.pair.,_2'(s(bool,X7),s(X6,X8)),X6,bool))) ) | $spl1),
  inference(superposition,[],[f918,f323])).
fof(f938,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,sK28(sK28(X14,X13,'type.pair.prod'(bool,X12)),X12,bool)) | s(bool,'F_0') = s(bool,sK28('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(bool,X12),X13),X14)),X12,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f926,f701])).
fof(f926,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X12),sK28(X14,X13,'type.pair.prod'(bool,X12))))) | s(bool,'F_0') = s(bool,sK28('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(bool,X12),X13),X14)),X12,bool))) ) | $spl1),
  inference(superposition,[],[f918,f701])).
fof(f937,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,sK28(X10,X8,bool)) | s(bool,'F_0') = s(bool,sK28('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(bool,X8),X9),'const.pair.,_2'(s('type.pair.prod'(bool,X8),X10),s(X9,X11)))),X8,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f925,f701])).
fof(f925,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X8),X10))) | s(bool,'F_0') = s(bool,sK28('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(bool,X8),X9),'const.pair.,_2'(s('type.pair.prod'(bool,X8),X10),s(X9,X11)))),X8,bool))) ) | $spl1),
  inference(superposition,[],[f918,f323])).
fof(f936,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,sK28(X7,X4,bool)) | s(bool,'F_0') = s(bool,sK28('const.pair.SND_1'(s('type.pair.prod'(X5,'type.pair.prod'(bool,X4)),'const.pair.,_2'(s(X5,X6),s('type.pair.prod'(bool,X4),X7)))),X4,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f924,f701])).
fof(f924,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),X7))) | s(bool,'F_0') = s(bool,sK28('const.pair.SND_1'(s('type.pair.prod'(X5,'type.pair.prod'(bool,X4)),'const.pair.,_2'(s(X5,X6),s('type.pair.prod'(bool,X4),X7)))),X4,bool))) ) | $spl1),
  inference(superposition,[],[f918,f322])).
fof(f935,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,sK28(X3,X2,bool)) | s(bool,'F_0') = s(bool,sK28(happ(s(fun('type.pair.prod'(bool,X2),'type.pair.prod'(bool,X2)),'_i_0_0'),s('type.pair.prod'(bool,X2),X3)),X2,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f923,f701])).
fof(f923,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'F_0') = s(bool,sK28(happ(s(fun('type.pair.prod'(bool,X2),'type.pair.prod'(bool,X2)),'_i_0_0'),s('type.pair.prod'(bool,X2),X3)),X2,bool))) ) | $spl1),
  inference(superposition,[],[f918,f279])).
fof(f934,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,sK28(X1,X0,bool)) | s(bool,'F_0') = s(bool,sK28('const.pair.,_2'(s(bool,sK28(X1,X0,bool)),s(X0,sK29(X1,X0,bool))),X0,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f922,f701])).
fof(f922,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) | s(bool,'F_0') = s(bool,sK28('const.pair.,_2'(s(bool,sK28(X1,X0,bool)),s(X0,sK29(X1,X0,bool))),X0,bool))) ) | $spl1),
  inference(superposition,[],[f918,f321])).
fof(f920,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,sK28(X17,X16,bool)) | s(bool,'F_0') = s(bool,sK28(X17,X16,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f919,f701])).
fof(f919,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X16),X17))) | s(bool,'F_0') = s(bool,sK28(X17,X16,bool))) ) | $spl1),
  inference(subsumption_resolution,[],[f900,f261])).
fof(f900,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | s(bool,X18) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X16),X17))) | s(bool,'F_0') = s(bool,sK28(X17,X16,bool))) ) | $spl1),
  inference(superposition,[],[f299,f831])).
fof(f918,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),X7))) | s(bool,'F_0') = s(bool,sK28(X7,X6,bool))) ) | $spl1),
  inference(subsumption_resolution,[],[f896,f261])).
fof(f896,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),X7))) | s(bool,'F_0') = s(bool,sK28(X7,X6,bool))) ) | $spl1),
  inference(superposition,[],[f269,f831])).
fof(f891,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,sK28('const.pair.,_2'(s(bool,X7),s(X6,X8)),X6,bool)) | s(bool,'T_0') = s(bool,X7)) ) | $spl1),
  inference(superposition,[],[f831,f323])).
fof(f916,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,sK28(sK28(X14,X13,'type.pair.prod'(bool,X12)),X12,bool)) | s(bool,'F_0') = s(bool,sK28('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(bool,X12),X13),X14)),X12,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f888,f701])).
fof(f888,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X12),sK28(X14,X13,'type.pair.prod'(bool,X12))))) | s(bool,'F_0') = s(bool,sK28('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(bool,X12),X13),X14)),X12,bool))) ) | $spl1),
  inference(superposition,[],[f831,f701])).
fof(f915,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,sK28(X10,X8,bool)) | s(bool,'F_0') = s(bool,sK28('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(bool,X8),X9),'const.pair.,_2'(s('type.pair.prod'(bool,X8),X10),s(X9,X11)))),X8,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f887,f701])).
fof(f887,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X8),X10))) | s(bool,'F_0') = s(bool,sK28('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(bool,X8),X9),'const.pair.,_2'(s('type.pair.prod'(bool,X8),X10),s(X9,X11)))),X8,bool))) ) | $spl1),
  inference(superposition,[],[f831,f323])).
fof(f914,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,sK28(X7,X4,bool)) | s(bool,'F_0') = s(bool,sK28('const.pair.SND_1'(s('type.pair.prod'(X5,'type.pair.prod'(bool,X4)),'const.pair.,_2'(s(X5,X6),s('type.pair.prod'(bool,X4),X7)))),X4,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f886,f701])).
fof(f886,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),X7))) | s(bool,'F_0') = s(bool,sK28('const.pair.SND_1'(s('type.pair.prod'(X5,'type.pair.prod'(bool,X4)),'const.pair.,_2'(s(X5,X6),s('type.pair.prod'(bool,X4),X7)))),X4,bool))) ) | $spl1),
  inference(superposition,[],[f831,f322])).
fof(f913,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,sK28(X3,X2,bool)) | s(bool,'F_0') = s(bool,sK28(happ(s(fun('type.pair.prod'(bool,X2),'type.pair.prod'(bool,X2)),'_i_0_0'),s('type.pair.prod'(bool,X2),X3)),X2,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f885,f701])).
fof(f885,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'F_0') = s(bool,sK28(happ(s(fun('type.pair.prod'(bool,X2),'type.pair.prod'(bool,X2)),'_i_0_0'),s('type.pair.prod'(bool,X2),X3)),X2,bool))) ) | $spl1),
  inference(superposition,[],[f831,f279])).
fof(f912,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK28(X1,X0,bool)) | s(bool,'F_0') = s(bool,sK28('const.pair.,_2'(s(bool,sK28(X1,X0,bool)),s(X0,sK29(X1,X0,bool))),X0,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f884,f701])).
fof(f884,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) | s(bool,'F_0') = s(bool,sK28('const.pair.,_2'(s(bool,sK28(X1,X0,bool)),s(X0,sK29(X1,X0,bool))),X0,bool))) ) | $spl1),
  inference(superposition,[],[f831,f321])).
fof(f872,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))))) ) | $spl1),
  inference(forward_demodulation,[],[f850,f391])).
fof(f850,plain,(
  ( ! [X2,X3] : (s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))))) ) | $spl1),
  inference(superposition,[],[f280,f595])).
fof(f871,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) ) | $spl1),
  inference(forward_demodulation,[],[f848,f391])).
fof(f848,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) ) | $spl1),
  inference(superposition,[],[f280,f595])).
fof(f803,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,X6) = s(X4,sK28('const.pair.,_2'(s(X4,X6),s(X5,X7)),X5,X4))) )),
  inference(superposition,[],[f701,f323])).
fof(f831,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) | s(bool,'F_0') = s(bool,sK28(X1,X0,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f801,f391])).
fof(f801,plain,(
  ( ! [X0,X1] : (s(bool,sK28(X1,X0,bool)) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1)))) ) | $spl1),
  inference(superposition,[],[f701,f595])).
fof(f830,plain,(
  ( ! [X19,X17,X20,X18] : (s(X17,sK28(sK28(X20,X19,'type.pair.prod'(X17,X18)),X18,X17)) = s(X17,sK28('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X17,X18),X19),X20)),X18,X17))) )),
  inference(forward_demodulation,[],[f800,f701])).
fof(f800,plain,(
  ( ! [X19,X17,X20,X18] : (s(X17,'const.pair.FST_1'(s('type.pair.prod'(X17,X18),sK28(X20,X19,'type.pair.prod'(X17,X18))))) = s(X17,sK28('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X17,X18),X19),X20)),X18,X17))) )),
  inference(superposition,[],[f701,f701])).
fof(f829,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(X12,sK28(X15,X13,X12)) = s(X12,sK28('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X12,X13),X14),'const.pair.,_2'(s('type.pair.prod'(X12,X13),X15),s(X14,X16)))),X13,X12))) )),
  inference(forward_demodulation,[],[f799,f701])).
fof(f799,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(X12,'const.pair.FST_1'(s('type.pair.prod'(X12,X13),X15))) = s(X12,sK28('const.pair.FST_1'(s('type.pair.prod'('type.pair.prod'(X12,X13),X14),'const.pair.,_2'(s('type.pair.prod'(X12,X13),X15),s(X14,X16)))),X13,X12))) )),
  inference(superposition,[],[f701,f323])).
fof(f828,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(X7,sK28(X11,X8,X7)) = s(X7,sK28('const.pair.SND_1'(s('type.pair.prod'(X9,'type.pair.prod'(X7,X8)),'const.pair.,_2'(s(X9,X10),s('type.pair.prod'(X7,X8),X11)))),X8,X7))) )),
  inference(forward_demodulation,[],[f798,f701])).
fof(f798,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X8),X11))) = s(X7,sK28('const.pair.SND_1'(s('type.pair.prod'(X9,'type.pair.prod'(X7,X8)),'const.pair.,_2'(s(X9,X10),s('type.pair.prod'(X7,X8),X11)))),X8,X7))) )),
  inference(superposition,[],[f701,f322])).
fof(f827,plain,(
  ( ! [X6,X4,X5] : (s(X4,sK28(X6,X5,X4)) = s(X4,sK28(happ(s(fun('type.pair.prod'(X4,X5),'type.pair.prod'(X4,X5)),'_i_0_0'),s('type.pair.prod'(X4,X5),X6)),X5,X4))) )),
  inference(forward_demodulation,[],[f797,f701])).
fof(f797,plain,(
  ( ! [X6,X4,X5] : (s(X4,'const.pair.FST_1'(s('type.pair.prod'(X4,X5),X6))) = s(X4,sK28(happ(s(fun('type.pair.prod'(X4,X5),'type.pair.prod'(X4,X5)),'_i_0_0'),s('type.pair.prod'(X4,X5),X6)),X5,X4))) )),
  inference(superposition,[],[f701,f279])).
fof(f826,plain,(
  ( ! [X2,X3,X1] : (s(X1,sK28(X3,X2,X1)) = s(X1,sK28('const.pair.,_2'(s(X1,sK28(X3,X2,X1)),s(X2,sK29(X3,X2,X1))),X2,X1))) )),
  inference(forward_demodulation,[],[f796,f701])).
fof(f796,plain,(
  ( ! [X2,X3,X1] : (s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X2),X3))) = s(X1,sK28('const.pair.,_2'(s(X1,sK28(X3,X2,X1)),s(X2,sK29(X3,X2,X1))),X2,X1))) )),
  inference(superposition,[],[f701,f321])).
fof(f825,plain,(
  ( ! [X0] : (s(sK1,sK28(X0,sK0,sK1)) = s(sK1,sK28(sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)),sK0,sK1))) ) | $spl1),
  inference(forward_demodulation,[],[f795,f701])).
fof(f795,plain,(
  ( ! [X0] : (s(sK1,'const.pair.FST_1'(s('type.pair.prod'(sK1,sK0),X0))) = s(sK1,sK28(sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)),sK0,sK1))) ) | $spl1),
  inference(superposition,[],[f701,f575])).
fof(f720,plain,(
  ( ! [X10,X8,X9] : (s(bool,X9) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X8),'const.pair.,_2'(s(bool,X9),s(X8,X10)))))) ) | $spl1),
  inference(superposition,[],[f595,f323])).
fof(f719,plain,(
  ( ! [X6,X7,X5] : (s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X5,bool),'const.pair.,_2'(s(X5,X6),s(bool,X7)))))) ) | $spl1),
  inference(superposition,[],[f595,f322])).
fof(f718,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4)))) ) | $spl1),
  inference(superposition,[],[f595,f279])).
fof(f710,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK28(X2,X1,X0)),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2)))))) )),
  inference(backward_demodulation,[],[f701,f307])).
fof(f700,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'(X1,bool),X0) = s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,sK28(X0,bool,X1)),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK29(X0,bool,X1))) )),
  inference(superposition,[],[f321,f264])).
fof(f699,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'(bool,X1),X0) = s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,sK29(X0,X1,bool)))) | s(bool,'T_0') = s(bool,sK28(X0,X1,bool))) )),
  inference(superposition,[],[f321,f264])).
fof(f661,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0))))) )),
  inference(subsumption_resolution,[],[f656,f263])).
fof(f656,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0))))) )),
  inference(superposition,[],[f288,f264])).
fof(f655,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK6(X0,bool))) )),
  inference(superposition,[],[f288,f264])).
fof(f595,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(subsumption_resolution,[],[f589,f263])).
fof(f589,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(superposition,[],[f468,f391])).
fof(f576,plain,(
  ( ! [X4,X2,X3] : (s(X2,X4) = s(X2,sK7(X4,X3,X2)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3)))) )),
  inference(subsumption_resolution,[],[f574,f263])).
fof(f574,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'F_0')) | s(X2,X4) = s(X2,sK7(X4,X3,X2)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3)))) )),
  inference(superposition,[],[f292,f264])).
fof(f575,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK1,sK0),X0) = s('type.pair.prod'(sK1,sK0),sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))) ) | $spl1),
  inference(subsumption_resolution,[],[f572,f263])).
fof(f572,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.pair.prod'(sK1,sK0),X0) = s('type.pair.prod'(sK1,sK0),sK7(X0,'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))) ) | $spl1),
  inference(superposition,[],[f292,f391])).
fof(f546,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2)))))) )),
  inference(superposition,[],[f323,f264])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2)))))) )),
  inference(superposition,[],[f322,f264])).
fof(f442,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(superposition,[],[f279,f264])).
fof(f397,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f264,f264])).
fof(f391,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | $spl1),
  inference(unit_resulting_resolution,[],[f367,f270])).
fof(f382,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | $spl1),
  inference(unit_resulting_resolution,[],[f367,f265])).
fof(f362,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),'_dst_x0000x0011x1111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X5,fun(X6,X1)),happ(s(fun(X2,fun(X5,fun(X6,X1))),X9),s(X2,X11))),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))) = s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),'_dst_x00000x10011x01111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X2,fun(X5,fun(X6,X1))),X9),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))),s(X2,X11)))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),'_dst_x0000x0011x1111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X5,fun(X6,X1)),happ(s(fun(X2,fun(X5,fun(X6,X1))),X9),s(X2,X11))),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))) = s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),'_dst_x00000x10011x01111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X2,fun(X5,fun(X6,X1))),X9),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))),s(X2,X11)))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ! [X8,X9,X10,X11] : s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),'_dst_x0000x0011x1111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X5,fun(X6,X1)),happ(s(fun(X2,fun(X5,fun(X6,X1))),X9),s(X2,X11))),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))) = s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),'_dst_x00000x10011x01111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X2,fun(X5,fun(X6,X1))),X9),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))),s(X2,X11)))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X29,X42,X32,X30,X31,X37,X43,X44] : ! [X33,X34,X38,X35] : s(fun(X30,fun(X31,fun(X37,fun(X43,X29)))),'_dst_x0000x0011x1111_3'(s(fun(X42,fun(X44,X29)),X33),s(fun(X37,fun(X43,X42)),happ(s(fun(X32,fun(X37,fun(X43,X42))),X34),s(X32,X35))),s(fun(X30,fun(X31,fun(X37,fun(X43,X44)))),X38))) = s(fun(X30,fun(X31,fun(X37,fun(X43,X29)))),happ(s(fun(X32,fun(X30,fun(X31,fun(X37,fun(X43,X29))))),'_dst_x00000x10011x01111_3'(s(fun(X42,fun(X44,X29)),X33),s(fun(X32,fun(X37,fun(X43,X42))),X34),s(fun(X30,fun(X31,fun(X37,fun(X43,X44)))),X38))),s(X32,X35)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x00000x10011x01111_ax')).
fof(f361,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x011x111_3'(s(fun(X1,fun(X5,X0)),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X2,fun(X3,fun(X4,X5))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9),s(X6,X10))))) = s(fun(X2,fun(X3,fun(X4,X0))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x0011x1111_3'(s(fun(X1,fun(X5,X0)),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9))),s(X6,X10)))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x011x111_3'(s(fun(X1,fun(X5,X0)),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X2,fun(X3,fun(X4,X5))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9),s(X6,X10))))) = s(fun(X2,fun(X3,fun(X4,X0))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x0011x1111_3'(s(fun(X1,fun(X5,X0)),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9))),s(X6,X10)))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10] : s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x011x111_3'(s(fun(X1,fun(X5,X0)),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X2,fun(X3,fun(X4,X5))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9),s(X6,X10))))) = s(fun(X2,fun(X3,fun(X4,X0))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x0011x1111_3'(s(fun(X1,fun(X5,X0)),X7),s(fun(X3,fun(X4,X1)),X8),s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9))),s(X6,X10)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X29,X43,X30,X31,X37,X42,X32] : ! [X33,X34,X38,X35] : s(fun(X30,fun(X31,fun(X37,X29))),'_dst_x000x011x111_3'(s(fun(X43,fun(X42,X29)),X33),s(fun(X31,fun(X37,X43)),X34),s(fun(X30,fun(X31,fun(X37,X42))),happ(s(fun(X32,fun(X30,fun(X31,fun(X37,X42)))),X38),s(X32,X35))))) = s(fun(X30,fun(X31,fun(X37,X29))),happ(s(fun(X32,fun(X30,fun(X31,fun(X37,X29)))),'_dst_x0000x0011x1111_3'(s(fun(X43,fun(X42,X29)),X33),s(fun(X31,fun(X37,X43)),X34),s(fun(X32,fun(X30,fun(X31,fun(X37,X42)))),X38))),s(X32,X35)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x0000x0011x1111_ax')).
fof(f360,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X4,fun(X1,fun(X5,X0))),'_dst_x000x010x101_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X8),s(X3,X10))),s(fun(X4,fun(X5,X6)),X9))) = s(fun(X4,fun(X1,fun(X5,X0))),happ(s(fun(X3,fun(X4,fun(X1,fun(X5,X0)))),'_dst_x0000x1010x0101_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X3,fun(X1,X2)),X8),s(fun(X4,fun(X5,X6)),X9))),s(X3,X10)))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X4,fun(X1,fun(X5,X0))),'_dst_x000x010x101_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X8),s(X3,X10))),s(fun(X4,fun(X5,X6)),X9))) = s(fun(X4,fun(X1,fun(X5,X0))),happ(s(fun(X3,fun(X4,fun(X1,fun(X5,X0)))),'_dst_x0000x1010x0101_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X3,fun(X1,X2)),X8),s(fun(X4,fun(X5,X6)),X9))),s(X3,X10)))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10] : s(fun(X4,fun(X1,fun(X5,X0))),'_dst_x000x010x101_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X8),s(X3,X10))),s(fun(X4,fun(X5,X6)),X9))) = s(fun(X4,fun(X1,fun(X5,X0))),happ(s(fun(X3,fun(X4,fun(X1,fun(X5,X0)))),'_dst_x0000x1010x0101_3'(s(fun(X2,fun(X6,X0)),X7),s(fun(X3,fun(X1,X2)),X8),s(fun(X4,fun(X5,X6)),X9))),s(X3,X10)))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X29,X31,X43,X32,X30,X37,X42] : ! [X33,X34,X38,X35] : s(fun(X30,fun(X31,fun(X37,X29))),'_dst_x000x010x101_3'(s(fun(X43,fun(X42,X29)),X33),s(fun(X31,X43),happ(s(fun(X32,fun(X31,X43)),X34),s(X32,X35))),s(fun(X30,fun(X37,X42)),X38))) = s(fun(X30,fun(X31,fun(X37,X29))),happ(s(fun(X32,fun(X30,fun(X31,fun(X37,X29)))),'_dst_x0000x1010x0101_3'(s(fun(X43,fun(X42,X29)),X33),s(fun(X32,fun(X31,X43)),X34),s(fun(X30,fun(X37,X42)),X38))),s(X32,X35)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x0000x1010x0101_ax')).
fof(f359,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X2,fun(X3,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X2,fun(X3,X1)),X7),s(fun(X2,fun(X3,X4)),happ(s(fun(X5,fun(X2,fun(X3,X4))),X8),s(X5,X9))))) = s(fun(X2,fun(X3,X0)),happ(s(fun(X5,fun(X2,fun(X3,X0))),'_dst_x000x011x111_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X2,fun(X3,X1)),X7),s(fun(X5,fun(X2,fun(X3,X4))),X8))),s(X5,X9)))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X2,fun(X3,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X2,fun(X3,X1)),X7),s(fun(X2,fun(X3,X4)),happ(s(fun(X5,fun(X2,fun(X3,X4))),X8),s(X5,X9))))) = s(fun(X2,fun(X3,X0)),happ(s(fun(X5,fun(X2,fun(X3,X0))),'_dst_x000x011x111_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X2,fun(X3,X1)),X7),s(fun(X5,fun(X2,fun(X3,X4))),X8))),s(X5,X9)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X2,fun(X3,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X2,fun(X3,X1)),X7),s(fun(X2,fun(X3,X4)),happ(s(fun(X5,fun(X2,fun(X3,X4))),X8),s(X5,X9))))) = s(fun(X2,fun(X3,X0)),happ(s(fun(X5,fun(X2,fun(X3,X0))),'_dst_x000x011x111_3'(s(fun(X1,fun(X4,X0)),X6),s(fun(X2,fun(X3,X1)),X7),s(fun(X5,fun(X2,fun(X3,X4))),X8))),s(X5,X9)))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X29,X37,X30,X31,X43,X32] : ! [X33,X34,X38,X35] : s(fun(X30,fun(X31,X29)),'_dst_x00x11x11_3'(s(fun(X37,fun(X43,X29)),X33),s(fun(X30,fun(X31,X37)),X34),s(fun(X30,fun(X31,X43)),happ(s(fun(X32,fun(X30,fun(X31,X43))),X38),s(X32,X35))))) = s(fun(X30,fun(X31,X29)),happ(s(fun(X32,fun(X30,fun(X31,X29))),'_dst_x000x011x111_3'(s(fun(X37,fun(X43,X29)),X33),s(fun(X30,fun(X31,X37)),X34),s(fun(X32,fun(X30,fun(X31,X43))),X38))),s(X32,X35)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x000x011x111_ax')).
fof(f358,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X1,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X2,fun(X4,X0)),X6),s(fun(X1,X2),X7),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X8),s(X5,X9))))) = s(fun(X1,fun(X3,X0)),happ(s(fun(X5,fun(X1,fun(X3,X0))),'_dst_x000x010x101_3'(s(fun(X2,fun(X4,X0)),X6),s(fun(X1,X2),X7),s(fun(X5,fun(X3,X4)),X8))),s(X5,X9)))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X1,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X2,fun(X4,X0)),X6),s(fun(X1,X2),X7),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X8),s(X5,X9))))) = s(fun(X1,fun(X3,X0)),happ(s(fun(X5,fun(X1,fun(X3,X0))),'_dst_x000x010x101_3'(s(fun(X2,fun(X4,X0)),X6),s(fun(X1,X2),X7),s(fun(X5,fun(X3,X4)),X8))),s(X5,X9)))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X1,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X2,fun(X4,X0)),X6),s(fun(X1,X2),X7),s(fun(X3,X4),happ(s(fun(X5,fun(X3,X4)),X8),s(X5,X9))))) = s(fun(X1,fun(X3,X0)),happ(s(fun(X5,fun(X1,fun(X3,X0))),'_dst_x000x010x101_3'(s(fun(X2,fun(X4,X0)),X6),s(fun(X1,X2),X7),s(fun(X5,fun(X3,X4)),X8))),s(X5,X9)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X29,X30,X37,X31,X43,X32] : ! [X33,X34,X38,X35] : s(fun(X30,fun(X31,X29)),'_dst_x00x10x01_3'(s(fun(X37,fun(X43,X29)),X33),s(fun(X30,X37),X34),s(fun(X31,X43),happ(s(fun(X32,fun(X31,X43)),X38),s(X32,X35))))) = s(fun(X30,fun(X31,X29)),happ(s(fun(X32,fun(X30,fun(X31,X29))),'_dst_x000x010x101_3'(s(fun(X37,fun(X43,X29)),X33),s(fun(X30,X37),X34),s(fun(X32,fun(X31,X43)),X38))),s(X32,X35)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x000x010x101_ax')).
fof(f357,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111_2'(s(fun(X4,X0),X6),s(fun(X1,fun(X2,fun(X3,X4))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7),s(X5,X8))))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X0)))),'_dst_x0000x1111_2'(s(fun(X4,X0),X6),s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7))),s(X5,X8)))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111_2'(s(fun(X4,X0),X6),s(fun(X1,fun(X2,fun(X3,X4))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7),s(X5,X8))))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X0)))),'_dst_x0000x1111_2'(s(fun(X4,X0),X6),s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7))),s(X5,X8)))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8] : s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111_2'(s(fun(X4,X0),X6),s(fun(X1,fun(X2,fun(X3,X4))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7),s(X5,X8))))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X0)))),'_dst_x0000x1111_2'(s(fun(X4,X0),X6),s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7))),s(X5,X8)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X29,X30,X31,X37,X43,X32] : ! [X33,X34,X35] : s(fun(X30,fun(X31,fun(X37,X29))),'_dst_x000x111_2'(s(fun(X43,X29),X33),s(fun(X30,fun(X31,fun(X37,X43))),happ(s(fun(X32,fun(X30,fun(X31,fun(X37,X43)))),X34),s(X32,X35))))) = s(fun(X30,fun(X31,fun(X37,X29))),happ(s(fun(X32,fun(X30,fun(X31,fun(X37,X29)))),'_dst_x0000x1111_2'(s(fun(X43,X29),X33),s(fun(X32,fun(X30,fun(X31,fun(X37,X43)))),X34))),s(X32,X35)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x0000x1111_ax')).
fof(f356,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X4,fun(X2,X1)),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8))) = s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),happ(s(fun(X4,fun(X2,X1)),X6),s(X4,X8))),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8)))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X4,fun(X2,X1)),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8))) = s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),happ(s(fun(X4,fun(X2,X1)),X6),s(X4,X8))),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8)))))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x11x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X4,fun(X2,X1)),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8))) = s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),happ(s(fun(X4,fun(X2,X1)),X6),s(X4,X8))),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X29,X31,X30,X37,X32] : ! [X33,X34,X38,X35] : s(fun(X30,X29),happ(s(fun(X32,fun(X30,X29)),'_dst_x00x11x11_3'(s(fun(X31,fun(X37,X29)),X33),s(fun(X32,fun(X30,X31)),X34),s(fun(X32,fun(X30,X37)),X38))),s(X32,X35))) = s(fun(X30,X29),'_dst_x0x1x1_3'(s(fun(X31,fun(X37,X29)),X33),s(fun(X30,X31),happ(s(fun(X32,fun(X30,X31)),X34),s(X32,X35))),s(fun(X30,X37),happ(s(fun(X32,fun(X30,X37)),X38),s(X32,X35)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x00x11x11_ax')).
fof(f355,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X29,X31,X32,X30,X37] : ! [X33,X34,X38,X35] : s(fun(X30,X29),'_dst_x0x0x1_3'(s(fun(X31,fun(X37,X29)),X33),s(X31,happ(s(fun(X32,X31),X34),s(X32,X35))),s(fun(X30,X37),X38))) = s(fun(X30,X29),happ(s(fun(X32,fun(X30,X29)),'_dst_x00x10x01_3'(s(fun(X31,fun(X37,X29)),X33),s(fun(X32,X31),X34),s(fun(X30,X37),X38))),s(X32,X35)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x00x10x01_ax')).
fof(f354,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun(X3,X0),X5),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X6),s(X4,X7))))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111_2'(s(fun(X3,X0),X5),s(fun(X4,fun(X1,fun(X2,X3))),X6))),s(X4,X7)))) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun(X3,X0),X5),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X6),s(X4,X7))))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111_2'(s(fun(X3,X0),X5),s(fun(X4,fun(X1,fun(X2,X3))),X6))),s(X4,X7)))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7] : s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun(X3,X0),X5),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X6),s(X4,X7))))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111_2'(s(fun(X3,X0),X5),s(fun(X4,fun(X1,fun(X2,X3))),X6))),s(X4,X7)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X29,X30,X31,X37,X32] : ! [X33,X34,X35] : s(fun(X30,fun(X31,X29)),'_dst_x00x11_2'(s(fun(X37,X29),X33),s(fun(X30,fun(X31,X37)),happ(s(fun(X32,fun(X30,fun(X31,X37))),X34),s(X32,X35))))) = s(fun(X30,fun(X31,X29)),happ(s(fun(X32,fun(X30,fun(X31,X29))),'_dst_x000x111_2'(s(fun(X37,X29),X33),s(fun(X32,fun(X30,fun(X31,X37))),X34))),s(X32,X35)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x000x111_ax')).
fof(f353,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X29,X30,X31,X32] : ! [X33,X34,X38,X35] : s(X29,happ(s(fun(X32,X29),'_dst_x0x1x1_3'(s(fun(X30,fun(X31,X29)),X33),s(fun(X32,X30),X34),s(fun(X32,X31),X38))),s(X32,X35))) = s(X29,happ(s(fun(X31,X29),happ(s(fun(X30,fun(X31,X29)),X33),s(X30,happ(s(fun(X32,X30),X34),s(X32,X35))))),s(X31,happ(s(fun(X32,X31),X38),s(X32,X35)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x0x1x1_ax')).
fof(f352,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X29,X30,X31,X32] : ! [X33,X34,X38,X35] : s(X29,happ(s(fun(X32,X29),'_dst_x0x0x1_3'(s(fun(X30,fun(X31,X29)),X33),s(X30,X34),s(fun(X32,X31),X38))),s(X32,X35))) = s(X29,happ(s(fun(X31,X29),happ(s(fun(X30,fun(X31,X29)),X33),s(X30,X34))),s(X31,happ(s(fun(X32,X31),X38),s(X32,X35)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x0x0x1_ax')).
fof(f351,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X29,X30,X31,X32] : ! [X33,X34,X35] : s(fun(X30,X29),'_dst_x0x1_2'(s(fun(X31,X29),X33),s(fun(X30,X31),happ(s(fun(X32,fun(X30,X31)),X34),s(X32,X35))))) = s(fun(X30,X29),happ(s(fun(X32,fun(X30,X29)),'_dst_x00x11_2'(s(fun(X31,X29),X33),s(fun(X32,fun(X30,X31)),X34))),s(X32,X35)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x00x11_ax')).
fof(f350,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),'_dst_x1x0_2'(s(fun(X0,fun(X2,X1)),X4),s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x01x10_2'(s(fun(X0,fun(X2,X1)),X4),s(fun(X3,X2),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X0,X1),'_dst_x1x0_2'(s(fun(X0,fun(X2,X1)),X4),s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x01x10_2'(s(fun(X0,fun(X2,X1)),X4),s(fun(X3,X2),X5))),s(X3,X6)))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X0,X1),'_dst_x1x0_2'(s(fun(X0,fun(X2,X1)),X4),s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x01x10_2'(s(fun(X0,fun(X2,X1)),X4),s(fun(X3,X2),X5))),s(X3,X6)))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X30,X29,X31,X32] : ! [X33,X34,X35] : s(fun(X30,X29),'_dst_x1x0_2'(s(fun(X30,fun(X31,X29)),X33),s(X31,happ(s(fun(X32,X31),X34),s(X32,X35))))) = s(fun(X30,X29),happ(s(fun(X32,fun(X30,X29)),'_dst_x01x10_2'(s(fun(X30,fun(X31,X29)),X33),s(fun(X32,X31),X34))),s(X32,X35)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x01x10_ax')).
fof(f380,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),X5),s(X1,X4),s(X2,X3))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X1,X2),X0),X5))),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3)))))) )),
  inference(forward_demodulation,[],[f349,f343])).
fof(f349,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X5),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X1,X2),X0),X5))),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3)))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X5),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X1,X2),X0),X5))),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3)))))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X5),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X1,X2),X0),X5))),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X11,X7,X6] : ! [X14,X9,X15] : s(X11,happ(s(fun('type.pair.prod'(X7,X6),X11),X15),s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X9),s(X6,X14))))) = s(X11,'const.pair.UNCURRY_2'(s(fun(X7,fun(X6,X11)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X7,X6),X11),X15))),s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X9),s(X6,X14)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.pair_2|49|_prove_ltb')).
fof(f348,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s(X1,X4),s(X2,X3)))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s(X1,X4),s(X2,X3)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s(X1,X4),s(X2,X3)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X11,X7,X6] : ! [X14,X9,X15] : s(X11,happ(s(fun(X6,X11),happ(s(fun(X7,fun(X6,X11)),X15),s(X7,X9))),s(X6,X14))) = s(X11,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X7,X6),X11),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,X11)),X15))),s(X7,X9),s(X6,X14)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.pair_|49|9_prove_ltb')).
fof(f347,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),X3),s(X1,X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X1,X2),X0),X3))),s(X1,X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),X3),s(X1,X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X1,X2),X0),X3))),s(X1,X4))),s(X2,X5)))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),X3),s(X1,X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X1,X2),X0),X3))),s(X1,X4))),s(X2,X5)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X11,X7,X6] : ! [X12,X13,X19] : s(X11,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X7,X6),X11),X12),s(X7,X13),s(X6,X19))) = s(X11,happ(s(fun(X6,X11),happ(s(fun(X7,fun(X6,X11)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X7,X6),X11),X12))),s(X7,X13))),s(X6,X19)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','ARITY_1_3_const.pair.CURRY')).
fof(f346,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X29,X32,X30] : ! [X33,X34,X35] : s(X29,happ(s(fun(X30,X29),happ(s(fun(X32,fun(X30,X29)),X33),s(X32,X35))),s(X30,X34))) = s(X29,happ(s(fun(X32,X29),'_dst_x1x0_2'(s(fun(X32,fun(X30,X29)),X33),s(X30,X34))),s(X32,X35)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x1x0_ax')).
fof(f345,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X11,X7,X6] : ! [X15,X9,X14] : s(X11,happ(s(fun(X6,X11),happ(s(fun(X7,fun(X6,X11)),X15),s(X7,X9))),s(X6,X14))) = s(X11,'const.pair.UNCURRY_2'(s(fun(X7,fun(X6,X11)),X15),s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X9),s(X6,X14)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.UNCURRY_DEF')).
fof(f344,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X29,X30,X32] : ! [X33,X34,X35] : s(X29,happ(s(fun(X30,X29),X33),s(X30,happ(s(fun(X32,X30),X34),s(X32,X35))))) = s(X29,happ(s(fun(X32,X29),'_dst_x0x1_2'(s(fun(X30,X29),X33),s(fun(X32,X30),X34))),s(X32,X35)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_dst_x0x1_ax')).
fof(f343,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),X3),s(X1,X4),s(X2,X5))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),X3),s(X1,X4),s(X2,X5))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X1,X2),X0),X3),s(X1,X4),s(X2,X5))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X11,X7,X6] : ! [X15,X9,X14] : s(X11,'const.pair.CURRY_3'(s(fun('type.pair.prod'(X7,X6),X11),X15),s(X7,X9),s(X6,X14))) = s(X11,happ(s(fun('type.pair.prod'(X7,X6),X11),X15),s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X9),s(X6,X14)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.CURRY_DEF')).
fof(f341,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))) != s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X4))) | s(fun('type.pair.prod'(X0,X1),X2),X3) = s(fun('type.pair.prod'(X0,X1),X2),X4)) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2,X3,X4] : ((s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))) != s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X4))) | s(fun('type.pair.prod'(X0,X1),X2),X3) = s(fun('type.pair.prod'(X0,X1),X2),X4)) & (s(fun('type.pair.prod'(X0,X1),X2),X3) != s(fun('type.pair.prod'(X0,X1),X2),X4) | s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))) = s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X4)))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))) = s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X4))) <=> s(fun('type.pair.prod'(X0,X1),X2),X3) = s(fun('type.pair.prod'(X0,X1),X2),X4))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : (s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))) = s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X4))) <=> s(fun('type.pair.prod'(X0,X1),X2),X3) = s(fun('type.pair.prod'(X0,X1),X2),X4))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X7,X6,X11] : ! [X24,X15] : (s(fun(X7,fun(X6,X11)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X7,X6),X11),X15))) = s(fun(X7,fun(X6,X11)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X7,X6),X11),X24))) <=> s(fun('type.pair.prod'(X7,X6),X11),X15) = s(fun('type.pair.prod'(X7,X6),X11),X24))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.CURRY_ONE_ONE_THM')).
fof(f342,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))) = s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X4))) | s(fun('type.pair.prod'(X0,X1),X2),X3) != s(fun('type.pair.prod'(X0,X1),X2),X4)) )),
  inference(cnf_transformation,[],[f258])).
fof(f339,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))) != s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X4))) | s(fun(X0,fun(X1,X2)),X3) = s(fun(X0,fun(X1,X2)),X4)) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2,X3,X4] : ((s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))) != s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X4))) | s(fun(X0,fun(X1,X2)),X3) = s(fun(X0,fun(X1,X2)),X4)) & (s(fun(X0,fun(X1,X2)),X3) != s(fun(X0,fun(X1,X2)),X4) | s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))) = s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X4)))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))) = s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X4))) <=> s(fun(X0,fun(X1,X2)),X3) = s(fun(X0,fun(X1,X2)),X4))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : (s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))) = s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X4))) <=> s(fun(X0,fun(X1,X2)),X3) = s(fun(X0,fun(X1,X2)),X4))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X7,X6,X11] : ! [X24,X15] : (s(fun('type.pair.prod'(X7,X6),X11),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,X11)),X15))) = s(fun('type.pair.prod'(X7,X6),X11),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,X11)),X24))) <=> s(fun(X7,fun(X6,X11)),X15) = s(fun(X7,fun(X6,X11)),X24))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.UNCURRY_ONE_ONE_THM')).
fof(f340,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X1,X2)),X3) != s(fun(X0,fun(X1,X2)),X4) | s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))) = s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X4)))) )),
  inference(cnf_transformation,[],[f257])).
fof(f336,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun('type.pair.prod'(X1,X2),X0),X3),s(fun(X1,fun(X2,'type.pair.prod'(X1,X2))),'const.pair.,_0'))),s('type.pair.prod'(X1,X2),X4)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun('type.pair.prod'(X1,X2),X0),X3),s(fun(X1,fun(X2,'type.pair.prod'(X1,X2))),'const.pair.,_0'))),s('type.pair.prod'(X1,X2),X4)))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun('type.pair.prod'(X1,X2),X0),X3),s(fun(X1,fun(X2,'type.pair.prod'(X1,X2))),'const.pair.,_0'))),s('type.pair.prod'(X1,X2),X4)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X11,X7,X6] : ! [X8,X5] : s(X11,happ(s(fun('type.pair.prod'(X7,X6),X11),X8),s('type.pair.prod'(X7,X6),X5))) = s(X11,'const.pair.UNCURRY_2'(s(fun(X7,fun(X6,X11)),'_dst_x00x11_2'(s(fun('type.pair.prod'(X7,X6),X11),X8),s(fun(X7,fun(X6,'type.pair.prod'(X7,X6))),'const.pair.,_0'))),s('type.pair.prod'(X7,X6),X5)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.LAMBDA_PROD')).
fof(f334,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),sK35(X3,X2,X1,X0)),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),sK35(X3,X2,X1,X0)),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK35])],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X4),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : ! [X3] : ? [X4] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X4),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X11,X7,X6] : ! [X15] : ? [X27] : ! [X9,X14] : s(X11,happ(s(fun(X6,X11),happ(s(fun(X7,fun(X6,X11)),X15),s(X7,X9))),s(X6,X14))) = s(X11,happ(s(fun('type.pair.prod'(X7,X6),X11),X27),s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X9),s(X6,X14)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.pair_Axiom')).
fof(f333,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun('type.pair.prod'(X0,X1),X2),X3) = s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X3)))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : s(fun('type.pair.prod'(X0,X1),X2),X3) = s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X3)))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : ! [X3] : s(fun('type.pair.prod'(X0,X1),X2),X3) = s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),X3)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X7,X6,X11] : ! [X15] : s(fun('type.pair.prod'(X7,X6),X11),X15) = s(fun('type.pair.prod'(X7,X6),X11),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,X11)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X7,X6),X11),X15)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.UNCURRY_CURRY_THM')).
fof(f332,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X1,X2)),X3) = s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3)))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3] : s(fun(X0,fun(X1,X2)),X3) = s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3)))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : ! [X3] : s(fun(X0,fun(X1,X2)),X3) = s(fun(X0,fun(X1,X2)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X7,X6,X11] : ! [X15] : s(fun(X7,fun(X6,X11)),X15) = s(fun(X7,fun(X6,X11)),'const.pair.CURRY_1'(s(fun('type.pair.prod'(X7,X6),X11),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,X11)),X15)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.CURRY_UNCURRY_THM')).
fof(f325,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK30(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK30(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK30(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK30(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK30])],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pair.ELIM_PEXISTS_EVAL',aHL_EXT)).
fof(f317,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK23(X2,X1,X0)))),s(X1,sK24(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X8))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X8))))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((! [X8] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X8))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X8)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK23(X2,X1,X0)))),s(X1,sK24(X2,X1,X0)))))) & (! [X4,X5] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),sK25(X2,X1,X0)))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),sK25(X2,X1,X0)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23,sK24,sK25])],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((! [X8] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X8))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X8)))))) | ? [X6,X7] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7))))) & (! [X4,X5] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ? [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))))))))),
  inference(rectify,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((! [X3] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))) | ? [X4,X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5))))) & (! [X4,X5] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ? [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))))))))),
  inference(nnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : (? [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))) <=> ? [X4,X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ! [X2] : (? [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))) <=> ? [X4,X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7,X6] : ! [X8] : (? [X16] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X6),X16))))),s(X6,'const.pair.SND_1'(s('type.pair.prod'(X7,X6),X16)))))) <=> ? [X17,X18] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,X17))),s(X6,X18)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.ELIM_PEXISTS')).
fof(f318,plain,(
  ( ! [X4,X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),sK25(X2,X1,X0)))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),sK25(X2,X1,X0))))))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f315,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK20(X2,X1,X0)),s(X1,sK21(X2,X1,X0))))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X8))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((! [X8] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X8)))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK20(X2,X1,X0)),s(X1,sK21(X2,X1,X0)))))))) & (! [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK22(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22])],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((! [X8] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X8)))) | ? [X6,X7] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))))))) & (! [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ? [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(rectify,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((! [X3] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | ? [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))))))) & (! [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ? [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : (? [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) <=> ? [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ! [X2] : (? [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) <=> ? [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X7,X6] : ! [X8] : (? [X16] : p(s(bool,happ(s(fun('type.pair.prod'(X7,X6),bool),X8),s('type.pair.prod'(X7,X6),X16)))) <=> ? [X20,X21] : p(s(bool,happ(s(fun('type.pair.prod'(X7,X6),bool),X8),s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X20),s(X6,X21)))))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.EXISTS_PROD')).
fof(f316,plain,(
  ( ! [X4,X2,X0,X5,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK22(X2,X1,X0))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f311,plain,(
  ( ! [X6,X2,X0,X7,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),sK15(X2,X1,X0)))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),sK15(X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),sK15(X2,X1,X0)))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),sK15(X2,X1,X0))))))) | ! [X6,X7] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7))))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK16(X2,X1,X0)))),s(X1,sK17(X2,X1,X0))))) | ! [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16,sK17])],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((? [X8] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X8))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X8)))))) | ! [X6,X7] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7))))) & (? [X4,X5] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ! [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))))))))),
  inference(rectify,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((? [X3] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))) | ! [X4,X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5))))) & (? [X4,X5] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ! [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))))))))),
  inference(nnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : (! [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))) <=> ! [X4,X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : ! [X2] : (! [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))) <=> ! [X4,X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X7,X6] : ! [X8] : (! [X16] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X6),X16))))),s(X6,'const.pair.SND_1'(s('type.pair.prod'(X7,X6),X16)))))) <=> ! [X17,X18] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,X17))),s(X6,X18)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.ELIM_PFORALL')).
fof(f312,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK16(X2,X1,X0)))),s(X1,sK17(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f309,plain,(
  ( ! [X6,X2,X0,X7,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK12(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK12(X2,X1,X0))))) | ! [X6,X7] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))))))) & (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK13(X2,X1,X0)),s(X1,sK14(X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((? [X8] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X8)))) | ! [X6,X7] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))))))) & (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(rectify,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((? [X3] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | ! [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))))))) & (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : (! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) <=> ! [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ! [X2] : (! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) <=> ! [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7,X6] : ! [X8] : (! [X16] : p(s(bool,happ(s(fun('type.pair.prod'(X7,X6),bool),X8),s('type.pair.prod'(X7,X6),X16)))) <=> ! [X20,X21] : p(s(bool,happ(s(fun('type.pair.prod'(X7,X6),bool),X8),s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X20),s(X6,X21)))))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.FORALL_PROD')).
fof(f310,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK13(X2,X1,X0)),s(X1,sK14(X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f308,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK10(X2,X1,X0)),s(X1,sK11(X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK10(X2,X1,X0)),s(X1,sK11(X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))))),
  inference(rectify,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : (? [X3,X4] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) | ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : (! [X3,X4] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) => ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : ! [X2] : (! [X3,X4] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) => ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X7,X6] : ! [X8] : (! [X20,X21] : p(s(bool,happ(s(fun('type.pair.prod'(X7,X6),bool),X8),s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X20),s(X6,X21)))))) => ! [X16] : p(s(bool,happ(s(fun('type.pair.prod'(X7,X6),bool),X8),s('type.pair.prod'(X7,X6),X16)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.pair_induction')).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2)))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2)))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X7,X6] : ! [X9] : s('type.pair.prod'(X7,X6),X9) = s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X6),X9))),s(X6,'const.pair.SND_1'(s('type.pair.prod'(X7,X6),X9)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.PAIR')).
fof(f302,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X1,X0)))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X1,X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f225])).
fof(f225,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X39,X46] : (p(s(bool,happ(s(fun(fun(X39,bool),bool),'?_1'),s(fun(X39,bool),X46)))) <=> ? [X40] : p(s(bool,happ(s(fun(X39,bool),X46),s(X39,X40)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','?_HL_QUANT')).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f300,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f222])).
fof(f222,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X39,X46] : (p(s(bool,happ(s(fun(fun(X39,bool),bool),'!_1'),s(fun(X39,bool),X46)))) <=> ! [X40] : p(s(bool,happ(s(fun(X39,bool),X46),s(X39,X40)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','!_HL_QUANT')).
fof(f301,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f284,plain,(
  ( ! [X4,X0,X5,X1] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | s(X0,X4) = s(X0,X5)) )),
  inference(cnf_transformation,[],[f213])).
fof(f285,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X1,X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f213])).
fof(f286,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X1,X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f213])).
fof(f287,plain,(
  ( ! [X0,X1] : (s(X0,sK4(X1,X0)) != s(X0,sK5(X1,X0)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f213])).
fof(f280,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X7] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X7,bool),X5))) = s(bool,happ(s(fun(X7,bool),X5),s(X7,'const.min.@_1'(s(fun(X7,bool),X5)))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.bool.EXISTS_DEF')).
fof(f279,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X32] : ! [X34] : s(X32,X34) = s(X32,happ(s(fun(X32,X32),'_i_0_0'),s(X32,X34)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','_i0_ax')).
fof(f269,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.bool.EQ_CLAUSES|split|2')).
fof(f270,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f266,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f8372,plain,(
  $false | ($spl1 | $spl28)),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f266,f265,f270,f269,f279,f280,f287,f286,f285,f284,f283,f289,f288,f299,f296,f301,f300,f303,f302,f307,f308,f310,f309,f312,f311,f316,f315,f318,f317,f321,f322,f323,f325,f332,f333,f334,f335,f336,f337,f340,f339,f342,f341,f343,f344,f345,f346,f347,f348,f380,f350,f351,f352,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f382,f381,f391,f397,f442,f468,f509,f546,f575,f576,f595,f611,f655,f661,f699,f700,f701,f709,f710,f702,f712,f718,f719,f720,f825,f826,f827,f828,f829,f830,f831,f803,f871,f872,f912,f913,f914,f915,f916,f891,f918,f920,f934,f935,f936,f937,f938,f930,f993,f998,f999,f1003,f1051,f1030,f1056,f1196,f1225,f1270,f1272,f1273,f1274,f1275,f1276,f1277,f1233,f1278,f1279,f1298,f1299,f1300,f1326,f1329,f326,f379,f378,f377,f376,f375,f1091,f939,f306,f305,f304,f973,f1127,f1295,f841,f1209,f331,f330,f329,f328,f327,f320,f371,f313,f372,f314,f319,f373,f374,f659,f660,f658,f968,f1122,f1265,f570,f823,f571,f569,f292,f1479,f1480,f1481,f1482,f1483,f1484,f1453,f1486,f1488,f1489,f1516,f1517,f1518,f1519,f1520,f1521,f1512,f1531,f1553,f1561,f1563,f1605,f1607,f1669,f1737,f1738,f1920,f1828,f1837,f1839,f1848,f1857,f1859,f1942,f1945,f2278,f2041,f2049,f2050,f2051,f2052,f2053,f2283,f2062,f2073,f2075,f2308,f2325,f2326,f2327,f2328,f2329,f2344,f2409,f2410,f2416,f2417,f2357,f2423,f2424,f2441,f2442,f2384,f2448,f2449,f2403,f2479,f2481,f2464,f2463,f2503,f2505,f2488,f2487,f2601,f2604,f2607,f2615,f2618,f2621,f2648,f2654,f2663,f2664,f2549,f2670,f2671,f2688,f2689,f2576,f2695,f2696,f2595,f2709,f2713,f2733,f2734,f2736,f2956,f3146,f3158,f3160,f3271,f3260,f3339,f3313,f3341,f3410,f3352,f3418,f3471,f3460,f3523,f3524,f3525,f3526,f3527,f3528,f3529,f3530,f3522,f3638,f3639,f3640,f3641,f3642,f3643,f3633,f4014,f4063,f4065,f4141,f4080,f4149,f4230,f4232,f4234,f4236,f4238,f4240,f3517,f2575,f2548,f2383,f2356,f1944,f1244,f854,f849,f811,f4416,f4418,f4420,f4422,f4424,f4426,f367,f4497,f4552,f4551,f4554,f4547,f4619,f4620,f4722,f4723,f4753,f4838,f4839,f4851,f5026,f5027,f5028,f5029,f5030,f5031,f5128,f5129,f5130,f5131,f5132,f5133,f5249,f5327,f5328,f5329,f5330,f5257,f5258,f5337,f5277,f5278,f5280,f5281,f5351,f5353,f5354,f5302,f5357,f5358,f5360,f5361,f5362,f5363,f5470,f5471,f5472,f5473,f5493,f5494,f5495,f5496,f5497,f5498,f5513,f5514,f5515,f5516,f5571,f5573,f5574,f5523,f5577,f5578,f5580,f5581,f5582,f5583,f5636,f5637,f5638,f5639,f5640,f5641,f5642,f5643,f5632,f5633,f5728,f5729,f5730,f5731,f5732,f5733,f5734,f5735,f5724,f5725,f5840,f5841,f5842,f5843,f5844,f5845,f5846,f5847,f5835,f5848,f5838,f6055,f6172,f6173,f6174,f6202,f6205,f6274,f6273,f6275,f6276,f6277,f6278,f6279,f6280,f6281,f6502,f6335,f6336,f6347,f6350,f6351,f6352,f6353,f6354,f6359,f6360,f6361,f6362,f6363,f6364,f6365,f6366,f6367,f6368,f6370,f6371,f6372,f6373,f6374,f6376,f6377,f6380,f6381,f6382,f6383,f6384,f6385,f6387,f6388,f6389,f6390,f6391,f6392,f6393,f6394,f6395,f6396,f6397,f6398,f6399,f6400,f6401,f6402,f6403,f6404,f6411,f6412,f6432,f6435,f6436,f6437,f6438,f6439,f6444,f6445,f6447,f6448,f6449,f6450,f6452,f6453,f6454,f6455,f6456,f6458,f6459,f6462,f6463,f6464,f6465,f6467,f6468,f6469,f6470,f6471,f6472,f6473,f6474,f6475,f6476,f6477,f6478,f6479,f6480,f6481,f6482,f6483,f6484,f6491,f6492,f6570,f6572,f6573,f6574,f6575,f6562,f6563,f6576,f6577,f6578,f6579,f6739,f6740,f6741,f6742,f6743,f6660,f6661,f6744,f6745,f6746,f6747,f6697,f6698,f6699,f6700,f6701,f6733,f6736,f6737,f6984,f6987,f6990,f7008,f7010,f7012,f7027,f7029,f7031,f7062,f7064,f7066,f7092,f6903,f6904,f6909,f6929,f6930,f7235,f7234,f7236,f7237,f7238,f7239,f7240,f7241,f7242,f7327,f7329,f7340,f7343,f7344,f7345,f7346,f7347,f7348,f7356,f7357,f7359,f7360,f7361,f7362,f7363,f7364,f7365,f7366,f7367,f7368,f7369,f7370,f7371,f7372,f7373,f7374,f7375,f7376,f7377,f7378,f7387,f7388,f7398,f7401,f7402,f7403,f7404,f7405,f7406,f7429,f7430,f7431,f7432,f7433,f7434,f7435,f7436,f7437,f7438,f7439,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7448,f7457,f7458,f7563,f7562,f7877,f7565,f7566,f7567,f7621,f7622,f7623,f7624,f7625,f7626,f7627,f7666,f7667,f7668,f7669,f7671,f7673,f7684,f7687,f7688,f7689,f7690,f7691,f7692,f7702,f7703,f7705,f7706,f7707,f7708,f7710,f7711,f7712,f7713,f7714,f7716,f7717,f7720,f7721,f7722,f7723,f7724,f7725,f7727,f7728,f7729,f7730,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7739,f7740,f7741,f7742,f7743,f7744,f7745,f7746,f7756,f7757,f7771,f7782,f7785,f7786,f7787,f7788,f7789,f7790,f7800,f7801,f7802,f7803,f7804,f7805,f7806,f7807,f7809,f7810,f7811,f7812,f7813,f7814,f7818,f7819,f7820,f7821,f7822,f7824,f7825,f7828,f7829,f7830,f7831,f7832,f7833,f7835,f7836,f7837,f7838,f7839,f7840,f7841,f7842,f7843,f7844,f7845,f7846,f7847,f7848,f7849,f7850,f7851,f7852,f7853,f7854,f7864,f7865,f7981,f7982,f7983,f7984,f7985,f7986,f7987,f7988,f7989,f7992,f7995,f8025,f8026,f8027,f8028,f8029,f8031,f8032,f8033,f8034,f8035,f8052,f8053,f8054,f8055,f8129,f8131,f8132,f8063,f8135,f8136,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8147,f8148,f4657,f5953,f2345])).
fof(f8371,plain,(
  $false | ($spl1 | $spl20 | $spl28)),
  inference(global_subsumption,[],[f4243,f260,f259,f261,f263,f264,f266,f265,f270,f269,f279,f280,f287,f286,f285,f284,f283,f289,f288,f299,f296,f301,f300,f303,f302,f307,f308,f310,f309,f312,f311,f316,f315,f318,f317,f321,f322,f323,f325,f332,f333,f334,f335,f336,f337,f340,f339,f342,f341,f343,f344,f345,f346,f347,f348,f380,f350,f351,f352,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f382,f381,f391,f397,f442,f468,f509,f546,f575,f576,f595,f611,f655,f661,f699,f700,f701,f709,f710,f702,f712,f718,f719,f720,f825,f826,f827,f828,f829,f830,f831,f803,f871,f872,f912,f913,f914,f915,f916,f891,f918,f920,f934,f935,f936,f937,f938,f930,f993,f998,f999,f1003,f1051,f1030,f1056,f1196,f1225,f1270,f1272,f1273,f1274,f1275,f1276,f1277,f1233,f1278,f1279,f1298,f1299,f1300,f1326,f1329,f326,f379,f378,f377,f376,f375,f1091,f939,f306,f305,f304,f973,f1127,f1295,f841,f1209,f331,f330,f329,f328,f327,f320,f371,f313,f372,f314,f319,f373,f374,f659,f660,f658,f968,f1122,f1265,f570,f823,f571,f569,f292,f1479,f1480,f1481,f1482,f1483,f1484,f1453,f1486,f1488,f1489,f1516,f1517,f1518,f1519,f1520,f1521,f1512,f1531,f1553,f1561,f1563,f1605,f1607,f1669,f1737,f1738,f1920,f1828,f1837,f1839,f1848,f1857,f1859,f1942,f1945,f2278,f2041,f2049,f2050,f2051,f2052,f2053,f2283,f2062,f2073,f2075,f2308,f2325,f2326,f2327,f2328,f2329,f2344,f2409,f2410,f2416,f2417,f2357,f2423,f2424,f2441,f2442,f2384,f2448,f2449,f2403,f2479,f2481,f2464,f2463,f2503,f2505,f2488,f2487,f2601,f2604,f2607,f2615,f2618,f2621,f2648,f2654,f2663,f2664,f2549,f2670,f2671,f2688,f2689,f2576,f2695,f2696,f2595,f2709,f2713,f2733,f2734,f2736,f2956,f3146,f3158,f3160,f3271,f3260,f3339,f3313,f3341,f3410,f3352,f3418,f3471,f3460,f3523,f3524,f3525,f3526,f3527,f3528,f3529,f3530,f3522,f3638,f3639,f3640,f3641,f3642,f3643,f3633,f4014,f4063,f4065,f4141,f4080,f4149,f4230,f4232,f4234,f4236,f4238,f4240,f3517,f2575,f2548,f2383,f2356,f1944,f1244,f854,f849,f811,f4416,f4418,f4420,f4422,f4424,f4426,f367,f4497,f4552,f4551,f4554,f4547,f4619,f4620,f4722,f4723,f4753,f4838,f4839,f4851,f5026,f5027,f5028,f5029,f5030,f5031,f5128,f5129,f5130,f5131,f5132,f5133,f5249,f5327,f5328,f5329,f5330,f5257,f5258,f5337,f5277,f5278,f5280,f5281,f5351,f5353,f5354,f5302,f5357,f5358,f5360,f5361,f5362,f5363,f5470,f5471,f5472,f5473,f5493,f5494,f5495,f5496,f5497,f5498,f5513,f5514,f5515,f5516,f5571,f5573,f5574,f5523,f5577,f5578,f5580,f5581,f5582,f5583,f5636,f5637,f5638,f5639,f5640,f5641,f5642,f5643,f5632,f5633,f5728,f5729,f5730,f5731,f5732,f5733,f5734,f5735,f5724,f5725,f5840,f5841,f5842,f5843,f5844,f5845,f5846,f5847,f5835,f5848,f5838,f6055,f6172,f6173,f6174,f6202,f6205,f6274,f6273,f6275,f6276,f6277,f6278,f6279,f6280,f6281,f6502,f6335,f6336,f6347,f6350,f6351,f6352,f6353,f6354,f6359,f6360,f6361,f6362,f6363,f6364,f6365,f6366,f6367,f6368,f6370,f6371,f6372,f6373,f6374,f6376,f6377,f6380,f6381,f6382,f6383,f6384,f6385,f6387,f6388,f6389,f6390,f6391,f6392,f6393,f6394,f6395,f6396,f6397,f6398,f6399,f6400,f6401,f6402,f6403,f6404,f6411,f6412,f6432,f6435,f6436,f6437,f6438,f6439,f6444,f6445,f6447,f6448,f6449,f6450,f6452,f6453,f6454,f6455,f6456,f6458,f6459,f6462,f6463,f6464,f6465,f6467,f6468,f6469,f6470,f6471,f6472,f6473,f6474,f6475,f6476,f6477,f6478,f6479,f6480,f6481,f6482,f6483,f6484,f6491,f6492,f6570,f6572,f6573,f6574,f6575,f6562,f6563,f6576,f6577,f6578,f6579,f6739,f6740,f6741,f6742,f6743,f6660,f6661,f6744,f6745,f6746,f6747,f6697,f6698,f6699,f6700,f6701,f6733,f6736,f6737,f6984,f6987,f6990,f7008,f7010,f7012,f7027,f7029,f7031,f7062,f7064,f7066,f7092,f6903,f6904,f6909,f6929,f6930,f7235,f7234,f7236,f7237,f7238,f7239,f7240,f7241,f7242,f7327,f7329,f7340,f7343,f7344,f7345,f7346,f7347,f7348,f7356,f7357,f7359,f7360,f7361,f7362,f7363,f7364,f7365,f7366,f7367,f7368,f7369,f7370,f7371,f7372,f7373,f7374,f7375,f7376,f7377,f7378,f7387,f7388,f7398,f7401,f7402,f7403,f7404,f7405,f7406,f7429,f7430,f7431,f7432,f7433,f7434,f7435,f7436,f7437,f7438,f7439,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7448,f7457,f7458,f7563,f7562,f7877,f7565,f7566,f7567,f7621,f7622,f7623,f7624,f7625,f7626,f7627,f7666,f7667,f7668,f7669,f7671,f7673,f7684,f7687,f7688,f7689,f7690,f7691,f7692,f7702,f7703,f7705,f7706,f7707,f7708,f7710,f7711,f7712,f7713,f7714,f7716,f7717,f7720,f7721,f7722,f7723,f7724,f7725,f7727,f7728,f7729,f7730,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7739,f7740,f7741,f7742,f7743,f7744,f7745,f7746,f7756,f7757,f7771,f7782,f7785,f7786,f7787,f7788,f7789,f7790,f7800,f7801,f7802,f7803,f7804,f7805,f7806,f7807,f7809,f7810,f7811,f7812,f7813,f7814,f7818,f7819,f7820,f7821,f7822,f7824,f7825,f7828,f7829,f7830,f7831,f7832,f7833,f7835,f7836,f7837,f7838,f7839,f7840,f7841,f7842,f7843,f7844,f7845,f7846,f7847,f7848,f7849,f7850,f7851,f7852,f7853,f7854,f7864,f7865,f7981,f7982,f7983,f7984,f7985,f7986,f7987,f7988,f7989,f7992,f7995,f8025,f8026,f8027,f8028,f8029,f8031,f8032,f8033,f8034,f8035,f8052,f8053,f8054,f8055,f8129,f8131,f8132,f8063,f8135,f8136,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8147,f8148,f4657,f5953])).
fof(f8370,plain,(
  $false | ($spl1 | $spl4 | $spl28)),
  inference(global_subsumption,[],[f370,f260,f259,f261,f263,f264,f266,f265,f270,f269,f279,f280,f287,f286,f285,f284,f283,f289,f288,f299,f296,f301,f300,f303,f302,f307,f308,f310,f309,f312,f311,f316,f315,f318,f317,f321,f322,f323,f325,f332,f333,f334,f335,f336,f337,f340,f339,f342,f341,f343,f344,f345,f346,f347,f348,f380,f350,f351,f352,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f382,f381,f391,f397,f442,f468,f509,f546,f575,f576,f595,f611,f655,f661,f699,f700,f701,f709,f710,f702,f712,f718,f719,f720,f825,f826,f827,f828,f829,f830,f831,f803,f871,f872,f912,f913,f914,f915,f916,f891,f918,f920,f934,f935,f936,f937,f938,f930,f993,f998,f999,f1003,f1051,f1030,f1056,f1196,f1225,f1270,f1272,f1273,f1274,f1275,f1276,f1277,f1233,f1278,f1279,f1298,f1299,f1300,f1326,f1329,f326,f379,f378,f377,f376,f375,f1091,f939,f306,f305,f304,f973,f1127,f1295,f841,f1209,f331,f330,f329,f328,f327,f320,f371,f313,f372,f314,f319,f373,f374,f659,f660,f658,f968,f1122,f1265,f570,f823,f571,f569,f292,f1479,f1480,f1481,f1482,f1483,f1484,f1453,f1486,f1488,f1489,f1516,f1517,f1518,f1519,f1520,f1521,f1512,f1531,f1553,f1561,f1563,f1605,f1607,f1669,f1737,f1738,f1920,f1828,f1837,f1839,f1848,f1857,f1859,f1942,f1945,f2278,f2041,f2049,f2050,f2051,f2052,f2053,f2283,f2062,f2073,f2075,f2308,f2325,f2326,f2327,f2328,f2329,f2344,f2409,f2410,f2416,f2417,f2357,f2423,f2424,f2441,f2442,f2384,f2448,f2449,f2403,f2479,f2481,f2464,f2463,f2503,f2505,f2488,f2487,f2601,f2604,f2607,f2615,f2618,f2621,f2648,f2654,f2663,f2664,f2549,f2670,f2671,f2688,f2689,f2576,f2695,f2696,f2595,f2709,f2713,f2733,f2734,f2736,f2956,f3146,f3158,f3160,f3271,f3260,f3339,f3313,f3341,f3410,f3352,f3418,f3471,f3460,f3523,f3524,f3525,f3526,f3527,f3528,f3529,f3530,f3522,f3638,f3639,f3640,f3641,f3642,f3643,f3633,f4014,f4063,f4065,f4141,f4080,f4149,f4230,f4232,f4234,f4236,f4238,f4240,f3517,f2575,f2548,f2383,f2356,f1944,f1244,f854,f849,f811,f4416,f4418,f4420,f4422,f4424,f4426,f367,f4497,f4552,f4551,f4554,f4547,f4619,f4620,f4722,f4723,f4753,f4838,f4839,f4851,f5026,f5027,f5028,f5029,f5030,f5031,f5128,f5129,f5130,f5131,f5132,f5133,f5249,f5327,f5328,f5329,f5330,f5257,f5258,f5337,f5277,f5278,f5280,f5281,f5351,f5353,f5354,f5302,f5357,f5358,f5360,f5361,f5362,f5363,f5470,f5471,f5472,f5473,f5493,f5494,f5495,f5496,f5497,f5498,f5513,f5514,f5515,f5516,f5571,f5573,f5574,f5523,f5577,f5578,f5580,f5581,f5582,f5583,f5636,f5637,f5638,f5639,f5640,f5641,f5642,f5643,f5632,f5633,f5728,f5729,f5730,f5731,f5732,f5733,f5734,f5735,f5724,f5725,f5840,f5841,f5842,f5843,f5844,f5845,f5846,f5847,f5835,f5848,f5838,f6055,f6172,f6173,f6174,f6202,f6205,f6274,f6273,f6275,f6276,f6277,f6278,f6279,f6280,f6281,f6502,f6335,f6336,f6347,f6350,f6351,f6352,f6353,f6354,f6359,f6360,f6361,f6362,f6363,f6364,f6365,f6366,f6367,f6368,f6370,f6371,f6372,f6373,f6374,f6376,f6377,f6380,f6381,f6382,f6383,f6384,f6385,f6387,f6388,f6389,f6390,f6391,f6392,f6393,f6394,f6395,f6396,f6397,f6398,f6399,f6400,f6401,f6402,f6403,f6404,f6411,f6412,f6432,f6435,f6436,f6437,f6438,f6439,f6444,f6445,f6447,f6448,f6449,f6450,f6452,f6453,f6454,f6455,f6456,f6458,f6459,f6462,f6463,f6464,f6465,f6467,f6468,f6469,f6470,f6471,f6472,f6473,f6474,f6475,f6476,f6477,f6478,f6479,f6480,f6481,f6482,f6483,f6484,f6491,f6492,f6570,f6572,f6573,f6574,f6575,f6562,f6563,f6576,f6577,f6578,f6579,f6739,f6740,f6741,f6742,f6743,f6660,f6661,f6744,f6745,f6746,f6747,f6697,f6698,f6699,f6700,f6701,f6733,f6736,f6737,f6984,f6987,f6990,f7008,f7010,f7012,f7027,f7029,f7031,f7062,f7064,f7066,f7092,f6903,f6904,f6909,f6929,f6930,f7235,f7234,f7236,f7237,f7238,f7239,f7240,f7241,f7242,f7327,f7329,f7340,f7343,f7344,f7345,f7346,f7347,f7348,f7356,f7357,f7359,f7360,f7361,f7362,f7363,f7364,f7365,f7366,f7367,f7368,f7369,f7370,f7371,f7372,f7373,f7374,f7375,f7376,f7377,f7378,f7387,f7388,f7398,f7401,f7402,f7403,f7404,f7405,f7406,f7429,f7430,f7431,f7432,f7433,f7434,f7435,f7436,f7437,f7438,f7439,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7448,f7457,f7458,f7563,f7562,f7877,f7565,f7566,f7567,f7621,f7622,f7623,f7624,f7625,f7626,f7627,f7666,f7667,f7668,f7669,f7671,f7673,f7684,f7687,f7688,f7689,f7690,f7691,f7692,f7702,f7703,f7705,f7706,f7707,f7708,f7710,f7711,f7712,f7713,f7714,f7716,f7717,f7720,f7721,f7722,f7723,f7724,f7725,f7727,f7728,f7729,f7730,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7739,f7740,f7741,f7742,f7743,f7744,f7745,f7746,f7756,f7757,f7771,f7782,f7785,f7786,f7787,f7788,f7789,f7790,f7800,f7801,f7802,f7803,f7804,f7805,f7806,f7807,f7809,f7810,f7811,f7812,f7813,f7814,f7818,f7819,f7820,f7821,f7822,f7824,f7825,f7828,f7829,f7830,f7831,f7832,f7833,f7835,f7836,f7837,f7838,f7839,f7840,f7841,f7842,f7843,f7844,f7845,f7846,f7847,f7848,f7849,f7850,f7851,f7852,f7853,f7854,f7864,f7865,f7981,f7982,f7983,f7984,f7985,f7986,f7987,f7988,f7989,f7992,f7995,f8025,f8026,f8027,f8028,f8029,f8031,f8032,f8033,f8034,f8035,f8052,f8053,f8054,f8055,f8129,f8131,f8132,f8063,f8135,f8136,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8147,f8148,f4657,f5953])).
fof(f8369,plain,(
  $false | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f8368,f263])).
fof(f8368,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f370,f8161])).
fof(f8161,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0)))))) ) | $spl2),
  inference(forward_demodulation,[],[f7208,f7200])).
fof(f7200,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(fun(sK0,bool),bool),'!_1'),s(fun(sK0,bool),sK8('?_1',fun(sK0,bool))))) | $spl2),
  inference(unit_resulting_resolution,[],[f6086,f270])).
fof(f6086,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,bool),bool),'!_1'),s(fun(sK0,bool),sK8('?_1',fun(sK0,bool)))))) | $spl2),
  inference(unit_resulting_resolution,[],[f6056,f300])).
fof(f6056,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),sK8('?_1',fun(sK0,bool))),s(sK0,X0))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f6003,f303])).
fof(f6003,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,bool),bool),'?_1'),s(fun(sK0,bool),sK8('?_1',fun(sK0,bool)))))) | $spl2),
  inference(unit_resulting_resolution,[],[f5958,f301])).
fof(f5958,plain,(
  ~p(s(bool,happ(s(fun(fun(fun(sK0,bool),bool),bool),'!_1'),s(fun(fun(sK0,bool),bool),'?_1')))) | $spl2),
  inference(unit_resulting_resolution,[],[f4674,f300])).
fof(f4674,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(fun(sK0,bool),bool),'?_1'),s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0))))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f4330,f302])).
fof(f4330,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0))),s(sK0,X1))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f368,f289])).
fof(f7208,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0))))) = s(bool,happ(s(fun(fun(sK0,bool),bool),'!_1'),s(fun(sK0,bool),sK8('?_1',fun(sK0,bool)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f4332,f6086,f468])).
fof(f4332,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f368,f265])).
fof(f8367,plain,(
  $false | ($spl2 | $spl20)),
  inference(subsumption_resolution,[],[f8366,f263])).
fof(f8366,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f4243,f8170])).
fof(f8170,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0)))))) ) | $spl2),
  inference(forward_demodulation,[],[f7215,f7200])).
fof(f7215,plain,(
  ( ! [X0] : (s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0))))) = s(bool,happ(s(fun(fun(sK0,bool),bool),'!_1'),s(fun(sK0,bool),sK8('?_1',fun(sK0,bool)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f4380,f6086,f468])).
fof(f4380,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f4331,f265])).
fof(f4331,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0))))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f368,f283])).
fof(f8106,plain,(
  $false | ($spl0 | $spl2 | $spl19)),
  inference(subsumption_resolution,[],[f8104,f263])).
fof(f8104,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl2 | $spl19)),
  inference(backward_demodulation,[],[f7991,f4541])).
fof(f4541,plain,(
  p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK6('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | $spl0),
  inference(forward_demodulation,[],[f4522,f335])).
fof(f4522,plain,(
  p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),sK6('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f366,f288])).
fof(f7991,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),X15)))) ) | ($spl2 | $spl19)),
  inference(superposition,[],[f712,f4698])).
fof(f4698,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0))),s(sK0,X1)))) ) | ($spl2 | $spl19)),
  inference(forward_demodulation,[],[f4684,f3919])).
fof(f3919,plain,(
  s(bool,'F_0') = s(bool,sK8('_i_0_0',bool)) | $spl19),
  inference(unit_resulting_resolution,[],[f3539,f264])).
fof(f3539,plain,(
  s(bool,'T_0') != s(bool,sK8('_i_0_0',bool)) | $spl19),
  inference(unit_resulting_resolution,[],[f3538,f265])).
fof(f4684,plain,(
  ( ! [X0,X1] : (s(bool,sK8('_i_0_0',bool)) = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0))),s(sK0,X1)))) ) | ($spl2 | $spl19)),
  inference(unit_resulting_resolution,[],[f3539,f4330,f468])).
fof(f4520,plain,(
  $false | ($spl2 | $spl4 | $spl19)),
  inference(subsumption_resolution,[],[f4519,f263])).
fof(f4519,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl4 | $spl19)),
  inference(forward_demodulation,[],[f370,f4348])).
fof(f4348,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0)))))) ) | ($spl2 | $spl19)),
  inference(forward_demodulation,[],[f4341,f3919])).
fof(f4341,plain,(
  ( ! [X0] : (s(bool,sK8('_i_0_0',bool)) = s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0)))))) ) | ($spl2 | $spl19)),
  inference(unit_resulting_resolution,[],[f3539,f368,f468])).
fof(f4518,plain,(
  $false | ($spl2 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f4517,f263])).
fof(f4517,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl19 | $spl20)),
  inference(forward_demodulation,[],[f4243,f4396])).
fof(f4396,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0)))))) ) | ($spl2 | $spl19)),
  inference(forward_demodulation,[],[f4389,f3919])).
fof(f4389,plain,(
  ( ! [X0] : (s(bool,sK8('_i_0_0',bool)) = s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0)))))) ) | ($spl2 | $spl19)),
  inference(unit_resulting_resolution,[],[f3539,f4331,f468])).
fof(f4497,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl28),
  inference(cnf_transformation,[],[f4497_D])).
fof(f4497_D,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4503,plain,(
  p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK4('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | $spl32),
  inference(cnf_transformation,[],[f4503_D])).
fof(f4503_D,plain,(
  p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK4('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4501,plain,(
  p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK4('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl0),
  inference(forward_demodulation,[],[f4500,f335])).
fof(f4500,plain,(
  p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),sK4('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl0),
  inference(subsumption_resolution,[],[f4461,f261])).
fof(f4461,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),sK4('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl0),
  inference(superposition,[],[f285,f4270])).
fof(f4270,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | $spl0),
  inference(unit_resulting_resolution,[],[f261,f366,f299])).
fof(f4499,plain,(
  p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | $spl30),
  inference(cnf_transformation,[],[f4499_D])).
fof(f4499_D,plain,(
  p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4495,plain,(
  p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl0),
  inference(forward_demodulation,[],[f4494,f335])).
fof(f4494,plain,(
  p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl0),
  inference(subsumption_resolution,[],[f4460,f261])).
fof(f4460,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | p(s(bool,'const.bool.?|33|_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl0),
  inference(superposition,[],[f286,f4270])).
fof(f4249,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),'_i_0_0')))) | $spl24),
  inference(cnf_transformation,[],[f4249_D])).
fof(f4249_D,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4251,plain,(
  p(s(bool,sK4('_i_0_0',bool))) | $spl26),
  inference(cnf_transformation,[],[f4251_D])).
fof(f4251_D,plain,(
  p(s(bool,sK4('_i_0_0',bool))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4247,plain,(
  p(s(bool,sK4('_i_0_0',bool))) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(forward_demodulation,[],[f4246,f279])).
fof(f4246,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK4('_i_0_0',bool))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(subsumption_resolution,[],[f4223,f261])).
fof(f4223,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK4('_i_0_0',bool))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(superposition,[],[f285,f623])).
fof(f623,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl6),
  inference(unit_resulting_resolution,[],[f616,f266])).
fof(f4243,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)))))) | $spl20),
  inference(cnf_transformation,[],[f4243_D])).
fof(f4243_D,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4245,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3))),s(sK0,sK4(happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)),sK0))))) | $spl22),
  inference(cnf_transformation,[],[f4245_D])).
fof(f4245_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3))),s(sK0,sK4(happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)),sK0))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4241,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3))),s(sK0,sK4(happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)),sK0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)))))) | $spl4),
  inference(subsumption_resolution,[],[f4222,f261])).
fof(f4222,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3))),s(sK0,sK4(happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)),sK0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)))))) | $spl4),
  inference(superposition,[],[f285,f385])).
fof(f385,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3))))) | $spl4),
  inference(unit_resulting_resolution,[],[f370,f266])).
fof(f3536,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | $spl16),
  inference(cnf_transformation,[],[f3536_D])).
fof(f3536_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3538,plain,(
  ~p(s(bool,sK8('_i_0_0',bool))) | $spl19),
  inference(cnf_transformation,[],[f3538_D])).
fof(f3538_D,plain,(
  ~p(s(bool,sK8('_i_0_0',bool))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f3522,plain,(
  ~p(s(bool,sK8('_i_0_0',bool))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),
  inference(superposition,[],[f301,f279])).
fof(f2710,plain,(
  ( ! [X4,X2,X3] : (s(X2,X3) != s(X2,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),'=_0'),s(X2,X3))),s(X2,X4)))) ) | $spl12),
  inference(cnf_transformation,[],[f2710_D])).
fof(f2710_D,plain,(
  ( ! [X4,X2,X3] : (s(X2,X3) != s(X2,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),'=_0'),s(X2,X3))),s(X2,X4)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2711,plain,(
  ( ! [X5] : (p(s(bool,X5))) ) | $spl14),
  inference(cnf_transformation,[],[f2711_D])).
fof(f2711_D,plain,(
  ( ! [X5] : (p(s(bool,X5))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2709,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,X5)) | s(X2,X3) != s(X2,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),'=_0'),s(X2,X3))),s(X2,X4)))) ) | $spl1),
  inference(subsumption_resolution,[],[f2708,f265])).
fof(f2708,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,X5)) | s(X2,X3) != s(X2,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),'=_0'),s(X2,X3))),s(X2,X4))) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(forward_demodulation,[],[f2596,f803])).
fof(f2596,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (p(s(bool,sK28('const.pair.,_2'(s(bool,X5),s(X6,X7)),X6,bool))) | s(X2,X3) != s(X2,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),'=_0'),s(X2,X3))),s(X2,X4))) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(superposition,[],[f331,f1051])).
fof(f1444,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1341,f1432])).
fof(f1432,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1413,f617])).
fof(f1413,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | s(X0,X1) != s(X0,X2)) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f1371,f331])).
fof(f1371,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) ) | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f382,f617,f468])).
fof(f1341,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f617,f270])).
fof(f1443,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1343,f1432])).
fof(f1343,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f263,f617,f296])).
fof(f1442,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1344,f1432])).
fof(f1344,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f617,f617,f296])).
fof(f1441,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1349,f263])).
fof(f1349,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f381,f617,f296])).
fof(f1440,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1351,f617])).
fof(f1351,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f382,f617,f296])).
fof(f1439,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1353,f1432])).
fof(f1353,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f263,f617,f296])).
fof(f1438,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1354,f1432])).
fof(f1354,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f617,f617,f296])).
fof(f1437,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1359,f263])).
fof(f1359,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f381,f617,f296])).
fof(f1436,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1361,f617])).
fof(f1361,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f382,f617,f296])).
fof(f1435,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1367,f381])).
fof(f1367,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f381,f617,f468])).
fof(f1434,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1369,f1432])).
fof(f1369,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f381,f617,f468])).
fof(f1433,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f1429,f617])).
fof(f1429,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f1371,f326])).
fof(f1372,plain,(
  $false | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f382,f382,f617,f468])).
fof(f1370,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f381,f381,f617,f468])).
fof(f1362,plain,(
  $false | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f617,f382,f617,f296])).
fof(f1360,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f263,f381,f617,f296])).
fof(f1358,plain,(
  $false | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f617,f382,f617,f296])).
fof(f1356,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f617,f381,f617,f296])).
fof(f1352,plain,(
  $false | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f617,f382,f617,f296])).
fof(f1350,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f263,f381,f617,f296])).
fof(f1348,plain,(
  $false | ($spl1 | $spl8)),
  inference(unit_resulting_resolution,[],[f617,f382,f617,f296])).
fof(f1346,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f617,f381,f617,f296])).
fof(f1342,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f381,f617,f270])).
fof(f1340,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f385,f617,f265])).
fof(f1338,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f261,f617])).
fof(f1337,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | $spl10),
  inference(cnf_transformation,[],[f1337_D])).
fof(f1337_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f617,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl8),
  inference(cnf_transformation,[],[f617_D])).
fof(f617_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1329,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f303,f279])).
fof(f616,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(cnf_transformation,[],[f616_D])).
fof(f616_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f611,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f289,f279])).
fof(f366,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl0),
  inference(cnf_transformation,[],[f366_D])).
fof(f366_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f370,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)))))) | $spl4),
  inference(cnf_transformation,[],[f370_D])).
fof(f370_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f259,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))))),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3))))))) & (~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | ! [X3] : ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f199])).
fof(f199,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) | ? [X4] : p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))))))) & (~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) | ! [X3] : ~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))))))))),
  inference(rectify,[],[f198])).
fof(f198,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) | ? [X3] : p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))))))) & (~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) | ! [X3] : ~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))))))))),
  inference(nnf_transformation,[],[f188])).
fof(f188,plain,(
  ? [X0,X1,X2] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) <~> ? [X3] : p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3)))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) <=> ? [X3] : p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3)))))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1] : ! [X2] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) <=> ? [X3] : p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X7,X6),bool),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,bool)),X8)))))) <=> ? [X9] : p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X7,X6),bool),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,bool)),X8)))))) <=> ? [X9] : p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,X9)))))))),
  file('thm.pair.ELIM_PEXISTS_EVAL','thm.pair.ELIM_PEXISTS_EVAL')).
fof(f367,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl1),
  inference(cnf_transformation,[],[f367_D])).
fof(f367_D,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f368,plain,(
  ( ! [X3] : (~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X3))))))) ) | $spl2),
  inference(cnf_transformation,[],[f368_D])).
fof(f368_D,plain,(
  ( ! [X3] : (~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X3))))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f260,plain,(
  ( ! [X3] : (~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X3)))))) | ~p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))))) )),
  inference(cnf_transformation,[],[f200])).
% SZS output end Proof for thm.pair
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 166
% Passive clauses: 1133
% Generated clauses: 9348
% Final active clauses: 67
% Final passive clauses: 660
% Input formulas: 68
% Initial clauses: 104
% 
% Duplicate literals: 76
% Trivial inequalities: 33
% Fw subsumption resolutions: 240
% Fw demodulations: 2434
% Bw demodulations: 131
% Condensations: 152
% Global subsumptions: 4
% 
% Simple tautologies: 35
% Equational tautologies: 75
% Forward subsumptions: 5099
% Fw demodulations to eq. taut.: 20
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 1037
% Factoring: 1
% Forward superposition: 2202
% Backward superposition: 2439
% Self superposition: 67
% Equality factoring: 10
% Equality resolution: 7
% 
% Split clauses: 10
% Split components: 20
% SAT solver clauses: 1196
% SAT solver unit clauses: 168
% SAT solver binary clauses: 743
% 
% Sat splits: 10
% Sat splitting refutations: 37
% 
% Memory used [KB]: 12153
% Time elapsed: 0.455 s
% ------------------------------
% ------------------------------
% Success in time 1.741 s
