Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numRing
% SZS output start Proof for thm.numRing
fof(f3013,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3012,f2943,f177,f275,f273,f176,f274,f272])).
fof(f272,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK1))) | $spl0),
  inference(cnf_transformation,[],[f272_D])).
fof(f272_D,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f274,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl2),
  inference(cnf_transformation,[],[f274_D])).
fof(f274_D,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f176,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) = s('type.num.num',sK1)) & (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) != s('type.num.num',sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f160])).
fof(f160,plain,(
  ? [X0,X1] : ((s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ? [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ~! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.numRing.num_rewrites|split|43','thm.numRing.num_rewrites|split|43')).
fof(f273,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK1))) | $spl1),
  inference(cnf_transformation,[],[f273_D])).
fof(f273_D,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f275,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | $spl3),
  inference(cnf_transformation,[],[f275_D])).
fof(f275_D,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f177,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f161])).
fof(f2943,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2898,f275])).
fof(f2898,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(equality_resolution,[],[f2378])).
fof(f2378,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) | s('type.num.num',sK1) = s('type.num.num',X0)) ) | $spl0),
  inference(superposition,[],[f266,f272])).
fof(f266,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X10,X7] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10))) <=> s('type.num.num',X7) = s('type.num.num',X10))),
  file('thm.numRing.num_rewrites|split|43','thm.numeral.numeral_eq|split|7')).
fof(f3012,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f3011])).
fof(f3011,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f273,f274])).
% SZS output end Proof for thm.numRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 94
% Passive clauses: 326
% Generated clauses: 2810
% Final active clauses: 71
% Final passive clauses: 180
% Input formulas: 81
% Initial clauses: 95
% 
% Duplicate literals: 1
% Trivial inequalities: 14
% Fw subsumption resolutions: 29
% Fw demodulations: 502
% Bw demodulations: 11
% 
% Simple tautologies: 8
% Equational tautologies: 9
% Forward subsumptions: 1923
% Fw demodulations to eq. taut.: 8
% 
% Binary resolution: 8
% Forward superposition: 1216
% Backward superposition: 886
% Self superposition: 26
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 22
% Split components: 44
% SAT solver clauses: 28
% SAT solver unit clauses: 2
% SAT solver binary clauses: 18
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 22
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 26
% TWLsolver calls for satisfiability: 5
% 
% Memory used [KB]: 2046
% Time elapsed: 0.061 s
% ------------------------------
% ------------------------------
% Success in time 0.065 s
