Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.rat
% (70360)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 723
% Passive clauses: 8878
% Generated clauses: 21031
% Final active clauses: 714
% Final passive clauses: 7713
% Input formulas: 54
% Initial clauses: 210
% 
% Duplicate literals: 464
% Trivial inequalities: 94
% Fw subsumption resolutions: 2196
% Fw demodulations: 563
% Bw demodulations: 7
% 
% Simple tautologies: 293
% Equational tautologies: 61
% Forward subsumptions: 7806
% 
% Binary resolution: 10493
% Factoring: 1276
% Forward superposition: 3093
% Backward superposition: 2318
% Self superposition: 33
% Equality factoring: 30
% Equality resolution: 32
% 
% Split clauses: 590
% Split components: 1187
% SAT solver clauses: 1051
% SAT solver unit clauses: 30
% SAT solver binary clauses: 201
% 
% Sat splits: 590
% 
% TWLsolver clauses: 1051
% TWLsolver calls for satisfiability: 405
% 
% Memory used [KB]: 9083
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (70377)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 393
% Passive clauses: 2449
% Generated clauses: 10556
% Final active clauses: 385
% Final passive clauses: 2036
% Input formulas: 54
% Initial clauses: 210
% 
% Duplicate literals: 245
% Trivial inequalities: 20
% Fw subsumption resolutions: 845
% Fw demodulations: 199
% Bw demodulations: 23
% Condensations: 265
% Global subsumptions: 22
% 
% Simple tautologies: 257
% Equational tautologies: 21
% Forward subsumptions: 3363
% 
% Binary resolution: 3616
% Factoring: 350
% Forward superposition: 1174
% Backward superposition: 1063
% Self superposition: 32
% Equality factoring: 26
% Equality resolution: 6
% 
% Split clauses: 25
% Split components: 55
% SAT solver clauses: 3079
% SAT solver unit clauses: 215
% SAT solver binary clauses: 93
% SAT solver learnt clauses: 45
% SAT solver learnt literals: 11
% 
% Sat splits: 25
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2990
% TWLsolver calls for satisfiability: 1753
% 
% Memory used [KB]: 12792
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f8946,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f276,f500,f498,f275,f499,f497,f504,f1775,f1779,f1783,f1786,f1789,f1792,f1795,f1798,f1801,f1804,f1807,f1810,f1813,f1816,f1819,f1822,f1825,f1828,f1831,f1834,f1869,f1776,f1871,f1780,f1874,f1784,f1876,f1787,f1878,f1790,f1880,f1793,f1882,f1796,f1884,f1799,f1886,f1802,f1888,f1805,f1890,f1808,f1892,f1811,f1894,f1814,f1896,f1817,f1898,f1820,f1900,f1823,f1902,f1826,f1904,f1829,f1906,f1832,f1908,f1777,f1835,f1939,f1940,f1944,f2413,f2414,f2421,f2418,f2415,f2416,f2424,f2417,f2426,f2419,f2423,f2443,f2444,f2456,f2454,f3051,f3054,f3052,f3056,f3057,f3058,f3060,f3065,f3067,f3071,f3073,f3075,f3077,f3085,f3087,f3095,f3097,f3104,f3111,f3123,f3140,f3168,f3170,f3222,f3226,f3246,f3253,f3261,f3268,f3270,f3277,f3283,f3290,f3296,f3298,f3300,f3302,f3304,f3875,f3876,f3317,f3319,f3321,f3323,f3340,f3342,f3348,f3354,f3363,f3370,f3372,f3383,f3391,f3393,f3402,f3408,f3415,f3421,f3423,f3433,f3435,f3441,f3443,f3454,f3456,f3462,f3464,f3475,f3477,f3483,f3485,f3492,f3498,f3505,f3511,f3520,f3529,f3539,f3551,f3554,f3562,f3564,f3566,f3571,f3573,f3578,f3580,f3585,f3587,f3592,f3594,f3604,f3612,f3621,f3628,f3630,f3646,f3648,f3650,f3652,f3662,f3664,f3674,f3676,f3687,f3689,f3700,f3702,f3704,f3710,f3715,f3720,f3722,f3727,f3729,f3737,f3745,f3754,f3762,f3767,f3769,f3771,f3773,f3778,f3780,f3782,f3784,f3786,f3794,f3799,f3804,f3807,f3812,f3815,f3822,f3827,f3832,f3837,f3839,f3841,f3846,f3848,f3853,f3877,f3878,f3879,f3880,f3881,f3882,f3883,f3884,f3885,f3886,f3887,f3888,f3889,f3890,f3891,f3892,f3893,f3894,f3895,f3896,f3897,f3898,f3899,f3900,f3901,f3902,f3903,f3904,f3905,f3906,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f3947,f3948,f3949,f3950,f3951,f3952,f3953,f3954,f3956,f3958,f3960,f3962,f3963,f3964,f3965,f3966,f3967,f3968,f3972,f3973,f3974,f3975,f3976,f3977,f3978,f3979,f3980,f3981,f3982,f3983,f3984,f3985,f3986,f3987,f3988,f3989,f3990,f3991,f3992,f3993,f3994,f3995,f3996,f3997,f3998,f3999,f4000,f4001,f4002,f4003,f4004,f4005,f4034,f4036,f4035,f4038,f4039,f4316,f4319,f4317,f5070,f5072,f5071,f8471,f8473,f8475,f8477,f8496,f8498,f8502,f8503,f8507,f8508,f8586,f8472,f8587,f8474,f8588,f8476,f8589,f8478,f8607,f8497,f8608,f8499,f8908,f8913,f3053,f8909,f8941,f8930])).
fof(f8930,plain,(
  $false | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f497,f8919])).
fof(f8919,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',X0))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',X0)))) ) | $spl3),
  inference(superposition,[],[f8885,f358])).
fof(f358,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X20,X21] : s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X21),s('type.rat.rat',X20)))),
  file('thm.rat.RAT_EQ_RADD','thm.rat.RAT_ADD_COMM')).
fof(f8885,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',sK12))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',X0)))) ) | $spl3),
  inference(superposition,[],[f8838,f358])).
fof(f8838,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',sK12))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',sK13)))) ) | $spl3),
  inference(forward_demodulation,[],[f8837,f359])).
fof(f359,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',X1)))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',X1)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X25,X26] : s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X25),s('type.rat.rat',X26))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X25))),s('type.rat.rat',X26)))),
  file('thm.rat.RAT_EQ_RADD','ARITY_1_2_const.rat.rat_add')).
fof(f8837,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',sK13))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',sK12)))) ) | $spl3),
  inference(forward_demodulation,[],[f8816,f359])).
fof(f8816,plain,(
  ( ! [X0] : (s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',sK12))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',sK13)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f500,f1332])).
fof(f1332,plain,(
  ( ! [X2,X0,X1] : (s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',X1))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',X2))) | s('type.rat.rat',X1) = s('type.rat.rat',X2)) )),
  inference(subsumption_resolution,[],[f1330,f277])).
fof(f277,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.RAT_EQ_RADD',aHL_TRUTH)).
fof(f1330,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',X1))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',X2))) | s('type.rat.rat',X1) = s('type.rat.rat',X2)) )),
  inference(superposition,[],[f401,f547])).
fof(f547,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f284,f298])).
fof(f298,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.rat.RAT_EQ_RADD','thm.bool.EQ_CLAUSES|split|0')).
fof(f284,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6] : p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6))))))),
  file('thm.rat.RAT_EQ_RADD','thm.rat.RAT_ADD_ONE_ONE')).
fof(f401,plain,(
  ( ! [X6,X2,X0,X5,X1] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X5,X6] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))) & ((s(X0,happ(s(fun(X1,X0),X2),s(X1,sK17(X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK18(X2,X1,X0)))) & s(X1,sK17(X2,X1,X0)) != s(X1,sK18(X2,X1,X0))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17,sK18])],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X5,X6] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))) & (? [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) & s(X1,X3) != s(X1,X4)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(rectify,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) | s(X1,X3) = s(X1,X4))) & (? [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) & s(X1,X3) != s(X1,X4)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) | s(X1,X3) = s(X1,X4)))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) => s(X1,X3) = s(X1,X4)))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ! [X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) => s(X1,X3) = s(X1,X4)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X15,X12] : ! [X17] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X12,X15),X17)))) <=> ! [X18,X19] : (s(X15,happ(s(fun(X12,X15),X17),s(X12,X18))) = s(X15,happ(s(fun(X12,X15),X17),s(X12,X19))) => s(X12,X18) = s(X12,X19)))),
  file('thm.rat.RAT_EQ_RADD','thm.bool.ONE_ONE_THM')).
fof(f8941,plain,(
  $false | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f8938])).
fof(f8938,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) | ($spl0 | $spl3)),
  inference(superposition,[],[f8919,f497])).
fof(f8909,plain,(
  ( ! [X35,X31,X36,X34,X32] : (s(bool,X34) != s(bool,sK17(X32,bool,X31)) | s(bool,'T_0') = s(bool,X36) | s(bool,X35) = s(bool,X36) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,'T_0') = s(bool,sK18(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34)) ) | $spl94),
  inference(cnf_transformation,[],[f8909_D])).
fof(f8909_D,plain,(
  ( ! [X35,X31,X36,X34,X32] : (s(bool,X34) != s(bool,sK17(X32,bool,X31)) | s(bool,'T_0') = s(bool,X36) | s(bool,X35) = s(bool,X36) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,'T_0') = s(bool,sK18(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f3053,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) | $spl64),
  inference(cnf_transformation,[],[f3053_D])).
fof(f3053_D,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f8913,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (p(s(bool,X57)) | s(bool,X58) != s(bool,sK17(X56,bool,X55)) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,sK18(X56,bool,X55)) | s(bool,X59) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X55),X56))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f8903,f297])).
fof(f297,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f178])).
fof(f8903,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (p(s(bool,X57)) | s(bool,X58) != s(bool,sK17(X56,bool,X55)) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,sK18(X56,bool,X55)) | s(bool,X59) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X55),X56))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f1437,f1372])).
fof(f1372,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f653,f653])).
fof(f653,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f286,f286])).
fof(f286,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rat.RAT_EQ_RADD',aHL_BOOL_CASES)).
fof(f1437,plain,(
  ( ! [X156,X157,X155] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X156),X155)))) | s(bool,X157) != s(bool,sK17(X155,bool,X156)) | s(bool,'T_0') = s(bool,X157) | s(bool,'T_0') = s(bool,sK18(X155,bool,X156))) )),
  inference(superposition,[],[f403,f653])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (s(X1,sK17(X2,X1,X0)) != s(X1,sK18(X2,X1,X0)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f8908,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,X33)) | s(bool,X34) != s(bool,sK17(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,sK18(X32,bool,X31)) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f8899,f297])).
fof(f8899,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,X33)) | s(bool,X34) != s(bool,sK17(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,sK18(X32,bool,X31)) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,'T_0') = s(bool,X33) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f1437,f1372])).
fof(f8499,plain,(
  ( ! [X565,X563,X566,X562] : (s(bool,'T_0') = s(bool,sK18(X563,bool,X562)) | s(bool,'T_0') = s(bool,X566) | s(bool,X565) = s(bool,X566) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,'T_0') = s(bool,sK17(X563,bool,X562))) ) | $spl88),
  inference(cnf_transformation,[],[f8499_D])).
fof(f8499_D,plain,(
  ( ! [X565,X563,X566,X562] : (s(bool,'T_0') = s(bool,sK18(X563,bool,X562)) | s(bool,'T_0') = s(bool,X566) | s(bool,X565) = s(bool,X566) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,'T_0') = s(bool,sK17(X563,bool,X562))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f8608,plain,(
  ( ! [X565,X563,X564,X566,X562] : (p(s(bool,X564)) | s(bool,'T_0') = s(bool,sK18(X563,bool,X562)) | s(bool,'T_0') = s(bool,sK17(X563,bool,X562)) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,X565) = s(bool,X566) | s(bool,'T_0') = s(bool,X566)) )),
  inference(subsumption_resolution,[],[f6344,f297])).
fof(f6344,plain,(
  ( ! [X565,X563,X564,X566,X562] : (p(s(bool,X564)) | s(bool,'T_0') = s(bool,sK18(X563,bool,X562)) | s(bool,'T_0') = s(bool,sK17(X563,bool,X562)) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,X565) = s(bool,X566) | s(bool,'T_0') = s(bool,X566) | s(bool,'T_0') = s(bool,X564)) )),
  inference(superposition,[],[f2340,f1372])).
fof(f2340,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X1),X0)))) | s(bool,'T_0') = s(bool,sK18(X0,bool,X1)) | s(bool,'T_0') = s(bool,sK17(X0,bool,X1))) )),
  inference(condensation,[],[f2339])).
fof(f2339,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X7),X6)))) | s(bool,'T_0') = s(bool,sK18(X6,bool,X7)) | s(bool,'T_0') = s(bool,sK17(X6,bool,X7)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f2335,f286])).
fof(f2335,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X7),X6)))) | s(bool,'T_0') = s(bool,sK18(X6,bool,X7)) | s(bool,'T_0') = s(bool,sK17(X6,bool,X7)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1174,f653])).
fof(f1174,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,sK17(X0,bool,X1)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X1),X0)))) | s(bool,'T_0') = s(bool,sK18(X0,bool,X1))) )),
  inference(superposition,[],[f403,f286])).
fof(f8497,plain,(
  ( ! [X561,X557,X560,X556,X558] : (s(X557,happ(s(fun(X556,X557),X558),s(X556,sK17(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK18(X558,X556,X557)))) | s(bool,'T_0') = s(bool,X561) | s(bool,X560) = s(bool,X561) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558)))) ) | $spl86),
  inference(cnf_transformation,[],[f8497_D])).
fof(f8497_D,plain,(
  ( ! [X561,X557,X560,X556,X558] : (s(X557,happ(s(fun(X556,X557),X558),s(X556,sK17(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK18(X558,X556,X557)))) | s(bool,'T_0') = s(bool,X561) | s(bool,X560) = s(bool,X561) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f8607,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,X559)) | s(X557,happ(s(fun(X556,X557),X558),s(X556,sK17(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK18(X558,X556,X557)))) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558))) | s(bool,X560) = s(bool,X561) | s(bool,'T_0') = s(bool,X561)) )),
  inference(subsumption_resolution,[],[f6343,f297])).
fof(f6343,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,X559)) | s(X557,happ(s(fun(X556,X557),X558),s(X556,sK17(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK18(X558,X556,X557)))) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558))) | s(bool,X560) = s(bool,X561) | s(bool,'T_0') = s(bool,X561) | s(bool,'T_0') = s(bool,X559)) )),
  inference(superposition,[],[f402,f1372])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,sK17(X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK18(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f8478,plain,(
  ( ! [X231,X235,X233,X236,X234] : (sP1(X231,X233) | s(bool,'T_0') = s(bool,X236) | s(bool,X235) = s(bool,X236) | s(bool,X231) = s(bool,X235) | s(bool,'T_0') = s(bool,X234) | s(bool,X233) = s(bool,X234)) ) | $spl84),
  inference(cnf_transformation,[],[f8478_D])).
fof(f8478_D,plain,(
  ( ! [X231,X235,X233,X236,X234] : (sP1(X231,X233) | s(bool,'T_0') = s(bool,X236) | s(bool,X235) = s(bool,X236) | s(bool,X231) = s(bool,X235) | s(bool,'T_0') = s(bool,X234) | s(bool,X233) = s(bool,X234)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f8589,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,X232)) | sP1(X231,X233) | s(bool,X233) = s(bool,X234) | s(bool,'T_0') = s(bool,X234) | s(bool,X231) = s(bool,X235) | s(bool,X235) = s(bool,X236) | s(bool,'T_0') = s(bool,X236)) )),
  inference(subsumption_resolution,[],[f6261,f297])).
fof(f6261,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,X232)) | sP1(X231,X233) | s(bool,X233) = s(bool,X234) | s(bool,'T_0') = s(bool,X234) | s(bool,X231) = s(bool,X235) | s(bool,X235) = s(bool,X236) | s(bool,'T_0') = s(bool,X236) | s(bool,'T_0') = s(bool,X232)) )),
  inference(superposition,[],[f1841,f1372])).
fof(f1841,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f1530,f277])).
fof(f1530,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f388,f653])).
fof(f388,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f227])).
fof(f227,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f8476,plain,(
  ( ! [X229,X227,X225,X230,X228] : (sP1(X227,X225) | s(bool,'T_0') = s(bool,X230) | s(bool,X229) = s(bool,X230) | s(bool,X225) = s(bool,X229) | s(bool,'T_0') = s(bool,X228) | s(bool,X227) = s(bool,X228)) ) | $spl82),
  inference(cnf_transformation,[],[f8476_D])).
fof(f8476_D,plain,(
  ( ! [X229,X227,X225,X230,X228] : (sP1(X227,X225) | s(bool,'T_0') = s(bool,X230) | s(bool,X229) = s(bool,X230) | s(bool,X225) = s(bool,X229) | s(bool,'T_0') = s(bool,X228) | s(bool,X227) = s(bool,X228)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f8588,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (p(s(bool,X226)) | sP1(X227,X225) | s(bool,X227) = s(bool,X228) | s(bool,'T_0') = s(bool,X228) | s(bool,X225) = s(bool,X229) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230)) )),
  inference(subsumption_resolution,[],[f6260,f297])).
fof(f6260,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (p(s(bool,X226)) | sP1(X227,X225) | s(bool,X227) = s(bool,X228) | s(bool,'T_0') = s(bool,X228) | s(bool,X225) = s(bool,X229) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230) | s(bool,'T_0') = s(bool,X226)) )),
  inference(superposition,[],[f1840,f1372])).
fof(f1840,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f1529,f277])).
fof(f1529,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f385,f653])).
fof(f385,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f228])).
fof(f8474,plain,(
  ( ! [X212,X210,X213] : (sP1(X210,'T_0') | s(bool,'T_0') = s(bool,X213) | s(bool,X212) = s(bool,X213) | s(bool,X210) = s(bool,X212)) ) | $spl80),
  inference(cnf_transformation,[],[f8474_D])).
fof(f8474_D,plain,(
  ( ! [X212,X210,X213] : (sP1(X210,'T_0') | s(bool,'T_0') = s(bool,X213) | s(bool,X212) = s(bool,X213) | s(bool,X210) = s(bool,X212)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f8587,plain,(
  ( ! [X212,X210,X213,X211] : (p(s(bool,X211)) | sP1(X210,'T_0') | s(bool,X210) = s(bool,X212) | s(bool,X212) = s(bool,X213) | s(bool,'T_0') = s(bool,X213)) ) | $spl8),
  inference(subsumption_resolution,[],[f6257,f297])).
fof(f6257,plain,(
  ( ! [X212,X210,X213,X211] : (p(s(bool,X211)) | sP1(X210,'T_0') | s(bool,X210) = s(bool,X212) | s(bool,X212) = s(bool,X213) | s(bool,'T_0') = s(bool,X213) | s(bool,'T_0') = s(bool,X211)) ) | $spl8),
  inference(superposition,[],[f1780,f1372])).
fof(f8472,plain,(
  ( ! [X208,X206,X209] : (sP1('T_0',X206) | s(bool,'T_0') = s(bool,X209) | s(bool,X208) = s(bool,X209) | s(bool,X206) = s(bool,X208)) ) | $spl78),
  inference(cnf_transformation,[],[f8472_D])).
fof(f8472_D,plain,(
  ( ! [X208,X206,X209] : (sP1('T_0',X206) | s(bool,'T_0') = s(bool,X209) | s(bool,X208) = s(bool,X209) | s(bool,X206) = s(bool,X208)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f8586,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X207)) | sP1('T_0',X206) | s(bool,X206) = s(bool,X208) | s(bool,X208) = s(bool,X209) | s(bool,'T_0') = s(bool,X209)) ) | $spl4),
  inference(subsumption_resolution,[],[f6256,f297])).
fof(f6256,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X207)) | sP1('T_0',X206) | s(bool,X206) = s(bool,X208) | s(bool,X208) = s(bool,X209) | s(bool,'T_0') = s(bool,X209) | s(bool,'T_0') = s(bool,X207)) ) | $spl4),
  inference(superposition,[],[f1776,f1372])).
fof(f8508,plain,(
  ( ! [X204,X202,X205,X203,X201] : (s(bool,X201) != s(bool,X202) | s(bool,'T_0') = s(bool,X205) | s(bool,X204) = s(bool,X205) | s(bool,X201) = s(bool,X204) | s(bool,'T_0') = s(bool,X203) | s(bool,X202) = s(bool,X203)) ) | $spl92),
  inference(cnf_transformation,[],[f8508_D])).
fof(f8508_D,plain,(
  ( ! [X204,X202,X205,X203,X201] : (s(bool,X201) != s(bool,X202) | s(bool,'T_0') = s(bool,X205) | s(bool,X204) = s(bool,X205) | s(bool,X201) = s(bool,X204) | s(bool,'T_0') = s(bool,X203) | s(bool,X202) = s(bool,X203)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f8507,plain,(
  ( ! [X204,X202,X200,X205,X203,X201] : (s(bool,X201) != s(bool,X202) | p(s(bool,X200)) | s(bool,X202) = s(bool,X203) | s(bool,'T_0') = s(bool,X203) | s(bool,X201) = s(bool,X204) | s(bool,X204) = s(bool,X205) | s(bool,'T_0') = s(bool,X205)) )),
  inference(subsumption_resolution,[],[f5880,f297])).
fof(f5880,plain,(
  ( ! [X204,X202,X200,X205,X203,X201] : (s(bool,X201) != s(bool,X202) | p(s(bool,X200)) | s(bool,X202) = s(bool,X203) | s(bool,'T_0') = s(bool,X203) | s(bool,X201) = s(bool,X204) | s(bool,'T_0') = s(bool,X200) | s(bool,X204) = s(bool,X205) | s(bool,'T_0') = s(bool,X205)) )),
  inference(superposition,[],[f1486,f1372])).
fof(f1486,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f297,f653])).
fof(f8503,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,X6)) ) | $spl90),
  inference(cnf_transformation,[],[f8503_D])).
fof(f8503_D,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,X6)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f8502,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f5844,f297])).
fof(f5844,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X4) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f297,f1372])).
fof(f8498,plain,(
  ( ! [X565,X563,X564,X566,X562] : (p(s(bool,X564)) | s(bool,'T_0') = s(bool,sK18(X563,bool,X562)) | s(bool,'T_0') = s(bool,sK17(X563,bool,X562)) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,X565) = s(bool,X566) | s(bool,'T_0') = s(bool,X566)) )),
  inference(subsumption_resolution,[],[f5831,f297])).
fof(f5831,plain,(
  ( ! [X565,X563,X564,X566,X562] : (p(s(bool,X564)) | s(bool,'T_0') = s(bool,sK18(X563,bool,X562)) | s(bool,'T_0') = s(bool,sK17(X563,bool,X562)) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,'T_0') = s(bool,X564) | s(bool,X565) = s(bool,X566) | s(bool,'T_0') = s(bool,X566)) )),
  inference(superposition,[],[f2340,f1372])).
fof(f8496,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,X559)) | s(X557,happ(s(fun(X556,X557),X558),s(X556,sK17(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK18(X558,X556,X557)))) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558))) | s(bool,X560) = s(bool,X561) | s(bool,'T_0') = s(bool,X561)) )),
  inference(subsumption_resolution,[],[f5830,f297])).
fof(f5830,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,X559)) | s(X557,happ(s(fun(X556,X557),X558),s(X556,sK17(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK18(X558,X556,X557)))) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558))) | s(bool,'T_0') = s(bool,X559) | s(bool,X560) = s(bool,X561) | s(bool,'T_0') = s(bool,X561)) )),
  inference(superposition,[],[f402,f1372])).
fof(f8477,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,X232)) | sP1(X231,X233) | s(bool,X233) = s(bool,X234) | s(bool,'T_0') = s(bool,X234) | s(bool,X231) = s(bool,X235) | s(bool,X235) = s(bool,X236) | s(bool,'T_0') = s(bool,X236)) )),
  inference(subsumption_resolution,[],[f5748,f297])).
fof(f5748,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,X232)) | sP1(X231,X233) | s(bool,X233) = s(bool,X234) | s(bool,'T_0') = s(bool,X234) | s(bool,X231) = s(bool,X235) | s(bool,'T_0') = s(bool,X232) | s(bool,X235) = s(bool,X236) | s(bool,'T_0') = s(bool,X236)) )),
  inference(superposition,[],[f1841,f1372])).
fof(f8475,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (p(s(bool,X226)) | sP1(X227,X225) | s(bool,X227) = s(bool,X228) | s(bool,'T_0') = s(bool,X228) | s(bool,X225) = s(bool,X229) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230)) )),
  inference(subsumption_resolution,[],[f5747,f297])).
fof(f5747,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (p(s(bool,X226)) | sP1(X227,X225) | s(bool,X227) = s(bool,X228) | s(bool,'T_0') = s(bool,X228) | s(bool,X225) = s(bool,X229) | s(bool,'T_0') = s(bool,X226) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230)) )),
  inference(superposition,[],[f1840,f1372])).
fof(f8473,plain,(
  ( ! [X212,X210,X213,X211] : (p(s(bool,X211)) | sP1(X210,'T_0') | s(bool,X210) = s(bool,X212) | s(bool,X212) = s(bool,X213) | s(bool,'T_0') = s(bool,X213)) ) | $spl8),
  inference(subsumption_resolution,[],[f5744,f297])).
fof(f5744,plain,(
  ( ! [X212,X210,X213,X211] : (p(s(bool,X211)) | sP1(X210,'T_0') | s(bool,X210) = s(bool,X212) | s(bool,'T_0') = s(bool,X211) | s(bool,X212) = s(bool,X213) | s(bool,'T_0') = s(bool,X213)) ) | $spl8),
  inference(superposition,[],[f1780,f1372])).
fof(f8471,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X207)) | sP1('T_0',X206) | s(bool,X206) = s(bool,X208) | s(bool,X208) = s(bool,X209) | s(bool,'T_0') = s(bool,X209)) ) | $spl4),
  inference(subsumption_resolution,[],[f5743,f297])).
fof(f5743,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X207)) | sP1('T_0',X206) | s(bool,X206) = s(bool,X208) | s(bool,'T_0') = s(bool,X207) | s(bool,X208) = s(bool,X209) | s(bool,'T_0') = s(bool,X209)) ) | $spl4),
  inference(superposition,[],[f1776,f1372])).
fof(f5071,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,sK18(X7,bool,X6)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X6),X7))) | s(bool,'T_0') = s(bool,sK17(X7,bool,X6))) ) | $spl76),
  inference(cnf_transformation,[],[f5071_D])).
fof(f5071_D,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,sK18(X7,bool,X6)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X6),X7))) | s(bool,'T_0') = s(bool,sK17(X7,bool,X6))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f5072,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,sK18(X10,bool,X9)) | s(bool,'T_0') = s(bool,sK17(X10,bool,X9)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X9),X10)))) )),
  inference(subsumption_resolution,[],[f5068,f297])).
fof(f5068,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,sK18(X10,bool,X9)) | s(bool,'T_0') = s(bool,sK17(X10,bool,X9)) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X9),X10)))) )),
  inference(superposition,[],[f2340,f653])).
fof(f5070,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,sK18(X7,bool,X6)) | s(bool,'T_0') = s(bool,sK17(X7,bool,X6)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X6),X7)))) )),
  inference(subsumption_resolution,[],[f5067,f297])).
fof(f5067,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,sK18(X7,bool,X6)) | s(bool,'T_0') = s(bool,sK17(X7,bool,X6)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X6),X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2340,f653])).
fof(f4317,plain,(
  ( ! [X138,X137] : (s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | $spl72),
  inference(cnf_transformation,[],[f4317_D])).
fof(f4317_D,plain,(
  ( ! [X138,X137] : (s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4319,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl75),
  inference(cnf_transformation,[],[f4319_D])).
fof(f4319_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f4316,plain,(
  ( ! [X138,X137] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | ($spl48 | $spl61)),
  inference(forward_demodulation,[],[f4279,f2469])).
fof(f2469,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2456,f701])).
fof(f701,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f698,f501])).
fof(f698,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f360,f286])).
fof(f360,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.rat.RAT_EQ_RADD','thm.bool.IMP_ANTISYM_AX')).
fof(f501,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f280,f277])).
fof(f280,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f30])).
fof(f30,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.rat.RAT_EQ_RADD','thm.bool.NOT_CLAUSES|split|1')).
fof(f4279,plain,(
  ( ! [X138,X137] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | $spl48),
  inference(duplicate_literal_removal,[],[f4278])).
fof(f4278,plain,(
  ( ! [X138,X137] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | $spl48),
  inference(superposition,[],[f404,f2419])).
fof(f404,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK19(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK19(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK19(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK19(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.rat.RAT_EQ_RADD',aHL_EXT)).
fof(f4039,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl70),
  inference(cnf_transformation,[],[f4039_D])).
fof(f4039_D,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4038,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f4029,f297])).
fof(f4029,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1841,f286])).
fof(f4035,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl68),
  inference(cnf_transformation,[],[f4035_D])).
fof(f4035_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4036,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4025,f297])).
fof(f4025,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1841,f653])).
fof(f4034,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f4024,f297])).
fof(f4024,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1841,f653])).
fof(f4005,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3061,f3069])).
fof(f3069,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f3053,f363])).
fof(f363,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f146])).
fof(f3061,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f301])).
fof(f301,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.rat.RAT_EQ_RADD','thm.bool.EQ_CLAUSES|split|3')).
fof(f4004,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3064,f277])).
fof(f3064,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f525,f3053,f363])).
fof(f525,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f501,f297])).
fof(f4003,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3070,f277])).
fof(f3070,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f525,f3053,f363])).
fof(f4002,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3079,f277])).
fof(f3079,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f541,f3053,f384])).
fof(f384,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f541,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f501,f277,f388])).
fof(f4001,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3080,f3053])).
fof(f3080,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f556,f3053,f384])).
fof(f556,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f501,f284,f388])).
fof(f4000,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3081,f3053])).
fof(f3081,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f3015,f3053,f384])).
fof(f3015,plain,(
  sP1('F_0',f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2779])).
fof(f2779,plain,(
  ( ! [X6] : (sP1(X6,f(s(bool,'T_0'))) | p(s(bool,X6))) ) | $spl61),
  inference(subsumption_resolution,[],[f2732,f277])).
fof(f2732,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1(X6,f(s(bool,'T_0')))) ) | $spl61),
  inference(superposition,[],[f388,f2469])).
fof(f3999,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3088,f277])).
fof(f3088,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f535,f3053,f384])).
fof(f535,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f277,f385])).
fof(f3998,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3092,f3053])).
fof(f3092,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2964,f3053,f384])).
fof(f2964,plain,(
  sP1(f(s(bool,'T_0')),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2778])).
fof(f2778,plain,(
  ( ! [X5] : (sP1(f(s(bool,'T_0')),X5) | p(s(bool,X5))) ) | $spl61),
  inference(subsumption_resolution,[],[f2731,f277])).
fof(f2731,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP1(f(s(bool,'T_0')),X5)) ) | $spl61),
  inference(superposition,[],[f385,f2469])).
fof(f3997,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3093,f3053])).
fof(f3093,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f555,f3053,f384])).
fof(f555,plain,(
  ( ! [X0] : (sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f501,f284,f385])).
fof(f3996,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3112,f277])).
fof(f3112,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f716,f501,f3053,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f241])).
fof(f241,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f716,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f525,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f123,f163,f162,f161,f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f161,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f162,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f123,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> s(bool,X22) = s(bool,X23)) <=> ((p(s(bool,X24)) | p(s(bool,X23)) | p(s(bool,X22))) & (p(s(bool,X24)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) & (p(s(bool,X23)) | ~p(s(bool,X22)) | ~p(s(bool,X24))) & (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)))))),
  file('thm.rat.RAT_EQ_RADD','thm.sat.sat_2|49|_prove_ltb')).
fof(f3995,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3113,f3053])).
fof(f3113,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2466,f501,f3053,f414])).
fof(f2466,plain,(
  sP5('F_0',f(s(bool,'T_0')),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2456,f430])).
fof(f3994,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3114,f3053])).
fof(f3114,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1156,f501,f3053,f414])).
fof(f1156,plain,(
  ( ! [X0] : (sP5('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f560,f511,f508,f501,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f508,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f501,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f244])).
fof(f244,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f162])).
fof(f511,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f501,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f247])).
fof(f247,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f161])).
fof(f560,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f284,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f160])).
fof(f3993,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3165,f277])).
fof(f3165,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f1119,f501,f3053,f414])).
fof(f1119,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f513,f510,f508,f501,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f510,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f248])).
fof(f513,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f250])).
fof(f3992,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3166,f3053])).
fof(f3166,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1120,f501,f3053,f414])).
fof(f1120,plain,(
  ( ! [X0] : (sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f559,f558,f508,f501,f416])).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))))) )),
  inference(unit_resulting_resolution,[],[f284,f423])).
fof(f559,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))))) )),
  inference(unit_resulting_resolution,[],[f284,f427])).
fof(f3991,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3179,f3339])).
fof(f3339,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f525,f3053,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f251])).
fof(f3179,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f514,f511,f507,f3053,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f507,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f245])).
fof(f514,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f277,f428])).
fof(f3990,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3196,f3339])).
fof(f3196,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f514,f511,f508,f3053,f415])).
fof(f3989,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3199,f732])).
fof(f732,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f277,f525,f432])).
fof(f3199,plain,(
  sP5('T_0','T_0','F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f514,f510,f508,f3053,f415])).
fof(f3988,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3202,f3339])).
fof(f3202,plain,(
  ( ! [X0] : (sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0','F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f514,f558,f508,f3053,f415])).
fof(f3987,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3219,f3269])).
fof(f3269,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f419])).
fof(f3219,plain,(
  ~sP4('T_0','F_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f514,f511,f732,f3053,f415])).
fof(f3986,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3220,f3297])).
fof(f3297,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f423])).
fof(f3220,plain,(
  ~sP3('T_0','F_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f514,f508,f732,f3053,f415])).
fof(f3985,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3221,f3303])).
fof(f3303,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f429])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f250])).
fof(f3221,plain,(
  ~sP2('F_0','T_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f511,f508,f732,f3053,f415])).
fof(f3984,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3223,f3269])).
fof(f3223,plain,(
  ~sP4('T_0','F_0',f(s(bool,'T_0'))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f513,f511,f2465,f3053,f415])).
fof(f2465,plain,(
  ~sP5('T_0',f(s(bool,'T_0')),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f277,f2456,f432])).
fof(f3983,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3224,f3297])).
fof(f3224,plain,(
  ~sP3(f(s(bool,'T_0')),'F_0','T_0') | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f513,f508,f2465,f3053,f415])).
fof(f3982,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3225,f3303])).
fof(f3225,plain,(
  ~sP2('F_0',f(s(bool,'T_0')),'T_0') | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f511,f508,f2465,f3053,f415])).
fof(f3981,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3240,f3053])).
fof(f3240,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f509,f501,f3053,f418])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f245])).
fof(f509,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f501,f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f3980,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3252,f3269])).
fof(f3252,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3053,f3053,f418])).
fof(f3979,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3254,f3053])).
fof(f3254,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f508,f501,f3053,f418])).
fof(f3978,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3260,f3269])).
fof(f3260,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f277,f3053,f418])).
fof(f3977,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3267,f3269])).
fof(f3267,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3053,f3053,f418])).
fof(f3976,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3272,f3053])).
fof(f3272,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f512,f501,f3053,f422])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f248])).
fof(f512,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f501,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f248])).
fof(f3975,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3282,f3297])).
fof(f3282,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3053,f3053,f422])).
fof(f3974,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3284,f3053])).
fof(f3284,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f511,f501,f3053,f422])).
fof(f3973,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3289,f3297])).
fof(f3289,plain,(
  ( ! [X0] : (~sP3('T_0',X0,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f277,f3053,f422])).
fof(f3972,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3295,f3297])).
fof(f3295,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3053,f3053,f422])).
fof(f3968,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3310,f547])).
fof(f3310,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f751,f3053,f431])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f251])).
fof(f751,plain,(
  ( ! [X0] : (~sP5('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f501,f547,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f251])).
fof(f3967,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3312,f2469])).
fof(f3312,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2708,f3053,f431])).
fof(f2708,plain,(
  ~sP5('F_0',f(s(bool,'T_0')),'T_0') | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2469,f433])).
fof(f3966,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3314,f2469])).
fof(f3314,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2710,f3053,f431])).
fof(f2710,plain,(
  ~sP5('F_0','T_0',f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2469,f433])).
fof(f3965,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3315,f547])).
fof(f3315,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f753,f3053,f431])).
fof(f753,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f501,f547,f433])).
fof(f3964,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3324,f525])).
fof(f3324,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f716,f3053,f432])).
fof(f3963,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3325,f525])).
fof(f3325,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f715,f3053,f432])).
fof(f715,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f501,f525,f430])).
fof(f3962,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3961,f525])).
fof(f3961,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl64),
  inference(forward_demodulation,[],[f3326,f547])).
fof(f3326,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1123,f3053,f432])).
fof(f1123,plain,(
  ( ! [X0] : (sP5('F_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f561,f512,f509,f501,f416])).
fof(f561,plain,(
  ( ! [X2,X0,X1] : (sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f284,f429])).
fof(f3960,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3959,f525])).
fof(f3959,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl61 | $spl64)),
  inference(forward_demodulation,[],[f3327,f2469])).
fof(f3327,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2471,f3053,f432])).
fof(f2471,plain,(
  sP5('F_0','F_0',f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2456,f430])).
fof(f3958,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3957,f525])).
fof(f3957,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl61 | $spl64)),
  inference(forward_demodulation,[],[f3328,f2469])).
fof(f3328,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2466,f3053,f432])).
fof(f3956,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3955,f525])).
fof(f3955,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl64),
  inference(forward_demodulation,[],[f3329,f547])).
fof(f3329,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1156,f3053,f432])).
fof(f3954,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3347,f3484])).
fof(f3484,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f3053,f3053,f447])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f125,f166,f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f166,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f125,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) & p(s(bool,X22)))) <=> ((p(s(bool,X24)) | ~p(s(bool,X23)) | ~p(s(bool,X22))) & (p(s(bool,X23)) | ~p(s(bool,X24))) & (p(s(bool,X22)) | ~p(s(bool,X24)))))),
  file('thm.rat.RAT_EQ_RADD','thm.sat.sat_23_prove_ltb')).
fof(f3347,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3053,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f253])).
fof(f253,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f166])).
fof(f3953,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3353,f3484])).
fof(f3353,plain,(
  ( ! [X0,X1] : (~sP7(X0,'F_0',X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3053,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f3952,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3356,f3053])).
fof(f3356,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f518,f579,f3053,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f254])).
fof(f579,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f277,f501,f436])).
fof(f518,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f501,f444])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f165])).
fof(f3951,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3357,f3053])).
fof(f3357,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f516,f579,f3053,f437])).
fof(f516,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f442])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f256])).
fof(f3950,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3362,f3422])).
fof(f3422,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f442])).
fof(f3362,plain,(
  ~sP6('F_0','T_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f579,f277,f3053,f437])).
fof(f3949,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3369,f3422])).
fof(f3369,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f575,f3053,f3053,f437])).
fof(f575,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f277,f501,f435])).
fof(f3948,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3371,f3422])).
fof(f3371,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3053,f3053,f437])).
fof(f3947,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3375,f3053])).
fof(f3375,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f517,f575,f3053,f437])).
fof(f517,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f501,f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f256])).
fof(f3946,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3376,f3053])).
fof(f3376,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f516,f575,f3053,f437])).
fof(f3945,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3382,f3422])).
fof(f3382,plain,(
  ~sP6('T_0','F_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f575,f277,f3053,f437])).
fof(f3944,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3390,f3422])).
fof(f3390,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f575,f3053,f3053,f437])).
fof(f3943,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3392,f3422])).
fof(f3392,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3053,f3053,f437])).
fof(f3942,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3397,f3053])).
fof(f3397,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f518,f501,f3053,f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f256])).
fof(f3941,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3407,f3422])).
fof(f3407,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3053,f3053,f441])).
fof(f3940,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3409,f3053])).
fof(f3409,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f517,f501,f3053,f441])).
fof(f3939,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3414,f3422])).
fof(f3414,plain,(
  ( ! [X0] : (~sP6('T_0',X0,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f277,f3053,f441])).
fof(f3938,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3420,f3422])).
fof(f3420,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3053,f3053,f441])).
fof(f3937,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3425,f3053])).
fof(f3425,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f277,f3053,f447])).
fof(f3936,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3427,f3053])).
fof(f3427,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3053,f3053,f447])).
fof(f3935,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3428,f277])).
fof(f3428,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f863,f3053,f3053,f447])).
fof(f863,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f854,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f254])).
fof(f854,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f277,f277,f441])).
fof(f3934,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3429,f3053])).
fof(f3429,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1345,f3053,f3053,f447])).
fof(f1345,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1333,f434])).
fof(f1333,plain,(
  ( ! [X0] : (~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f277,f501,f859])).
fof(f859,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP6(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(subsumption_resolution,[],[f857,f277])).
fof(f857,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP6(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(superposition,[],[f441,f547])).
fof(f3933,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3430,f277])).
fof(f3430,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f277,f3053,f447])).
fof(f3932,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3436,f277])).
fof(f3436,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f575,f3053,f3053,f447])).
fof(f3931,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3437,f277])).
fof(f3437,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3053,f3053,f447])).
fof(f3930,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3444,f3053])).
fof(f3444,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f575,f277,f3053,f447])).
fof(f3929,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3446,f3053])).
fof(f3446,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f575,f3053,f3053,f447])).
fof(f3928,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3448,f277])).
fof(f3448,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f863,f3053,f3053,f447])).
fof(f3927,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3449,f277])).
fof(f3449,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f1345,f3053,f3053,f447])).
fof(f3926,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3450,f277])).
fof(f3450,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f575,f277,f3053,f447])).
fof(f3925,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3457,f277])).
fof(f3457,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f575,f3053,f3053,f447])).
fof(f3924,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3458,f277])).
fof(f3458,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3053,f3053,f447])).
fof(f3923,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3466,f277])).
fof(f3466,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f863,f277,f3053,f447])).
fof(f3922,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3467,f3053])).
fof(f3467,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f575,f3053,f3053,f447])).
fof(f3921,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3469,f277])).
fof(f3469,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f863,f3053,f3053,f447])).
fof(f3920,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3470,f277])).
fof(f3470,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f1345,f3053,f3053,f447])).
fof(f3919,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3472,f277])).
fof(f3472,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f863,f277,f3053,f447])).
fof(f3918,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3473,f3053])).
fof(f3473,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1345,f277,f3053,f447])).
fof(f3917,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3479,f3053])).
fof(f3479,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3053,f3053,f447])).
fof(f3916,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3480,f277])).
fof(f3480,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f863,f3053,f3053,f447])).
fof(f3915,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3481,f3053])).
fof(f3481,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1345,f3053,f3053,f447])).
fof(f3914,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3487,f3053])).
fof(f3487,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f594,f501,f3053,f450])).
fof(f450,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f258])).
fof(f594,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f501,f501,f446])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f3913,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3499,f3053])).
fof(f3499,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f585,f501,f3053,f450])).
fof(f585,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f501,f501,f445])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f3912,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3519,f3785])).
fof(f3785,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f521,f3053,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f454])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f260])).
fof(f260,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f521,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f501,f461])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f263])).
fof(f263,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f3519,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3053,f451])).
fof(f451,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f261])).
fof(f3911,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3528,f3785])).
fof(f3528,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3053,f452])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f261])).
fof(f3910,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3538,f3565])).
fof(f3565,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f460])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f264])).
fof(f3538,plain,(
  ~sP8('F_0','T_0','F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f605,f501,f3053,f455])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f261])).
fof(f605,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f277,f501,f452])).
fof(f3909,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3548,f3565])).
fof(f3548,plain,(
  ~sP8('F_0','F_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f601,f501,f3053,f456])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f261])).
fof(f601,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f501,f451])).
fof(f3908,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3561,f3565])).
fof(f3561,plain,(
  ( ! [X0] : (~sP8(X0,'F_0','F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f501,f3053,f458])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f264])).
fof(f3907,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3569,f277])).
fof(f3569,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f931,f3053,f463])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f127,f169,f168])).
fof(f127,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) | p(s(bool,X22)))) <=> ((p(s(bool,X24)) | ~p(s(bool,X23))) & (p(s(bool,X24)) | ~p(s(bool,X22))) & (p(s(bool,X23)) | p(s(bool,X22)) | ~p(s(bool,X24)))))),
  file('thm.rat.RAT_EQ_RADD','thm.sat.sat_25_prove_ltb')).
fof(f931,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f918,f453])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f261])).
fof(f918,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f501,f277,f458])).
fof(f3906,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3574,f277])).
fof(f3574,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f601,f3053,f463])).
fof(f3905,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3575,f3053])).
fof(f3575,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f605,f3053,f463])).
fof(f3904,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3583,f277])).
fof(f3583,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f931,f3053,f464])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f266])).
fof(f3903,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3588,f3053])).
fof(f3588,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f601,f3053,f464])).
fof(f3902,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3589,f277])).
fof(f3589,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f605,f3053,f464])).
fof(f3901,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3611,f3772])).
fof(f3772,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f524,f3053,f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f471])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f268])).
fof(f268,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f524,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f501,f478])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f271])).
fof(f271,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f3611,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3053,f469])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f269])).
fof(f3900,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3614,f3053])).
fof(f3614,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f524,f619,f3053,f472])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f269])).
fof(f619,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f277,f501,f469])).
fof(f3899,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3615,f3053])).
fof(f3615,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f523,f619,f3053,f472])).
fof(f523,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f277,f477])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f272])).
fof(f3898,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3620,f3703])).
fof(f3703,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f477])).
fof(f3620,plain,(
  ~sP10('F_0','T_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f619,f277,f3053,f472])).
fof(f3897,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3627,f3703])).
fof(f3627,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f618,f3053,f3053,f472])).
fof(f618,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f501,f501,f468])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f269])).
fof(f3896,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3629,f3703])).
fof(f3629,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f619,f3053,f3053,f472])).
fof(f3895,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3642,f3703])).
fof(f3642,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f618,f3053,f3053,f472])).
fof(f3894,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3643,f3703])).
fof(f3643,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f619,f3053,f3053,f472])).
fof(f3893,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3659,f3703])).
fof(f3659,plain,(
  ~sP10('F_0','F_0','F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f618,f501,f3053,f473])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f269])).
fof(f3892,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3671,f3703])).
fof(f3671,plain,(
  ~sP10('F_0','F_0','F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f618,f501,f3053,f474])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f269])).
fof(f3891,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3685,f3053])).
fof(f3685,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f524,f501,f3053,f475])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f3890,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3688,f3703])).
fof(f3688,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f501,f3053,f475])).
fof(f3889,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3697,f3053])).
fof(f3697,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f522,f501,f3053,f475])).
fof(f522,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f501,f476])).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f272])).
fof(f3888,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3699,f3703])).
fof(f3699,plain,(
  ( ! [X0] : (~sP10('T_0','F_0',X0)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f277,f501,f3053,f475])).
fof(f3887,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3701,f3703])).
fof(f3701,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f501,f3053,f475])).
fof(f3886,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3718,f277])).
fof(f3718,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f978,f3053,f481])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f155,f172,f171])).
fof(f155,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) => p(s(bool,X22)))) <=> ((p(s(bool,X24)) | p(s(bool,X23))) & (p(s(bool,X24)) | ~p(s(bool,X22))) & (~p(s(bool,X23)) | p(s(bool,X22)) | ~p(s(bool,X24)))))),
  file('thm.rat.RAT_EQ_RADD','thm.sat.sat_27_prove_ltb')).
fof(f978,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f966,f470])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f966,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f277,f501,f277,f475])).
fof(f3885,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3723,f3053])).
fof(f3723,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f618,f3053,f481])).
fof(f3884,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3724,f277])).
fof(f3724,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f619,f3053,f481])).
fof(f3883,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3730,f277])).
fof(f3730,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f621,f501,f3053,f482])).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f274])).
fof(f621,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f501,f277,f480])).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f274])).
fof(f3882,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3732,f3053])).
fof(f3732,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1237,f501,f3053,f482])).
fof(f1237,plain,(
  ( ! [X0,X1] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f501,f786])).
fof(f786,plain,(
  ( ! [X26,X27,X25] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X25)))),X27,X26) | p(s(bool,X26))) )),
  inference(subsumption_resolution,[],[f771,f277])).
fof(f771,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X25)))),X27,X26)) )),
  inference(superposition,[],[f480,f547])).
fof(f3881,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3747,f277])).
fof(f3747,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f958,f501,f3053,f482])).
fof(f958,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f524,f277,f501,f474])).
fof(f3880,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3763,f522])).
fof(f3763,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f618,f3053,f485])).
fof(f3879,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3764,f523])).
fof(f3764,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f619,f3053,f485])).
fof(f3878,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3774,f519])).
fof(f519,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f459])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f264])).
fof(f3774,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f601,f3053,f486])).
fof(f3877,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3775,f520])).
fof(f520,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f277,f460])).
fof(f3775,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f605,f3053,f486])).
fof(f3853,plain,(
  $false | ($spl34 | $spl64)),
  inference(unit_resulting_resolution,[],[f621,f501,f3053,f1820])).
fof(f3848,plain,(
  $false | ($spl32 | $spl64)),
  inference(unit_resulting_resolution,[],[f978,f3053,f1817])).
fof(f3846,plain,(
  $false | ($spl28 | $spl64)),
  inference(unit_resulting_resolution,[],[f524,f501,f3053,f1811])).
fof(f3841,plain,(
  $false | ($spl24 | $spl64)),
  inference(unit_resulting_resolution,[],[f931,f3053,f1805])).
fof(f3839,plain,(
  $false | ($spl22 | $spl64)),
  inference(unit_resulting_resolution,[],[f931,f3053,f1802])).
fof(f3837,plain,(
  $false | ($spl18 | $spl64)),
  inference(unit_resulting_resolution,[],[f594,f501,f3053,f1796])).
fof(f3832,plain,(
  $false | ($spl16 | $spl64)),
  inference(unit_resulting_resolution,[],[f518,f501,f3053,f1793])).
fof(f3827,plain,(
  $false | ($spl14 | $spl64)),
  inference(unit_resulting_resolution,[],[f512,f501,f3053,f1790])).
fof(f3822,plain,(
  $false | ($spl12 | $spl64)),
  inference(unit_resulting_resolution,[],[f509,f501,f3053,f1787])).
fof(f3815,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f1237,f3053,f995])).
fof(f995,plain,(
  ( ! [X6,X8,X7] : (~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8,X7) | p(s(bool,X8)) | ~p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f993,f277])).
fof(f993,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8,X7)) )),
  inference(superposition,[],[f482,f547])).
fof(f3812,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f524,f3053,f971])).
fof(f971,plain,(
  ( ! [X6,X8,X7] : (~sP10(X8,X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6))))) | ~p(s(bool,X8)) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f969,f277])).
fof(f969,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~p(s(bool,X8)) | ~sP10(X8,X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))))) )),
  inference(superposition,[],[f475,f547])).
fof(f3807,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f594,f3053,f890])).
fof(f890,plain,(
  ( ! [X6,X8,X7] : (~sP7(X8,X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6))))) | p(s(bool,X8)) | ~p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f889,f277])).
fof(f889,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP7(X8,X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))))) )),
  inference(superposition,[],[f450,f547])).
fof(f3804,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f518,f3053,f859])).
fof(f3799,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f512,f3053,f816])).
fof(f816,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP3(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(subsumption_resolution,[],[f814,f277])).
fof(f814,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP3(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(superposition,[],[f422,f547])).
fof(f3794,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f509,f3053,f799])).
fof(f799,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP4(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(subsumption_resolution,[],[f797,f277])).
fof(f797,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP4(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(superposition,[],[f418,f547])).
fof(f3786,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f601,f521,f3053,f486])).
fof(f3784,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f601,f564,f3053,f486])).
fof(f564,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f284,f460])).
fof(f3782,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f520,f3053,f486])).
fof(f3780,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f563,f3053,f486])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))))) )),
  inference(unit_resulting_resolution,[],[f284,f459])).
fof(f3778,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f601,f519,f3053,f486])).
fof(f3773,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f524,f3053,f485])).
fof(f3771,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f565,f3053,f485])).
fof(f565,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f284,f477])).
fof(f3769,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f619,f523,f3053,f485])).
fof(f3767,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f522,f3053,f485])).
fof(f3762,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f621,f501,f3053,f3053,f482])).
fof(f3754,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f958,f501,f277,f3053,f482])).
fof(f3745,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f621,f501,f3053,f3053,f482])).
fof(f3737,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f621,f501,f277,f3053,f482])).
fof(f3729,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f3053,f3053,f481])).
fof(f3727,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f619,f277,f3053,f481])).
fof(f3722,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f3053,f3053,f481])).
fof(f3720,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f978,f277,f3053,f481])).
fof(f3715,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f501,f3053,f480])).
fof(f3710,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f501,f501,f3053,f479])).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f274])).
fof(f3704,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f966,f3053,f477])).
fof(f3702,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f522,f3053,f501,f3053,f475])).
fof(f3700,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f522,f277,f501,f3053,f475])).
fof(f3689,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f522,f3053,f501,f3053,f475])).
fof(f3687,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f524,f277,f501,f3053,f475])).
fof(f3676,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f524,f501,f3053,f474])).
fof(f3674,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f522,f501,f3053,f474])).
fof(f3664,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f524,f501,f3053,f473])).
fof(f3662,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f522,f501,f3053,f473])).
fof(f3652,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f524,f3053,f3053,f472])).
fof(f3650,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f565,f3053,f3053,f472])).
fof(f3648,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f619,f523,f3053,f3053,f472])).
fof(f3646,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f618,f522,f3053,f3053,f472])).
fof(f3630,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f524,f619,f3053,f3053,f472])).
fof(f3628,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f524,f618,f3053,f3053,f472])).
fof(f3621,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f524,f619,f277,f3053,f472])).
fof(f3612,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f958,f501,f3053,f469])).
fof(f3604,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f909,f501,f501,f3053,f465])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f266])).
fof(f909,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f521,f501,f501,f457])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f261])).
fof(f3594,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f601,f3053,f3053,f464])).
fof(f3592,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f277,f3053,f464])).
fof(f3587,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f601,f3053,f3053,f464])).
fof(f3585,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f931,f277,f3053,f464])).
fof(f3580,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f601,f3053,f3053,f463])).
fof(f3578,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f601,f277,f3053,f463])).
fof(f3573,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f601,f3053,f3053,f463])).
fof(f3571,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f931,f277,f3053,f463])).
fof(f3566,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f918,f3053,f460])).
fof(f3564,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f918,f3053,f459])).
fof(f3562,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f521,f501,f501,f3053,f458])).
fof(f3554,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f601,f521,f501,f3053,f456])).
fof(f3551,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f601,f519,f501,f3053,f456])).
fof(f3539,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f521,f605,f501,f3053,f455])).
fof(f3529,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f909,f501,f3053,f452])).
fof(f3520,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f909,f501,f3053,f451])).
fof(f3511,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f585,f501,f3053,f3053,f450])).
fof(f3505,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f585,f501,f277,f3053,f450])).
fof(f3498,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f585,f501,f3053,f3053,f450])).
fof(f3492,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f594,f501,f277,f3053,f450])).
fof(f3485,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f575,f3053,f3053,f3053,f447])).
fof(f3483,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f579,f277,f3053,f3053,f447])).
fof(f3477,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f575,f3053,f277,f3053,f447])).
fof(f3475,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f863,f277,f277,f3053,f447])).
fof(f3464,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f575,f3053,f3053,f3053,f447])).
fof(f3462,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f575,f277,f3053,f3053,f447])).
fof(f3456,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f575,f3053,f277,f3053,f447])).
fof(f3454,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f575,f277,f277,f3053,f447])).
fof(f3443,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f575,f3053,f3053,f3053,f447])).
fof(f3441,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f575,f277,f3053,f3053,f447])).
fof(f3435,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3053,f277,f3053,f447])).
fof(f3433,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f579,f277,f277,f3053,f447])).
fof(f3423,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f854,f3053,f442])).
fof(f3421,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f517,f501,f3053,f3053,f441])).
fof(f3415,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f517,f501,f277,f3053,f441])).
fof(f3408,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f517,f501,f3053,f3053,f441])).
fof(f3402,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f518,f501,f277,f3053,f441])).
fof(f3393,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f518,f579,f3053,f3053,f437])).
fof(f3391,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f517,f575,f3053,f3053,f437])).
fof(f3383,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f517,f575,f277,f3053,f437])).
fof(f3372,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f518,f579,f3053,f3053,f437])).
fof(f3370,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f517,f575,f3053,f3053,f437])).
fof(f3363,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f518,f579,f277,f3053,f437])).
fof(f3354,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f594,f501,f3053,f436])).
fof(f3348,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f585,f501,f3053,f435])).
fof(f3342,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f715,f525,f3053,f432])).
fof(f3340,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f716,f525,f3053,f432])).
fof(f3323,plain,(
  $false | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2710,f2469,f3053,f431])).
fof(f3321,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f753,f547,f3053,f431])).
fof(f3319,plain,(
  $false | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2708,f2469,f3053,f431])).
fof(f3317,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f751,f547,f3053,f431])).
fof(f3876,plain,(
  $false | $spl64),
  inference(trivial_inequality_removal,[],[f3313])).
fof(f3313,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f835,f3053,f431])).
fof(f835,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f826,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f826,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f501,f501,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f250])).
fof(f3875,plain,(
  $false | $spl64),
  inference(trivial_inequality_removal,[],[f3311])).
fof(f3311,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f820,f3053,f431])).
fof(f820,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f811,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f811,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f277,f277,f422])).
fof(f3304,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f826,f3053,f429])).
fof(f3302,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f826,f3053,f428])).
fof(f3300,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f826,f3053,f427])).
fof(f3298,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f811,f3053,f423])).
fof(f3296,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f511,f501,f3053,f3053,f422])).
fof(f3290,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f511,f501,f277,f3053,f422])).
fof(f3283,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f511,f501,f3053,f3053,f422])).
fof(f3277,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f512,f501,f277,f3053,f422])).
fof(f3270,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f794,f3053,f419])).
fof(f794,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f277,f277,f418])).
fof(f3268,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f508,f501,f3053,f3053,f418])).
fof(f3261,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f508,f501,f277,f3053,f418])).
fof(f3253,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f508,f501,f3053,f3053,f418])).
fof(f3246,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f509,f501,f277,f3053,f418])).
fof(f3226,plain,(
  $false | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f513,f511,f508,f2465,f3053,f415])).
fof(f3222,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f514,f511,f508,f732,f3053,f415])).
fof(f3170,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f716,f3053,f501,f3053,f414])).
fof(f3168,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f1119,f277,f501,f3053,f414])).
fof(f3140,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f716,f501,f3053,f3053,f414])).
fof(f3123,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f716,f501,f277,f3053,f414])).
fof(f3111,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f532,f501,f3053,f388])).
fof(f532,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f501,f383])).
fof(f383,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f228])).
fof(f3104,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f532,f501,f3053,f385])).
fof(f3097,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f535,f3053,f3053,f384])).
fof(f3095,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f535,f277,f3053,f384])).
fof(f3087,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f535,f3053,f3053,f384])).
fof(f3085,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f541,f277,f3053,f384])).
fof(f3077,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f546,f3053,f370])).
fof(f370,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f546,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f501,f502])).
fof(f502,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f368,f501])).
fof(f368,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f3075,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f546,f3053,f369])).
fof(f369,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f3073,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f525,f3053,f363])).
fof(f3071,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f277,f525,f3053,f363])).
fof(f3067,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f3053,f525,f3053,f363])).
fof(f3065,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f277,f525,f3053,f363])).
fof(f3060,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f525,f3053,f298])).
fof(f3058,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3053])).
fof(f3057,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl66),
  inference(cnf_transformation,[],[f3057_D])).
fof(f3057_D,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3056,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f3046,f297])).
fof(f3046,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1840,f286])).
fof(f3052,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl62),
  inference(cnf_transformation,[],[f3052_D])).
fof(f3052_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3054,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f3042,f297])).
fof(f3042,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1840,f653])).
fof(f3051,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f3041,f297])).
fof(f3041,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1840,f653])).
fof(f2454,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK19(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X0,X1,bool,bool))))) ) | $spl58),
  inference(cnf_transformation,[],[f2454_D])).
fof(f2454_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK19(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X0,X1,bool,bool))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2456,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl61),
  inference(cnf_transformation,[],[f2456_D])).
fof(f2456_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f2444,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK19(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1289,f653])).
fof(f1289,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK19(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK19(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f404,f286])).
fof(f2443,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK19(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1289,f653])).
fof(f2423,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl52),
  inference(cnf_transformation,[],[f2423_D])).
fof(f2423_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2419,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK19(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl48),
  inference(cnf_transformation,[],[f2419_D])).
fof(f2419_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK19(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2426,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl57),
  inference(cnf_transformation,[],[f2426_D])).
fof(f2426_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f2417,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK19(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1288,f286])).
fof(f1288,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK19(X0,X1,X2,bool))) )),
  inference(superposition,[],[f404,f286])).
fof(f2424,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl54),
  inference(cnf_transformation,[],[f2424_D])).
fof(f2424_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2416,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK19(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1288,f653])).
fof(f2415,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK19(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1288,f653])).
fof(f2418,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl46),
  inference(cnf_transformation,[],[f2418_D])).
fof(f2418_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2421,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl51),
  inference(cnf_transformation,[],[f2421_D])).
fof(f2421_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f2414,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK19(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1288,f653])).
fof(f2413,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK19(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1288,f653])).
fof(f1944,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f276,f275,f277,f501,f284,f286,f298,f297,f302,f301,f357,f356,f358,f359,f363,f360,f370,f369,f502,f388,f385,f384,f383,f403,f402,f401,f404,f417,f416,f415,f414,f413,f412,f411,f421,f420,f419,f418,f425,f424,f423,f422,f429,f428,f427,f426,f433,f432,f431,f430,f487,f437,f436,f435,f434,f444,f443,f442,f441,f450,f447,f446,f445,f457,f456,f455,f486,f453,f452,f451,f461,f460,f459,f458,f465,f464,f463,f462,f474,f473,f472,f485,f470,f469,f468,f478,f477,f476,f475,f482,f481,f480,f479,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f532,f533,f535,f541,f546,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f551,f550,f549,f548,f547,f575,f579,f585,f594,f601,f605,f608,f611,f618,f619,f621,f628,f653,f701,f712,f716,f715,f723,f726,f732,f730,f754,f753,f752,f751,f782,f783,f784,f785,f786,f787,f788,f789,f794,f799,f811,f816,f820,f826,f835,f840,f839,f854,f859,f863,f890,f909,f918,f924,f931,f950,f958,f966,f971,f978,f995,f996,f1066,f1094,f1088,f1082,f1079,f1100,f1123,f1120,f1119,f1156,f1151,f1174,f1237,f1251,f1253,f1257,f1267,f1288,f1289,f1290,f1298,f1302,f1310,f1329,f1332,f1333,f1341,f1345,f1372,f1434,f1435,f1436,f1437,f1438,f1486,f1775,f1779,f1783,f1786,f1789,f1792,f1795,f1798,f1801,f1804,f1807,f1810,f1813,f1816,f1819,f1822,f1825,f1828,f1831,f1834,f1836,f1837,f1838,f1839,f1840,f1841,f1842,f1843,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1560,f1561,f1562,f1563,f1564,f1869,f1871,f1874,f1876,f1878,f1880,f1882,f1884,f1886,f1888,f1890,f1892,f1894,f1896,f1898,f1900,f1902,f1904,f1906,f1908,f1777,f1938,f1943])).
fof(f1943,plain,(
  ( ! [X6,X2,X0,X5,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f1928,f277])).
fof(f1928,plain,(
  ( ! [X6,X2,X0,X5,X1] : (~p(s(bool,'T_0')) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)) ) | $spl6),
  inference(backward_demodulation,[],[f1777,f401])).
fof(f1938,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f1909,f1777])).
fof(f1909,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl6),
  inference(backward_demodulation,[],[f1777,f301])).
fof(f1564,plain,(
  ( ! [X156,X154,X157,X155] : (s(X156,f(s(bool,'T_0'))) != s(X156,g(s(bool,'T_0'))) | s(fun(bool,X156),X154) = s(fun(bool,X156),X155) | s(bool,X157) = s(bool,sK19(X154,X155,X156,bool)) | s(bool,'T_0') = s(bool,X157)) )),
  inference(superposition,[],[f404,f653])).
fof(f1563,plain,(
  ( ! [X152,X151,X153] : (s(bool,'T_0') != s(bool,sK17(X151,bool,X152)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X152),X151)))) | s(bool,X153) = s(bool,sK18(X151,bool,X152)) | s(bool,'T_0') = s(bool,X153)) )),
  inference(superposition,[],[f403,f653])).
fof(f1562,plain,(
  ( ! [X146,X147,X148] : (~p(s(bool,happ(s(fun(bool,bool),X146),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X146),s(bool,X147)))) | s(bool,X148) = s(bool,sK16(X146,bool)) | s(bool,'T_0') = s(bool,X148)) )),
  inference(superposition,[],[f357,f653])).
fof(f1561,plain,(
  ( ! [X144,X145,X143] : (~p(s(bool,happ(s(fun(bool,bool),X143),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X143),s(bool,X144)))) | s(bool,X145) = s(bool,sK15(X143,bool)) | s(bool,'T_0') = s(bool,X145)) )),
  inference(superposition,[],[f356,f653])).
fof(f1560,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,'T_0') != s(bool,f(s(X139,sK19(X140,X141,bool,X139)))) | s(fun(X139,bool),X140) = s(fun(X139,bool),X141) | s(bool,X142) = s(bool,g(s(X139,sK19(X140,X141,bool,X139)))) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f404,f653])).
fof(f1863,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (s(X127,happ(s(fun(X126,X127),X128),s(X126,X129))) != s(X127,happ(s(fun(X126,X127),X128),s(X126,X130))) | s(X126,X129) = s(X126,X130) | s(bool,X131) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X126,X127),X128))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(subsumption_resolution,[],[f1558,f277])).
fof(f1558,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (~p(s(bool,'T_0')) | s(X127,happ(s(fun(X126,X127),X128),s(X126,X129))) != s(X127,happ(s(fun(X126,X127),X128),s(X126,X130))) | s(X126,X129) = s(X126,X130) | s(bool,X131) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X126,X127),X128))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(superposition,[],[f401,f653])).
fof(f1862,plain,(
  ( ! [X123,X121,X124,X122] : (p(s(bool,happ(s(fun(X121,bool),X122),s(X121,X123)))) | s(bool,X124) = s(bool,happ(s(fun(X121,bool),X122),s(X121,sK16(X122,X121)))) | s(bool,'T_0') = s(bool,X124)) )),
  inference(subsumption_resolution,[],[f1557,f277])).
fof(f1557,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X121,bool),X122),s(X121,X123)))) | s(bool,X124) = s(bool,happ(s(fun(X121,bool),X122),s(X121,sK16(X122,X121)))) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f357,f653])).
fof(f1861,plain,(
  ( ! [X118,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119)))) | s(bool,X120) = s(bool,happ(s(fun(X117,bool),X118),s(X117,sK15(X118,X117)))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f1556,f277])).
fof(f1556,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119)))) | s(bool,X120) = s(bool,happ(s(fun(X117,bool),X118),s(X117,sK15(X118,X117)))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f356,f653])).
fof(f1860,plain,(
  ( ! [X109,X107,X108,X106] : (p(s(bool,X107)) | ~sP6(X106,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X108)))),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f1550,f277])).
fof(f1550,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,'T_0')) | p(s(bool,X107)) | ~sP6(X106,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X108)))),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f859,f653])).
fof(f1859,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X103)) | ~sP3(X102,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X104)))),X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f1549,f277])).
fof(f1549,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | p(s(bool,X103)) | ~sP3(X102,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X104)))),X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f816,f653])).
fof(f1858,plain,(
  ( ! [X101,X99,X100,X98] : (p(s(bool,X99)) | ~sP4(X98,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X100)))),X99) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(subsumption_resolution,[],[f1548,f277])).
fof(f1548,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,'T_0')) | p(s(bool,X99)) | ~sP4(X98,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X100)))),X99) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f799,f653])).
fof(f1857,plain,(
  ( ! [X94,X97,X95,X96] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X95)))),X94,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f1547,f277])).
fof(f1547,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,'T_0')) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X95)))),X94,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f787,f653])).
fof(f1856,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1545,f277])).
fof(f1545,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f486,f653])).
fof(f1855,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1544,f277])).
fof(f1544,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f485,f653])).
fof(f1854,plain,(
  ( ! [X80,X81,X79,X82] : (~sP11(X79,X81,X80) | p(s(bool,X81)) | ~p(s(bool,X80)) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1543,f277])).
fof(f1543,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f482,f653])).
fof(f1853,plain,(
  ( ! [X78,X76,X77,X75] : (sP11(X75,X76,X77) | ~p(s(bool,X76)) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1542,f277])).
fof(f1542,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f481,f653])).
fof(f1852,plain,(
  ( ! [X74,X72,X71,X73] : (sP11(X71,X73,X72) | p(s(bool,X72)) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1541,f277])).
fof(f1541,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f480,f653])).
fof(f1851,plain,(
  ( ! [X70,X68,X69,X67] : (~sP10(X69,X68,X67) | ~p(s(bool,X69)) | p(s(bool,X68)) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1540,f277])).
fof(f1540,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f475,f653])).
fof(f1850,plain,(
  ( ! [X66,X64,X65,X63] : (~sP9(X63,X65,X64) | p(s(bool,X65)) | p(s(bool,X64)) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1539,f277])).
fof(f1539,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f465,f653])).
fof(f1849,plain,(
  ( ! [X61,X59,X62,X60] : (sP9(X59,X60,X61) | ~p(s(bool,X60)) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1538,f277])).
fof(f1538,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f464,f653])).
fof(f1848,plain,(
  ( ! [X57,X58,X56,X55] : (sP9(X55,X57,X56) | ~p(s(bool,X56)) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1537,f277])).
fof(f1537,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f463,f653])).
fof(f1847,plain,(
  ( ! [X54,X52,X53,X51] : (~sP8(X51,X52,X53) | p(s(bool,X53)) | p(s(bool,X52)) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1536,f277])).
fof(f1536,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f458,f653])).
fof(f1846,plain,(
  ( ! [X47,X50,X48,X49] : (~sP7(X49,X48,X47) | p(s(bool,X49)) | ~p(s(bool,X48)) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1535,f277])).
fof(f1535,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f450,f653])).
fof(f1845,plain,(
  ( ! [X45,X43,X46,X44] : (~sP6(X44,X43,X45) | p(s(bool,X45)) | ~p(s(bool,X44)) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1534,f277])).
fof(f1534,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f441,f653])).
fof(f1844,plain,(
  ( ! [X39,X41,X42,X40] : (~sP3(X40,X39,X41) | p(s(bool,X41)) | ~p(s(bool,X40)) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1533,f277])).
fof(f1533,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f422,f653])).
fof(f1843,plain,(
  ( ! [X37,X35,X38,X36] : (~sP4(X36,X35,X37) | p(s(bool,X37)) | ~p(s(bool,X36)) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1532,f277])).
fof(f1532,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f418,f653])).
fof(f1842,plain,(
  ( ! [X33,X31,X34,X32] : (~sP4(X32,X31,X33) | sP5(X32,X33,X31) | ~sP3(X33,X31,X32) | ~sP2(X31,X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1531,f277])).
fof(f1531,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | sP5(X32,X33,X31) | ~sP4(X32,X31,X33) | ~sP3(X33,X31,X32) | ~sP2(X31,X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f415,f653])).
fof(f1839,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1528,f277])).
fof(f1528,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f370,f653])).
fof(f1838,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1527,f277])).
fof(f1527,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f369,f653])).
fof(f1837,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1526,f277])).
fof(f1526,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f363,f653])).
fof(f1836,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1522,f277])).
fof(f1522,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f301,f653])).
fof(f1438,plain,(
  ( ! [X158,X161,X159,X160] : (s(X160,f(s(bool,X161))) != s(X160,g(s(bool,X161))) | s(fun(bool,X160),X158) = s(fun(bool,X160),X159) | s(bool,'T_0') = s(bool,X161) | s(bool,'T_0') = s(bool,sK19(X158,X159,X160,bool))) )),
  inference(superposition,[],[f404,f653])).
fof(f1436,plain,(
  ( ! [X151,X149,X150] : (~p(s(bool,happ(s(fun(bool,bool),X149),s(bool,X150)))) | p(s(bool,happ(s(fun(bool,bool),X149),s(bool,X151)))) | s(bool,'T_0') = s(bool,X150) | s(bool,'T_0') = s(bool,sK16(X149,bool))) )),
  inference(superposition,[],[f357,f653])).
fof(f1435,plain,(
  ( ! [X146,X147,X148] : (~p(s(bool,happ(s(fun(bool,bool),X146),s(bool,X147)))) | p(s(bool,happ(s(fun(bool,bool),X146),s(bool,X148)))) | s(bool,'T_0') = s(bool,X147) | s(bool,'T_0') = s(bool,sK15(X146,bool))) )),
  inference(superposition,[],[f356,f653])).
fof(f1434,plain,(
  ( ! [X144,X142,X145,X143] : (s(bool,X145) != s(bool,f(s(X142,sK19(X143,X144,bool,X142)))) | s(fun(X142,bool),X143) = s(fun(X142,bool),X144) | s(bool,'T_0') = s(bool,X145) | s(bool,'T_0') = s(bool,g(s(X142,sK19(X143,X144,bool,X142))))) )),
  inference(superposition,[],[f404,f653])).
fof(f1341,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | ~sP6('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X8)))),X7)) )),
  inference(subsumption_resolution,[],[f1339,f277])).
fof(f1339,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP6('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X8)))),X7)) )),
  inference(superposition,[],[f859,f547])).
fof(f1329,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | ~sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X8)))),X7)) )),
  inference(subsumption_resolution,[],[f1327,f277])).
fof(f1327,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X8)))),X7)) )),
  inference(superposition,[],[f816,f547])).
fof(f1310,plain,(
  ( ! [X2,X3,X1] : (s(X2,happ(s(fun(X1,X2),X3),s(X1,sK17(X3,X1,X2)))) = s(X2,happ(s(fun(X1,X2),X3),s(X1,sK18(X3,X1,X2)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X2),X3)))) )),
  inference(subsumption_resolution,[],[f1309,f501])).
fof(f1309,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'F_0')) | s(X2,happ(s(fun(X1,X2),X3),s(X1,sK17(X3,X1,X2)))) = s(X2,happ(s(fun(X1,X2),X3),s(X1,sK18(X3,X1,X2)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X2),X3)))) )),
  inference(superposition,[],[f402,f286])).
fof(f1302,plain,(
  ( ! [X0] : (~sP5('T_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1290,f413])).
fof(f1298,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | ~sP4('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X8)))),X7)) )),
  inference(subsumption_resolution,[],[f1296,f277])).
fof(f1296,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP4('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X8)))),X7)) )),
  inference(superposition,[],[f799,f547])).
fof(f1290,plain,(
  ( ! [X0] : (~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f277,f501,f799])).
fof(f1267,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK16(X0,bool))) )),
  inference(superposition,[],[f357,f286])).
fof(f1257,plain,(
  ( ! [X6,X8,X7] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X7)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(subsumption_resolution,[],[f1256,f277])).
fof(f1256,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X7)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(superposition,[],[f787,f547])).
fof(f1253,plain,(
  ( ! [X0,X1] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f277,f787])).
fof(f1251,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK15(X0,bool))) )),
  inference(superposition,[],[f356,f286])).
fof(f1151,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f514,f753,f507,f501,f417])).
fof(f1100,plain,(
  ( ! [X6,X8,X7] : (~sP4(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8) | sP5(X7,X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6))))) | ~sP3(X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X7)) )),
  inference(subsumption_resolution,[],[f1099,f561])).
fof(f1099,plain,(
  ( ! [X6,X8,X7] : (sP5(X7,X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6))))) | ~sP4(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8) | ~sP3(X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X7) | ~sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8,X7)) )),
  inference(subsumption_resolution,[],[f1098,f277])).
fof(f1098,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | sP5(X7,X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6))))) | ~sP4(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8) | ~sP3(X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X7) | ~sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8,X7)) )),
  inference(superposition,[],[f415,f547])).
fof(f1079,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f515,f510,f507,f277,f415])).
fof(f1082,plain,(
  ( ! [X0] : (sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f559,f558,f507,f277,f415])).
fof(f1088,plain,(
  ( ! [X0,X1] : (sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f559,f558,f557,f277,f415])).
fof(f1094,plain,(
  ( ! [X0] : (~sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f560,f557,f751,f277,f415])).
fof(f1066,plain,(
  ( ! [X0,X1] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f277,f785])).
fof(f996,plain,(
  ( ! [X0,X1] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f784])).
fof(f950,plain,(
  ( ! [X6,X8,X7] : (~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8,X7) | p(s(bool,X8)) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f947,f277])).
fof(f947,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | p(s(bool,X8)) | ~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8,X7)) )),
  inference(superposition,[],[f465,f547])).
fof(f924,plain,(
  ( ! [X6,X8,X7] : (~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X7,X8) | p(s(bool,X8)) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f921,f277])).
fof(f921,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | p(s(bool,X8)) | ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X7,X8)) )),
  inference(superposition,[],[f458,f547])).
fof(f839,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f516,f277,f277,f437])).
fof(f840,plain,(
  ( ! [X0] : (sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f562,f277,f277,f437])).
fof(f789,plain,(
  ( ! [X35,X36,X34] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X34)))),X35,X36) | ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X34)))),X35,X36)) )),
  inference(subsumption_resolution,[],[f774,f277])).
fof(f774,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X34)))),X35,X36) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X34)))),X35,X36)) )),
  inference(superposition,[],[f486,f547])).
fof(f788,plain,(
  ( ! [X33,X31,X32] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X31)))),X32,X33) | ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X31)))),X32,X33)) )),
  inference(subsumption_resolution,[],[f773,f277])).
fof(f773,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X31)))),X32,X33) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X31)))),X32,X33)) )),
  inference(superposition,[],[f485,f547])).
fof(f787,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X29)) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X28)))),X29,X30)) )),
  inference(subsumption_resolution,[],[f772,f277])).
fof(f772,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X28)))),X29,X30)) )),
  inference(superposition,[],[f481,f547])).
fof(f785,plain,(
  ( ! [X24,X23,X22] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X22)))),X23,X24) | ~p(s(bool,X23))) )),
  inference(subsumption_resolution,[],[f770,f277])).
fof(f770,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X22)))),X23,X24)) )),
  inference(superposition,[],[f464,f547])).
fof(f784,plain,(
  ( ! [X21,X19,X20] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X19)))),X21,X20) | ~p(s(bool,X20))) )),
  inference(subsumption_resolution,[],[f769,f277])).
fof(f769,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X19)))),X21,X20)) )),
  inference(superposition,[],[f463,f547])).
fof(f783,plain,(
  ( ! [X17,X18] : (sP1(X18,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X17))))) | p(s(bool,X18))) )),
  inference(subsumption_resolution,[],[f768,f277])).
fof(f768,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | sP1(X18,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X17)))))) )),
  inference(superposition,[],[f388,f547])).
fof(f782,plain,(
  ( ! [X15,X16] : (sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X15)))),X16) | p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f767,f277])).
fof(f767,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X15)))),X16)) )),
  inference(superposition,[],[f385,f547])).
fof(f752,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f547,f431])).
fof(f754,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f277,f547,f431])).
fof(f730,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f277,f525,f432])).
fof(f726,plain,(
  ( ! [X0,X1] : (sP11('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f277,f565,f485])).
fof(f723,plain,(
  ( ! [X0,X1] : (sP9('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f277,f564,f486])).
fof(f712,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))))) )),
  inference(unit_resulting_resolution,[],[f277,f563,f486])).
fof(f628,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f277,f277,f481])).
fof(f611,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f277,f277,f464])).
fof(f608,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f277,f463])).
fof(f548,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f284,f301])).
fof(f549,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))))) )),
  inference(unit_resulting_resolution,[],[f284,f369])).
fof(f550,plain,(
  ( ! [X0,X1] : (sP0('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f284,f370])).
fof(f551,plain,(
  ( ! [X0] : (~sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f284,f384])).
fof(f553,plain,(
  ( ! [X0] : (~sP1('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f277,f284,f384])).
fof(f554,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))))) )),
  inference(unit_resulting_resolution,[],[f284,f284,f384])).
fof(f557,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))))) )),
  inference(unit_resulting_resolution,[],[f284,f419])).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))))) )),
  inference(unit_resulting_resolution,[],[f284,f442])).
fof(f533,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f277,f277,f384])).
fof(f515,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f277,f429])).
fof(f506,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f277,f370])).
fof(f505,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f369])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f266])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f254])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f356,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16])],[f216])).
fof(f216,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f215])).
fof(f215,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X12] : ! [X14] : (~! [X13] : p(s(bool,happ(s(fun(X12,bool),X14),s(X12,X13)))) <=> ? [X13] : ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,X13)))))),
  file('thm.rat.RAT_EQ_RADD','thm.bool.NOT_FORALL_THM')).
fof(f357,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f302,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f180])).
fof(f1940,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f1911,f277])).
fof(f1911,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f1777,f501])).
fof(f1939,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f276,f275,f277,f501,f284,f286,f298,f297,f302,f301,f357,f356,f358,f359,f363,f360,f370,f369,f502,f388,f385,f384,f383,f403,f402,f401,f404,f417,f416,f415,f414,f413,f412,f411,f421,f420,f419,f418,f425,f424,f423,f422,f429,f428,f427,f426,f433,f432,f431,f430,f487,f437,f436,f435,f434,f444,f443,f442,f441,f450,f447,f446,f445,f457,f456,f455,f486,f453,f452,f451,f461,f460,f459,f458,f465,f464,f463,f462,f474,f473,f472,f485,f470,f469,f468,f478,f477,f476,f475,f482,f481,f480,f479,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f532,f533,f535,f541,f546,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f551,f550,f549,f548,f547,f575,f579,f585,f594,f601,f605,f608,f611,f618,f619,f621,f628,f653,f701,f712,f716,f715,f723,f726,f732,f730,f754,f753,f752,f751,f782,f783,f784,f785,f786,f787,f788,f789,f794,f799,f811,f816,f820,f826,f835,f840,f839,f854,f859,f863,f890,f909,f918,f924,f931,f950,f958,f966,f971,f978,f995,f996,f1066,f1094,f1088,f1082,f1079,f1100,f1123,f1120,f1119,f1156,f1151,f1174,f1237,f1251,f1253,f1257,f1267,f1288,f1289,f1290,f1298,f1302,f1310,f1329,f1332,f1333,f1341,f1345,f1372,f1434,f1435,f1436,f1437,f1438,f1486,f1775,f1779,f1783,f1786,f1789,f1792,f1795,f1798,f1801,f1804,f1807,f1810,f1813,f1816,f1819,f1822,f1825,f1828,f1831,f1834,f1836,f1837,f1838,f1839,f1840,f1841,f1842,f1843,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1560,f1561,f1562,f1563,f1564,f1869,f1871,f1874,f1876,f1878,f1880,f1882,f1884,f1886,f1888,f1890,f1892,f1894,f1896,f1898,f1900,f1902,f1904,f1906,f1908,f1777,f1938])).
fof(f1835,plain,(
  ( ! [X2,X1] : (~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1) | p(s(bool,X1))) ) | $spl44),
  inference(cnf_transformation,[],[f1835_D])).
fof(f1835_D,plain,(
  ( ! [X2,X1] : (~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1) | p(s(bool,X1))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1777,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f1777_D])).
fof(f1777_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1908,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1)) )),
  inference(condensation,[],[f1907])).
fof(f1907,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,X111)) | ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X112)))),X111) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f1597,f701])).
fof(f1597,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,X110)) | p(s(bool,X111)) | ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X112)))),X111) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f859,f653])).
fof(f1832,plain,(
  ( ! [X2,X1] : (~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1) | p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f1832_D])).
fof(f1832_D,plain,(
  ( ! [X2,X1] : (~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1) | p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1906,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1)) )),
  inference(condensation,[],[f1905])).
fof(f1905,plain,(
  ( ! [X109,X107,X108,X106] : (p(s(bool,X107)) | ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X108)))),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f1596,f701])).
fof(f1596,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,X106)) | p(s(bool,X107)) | ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X108)))),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f816,f653])).
fof(f1829,plain,(
  ( ! [X2,X1] : (~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1) | p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f1829_D])).
fof(f1829_D,plain,(
  ( ! [X2,X1] : (~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1) | p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1904,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1)) )),
  inference(condensation,[],[f1903])).
fof(f1903,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X103)) | ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X104)))),X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f1595,f701])).
fof(f1595,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,X102)) | p(s(bool,X103)) | ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X104)))),X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f799,f653])).
fof(f1826,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f1826_D])).
fof(f1826_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1902,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1901])).
fof(f1901,plain,(
  ( ! [X94,X92,X93,X91] : (~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f1592,f701])).
fof(f1592,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f486,f653])).
fof(f1823,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl36),
  inference(cnf_transformation,[],[f1823_D])).
fof(f1823_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1900,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1899])).
fof(f1899,plain,(
  ( ! [X90,X88,X87,X89] : (~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1591,f701])).
fof(f1591,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f485,f653])).
fof(f1820,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f1820_D])).
fof(f1820_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1898,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1897])).
fof(f1897,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1590,f701])).
fof(f1590,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f482,f653])).
fof(f1817,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f1817_D])).
fof(f1817_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1896,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1895])).
fof(f1895,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1589,f701])).
fof(f1589,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f481,f653])).
fof(f1814,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f1814_D])).
fof(f1814_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1894,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1893])).
fof(f1893,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1588,f701])).
fof(f1588,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f480,f653])).
fof(f1811,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f1811_D])).
fof(f1811_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1892,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f1891])).
fof(f1891,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1587,f701])).
fof(f1587,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f475,f653])).
fof(f1808,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f1808_D])).
fof(f1808_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1890,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1889])).
fof(f1889,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1586,f701])).
fof(f1586,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f465,f653])).
fof(f1805,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f1805_D])).
fof(f1805_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1888,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1887])).
fof(f1887,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1585,f701])).
fof(f1585,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f464,f653])).
fof(f1802,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f1802_D])).
fof(f1802_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1886,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1885])).
fof(f1885,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1584,f701])).
fof(f1584,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f463,f653])).
fof(f1799,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f1799_D])).
fof(f1799_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1884,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1883])).
fof(f1883,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1583,f701])).
fof(f1583,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f458,f653])).
fof(f1796,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f1796_D])).
fof(f1796_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1882,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f1881])).
fof(f1881,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1582,f701])).
fof(f1582,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f450,f653])).
fof(f1793,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f1793_D])).
fof(f1793_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1880,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1879])).
fof(f1879,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1581,f701])).
fof(f1581,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f441,f653])).
fof(f1790,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl14),
  inference(cnf_transformation,[],[f1790_D])).
fof(f1790_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1878,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1877])).
fof(f1877,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1580,f701])).
fof(f1580,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f422,f653])).
fof(f1787,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f1787_D])).
fof(f1787_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1876,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1875])).
fof(f1875,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1579,f701])).
fof(f1579,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f418,f653])).
fof(f1784,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) | $spl10),
  inference(cnf_transformation,[],[f1784_D])).
fof(f1784_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1874,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f1873,f515])).
fof(f1873,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f1872])).
fof(f1872,plain,(
  ( ! [X37,X35,X38,X36] : (sP5(X36,X37,'T_0') | ~sP4(X36,'T_0',X37) | ~sP3(X37,'T_0',X36) | ~sP2('T_0',X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1578,f701])).
fof(f1578,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | sP5(X36,X37,'T_0') | ~sP4(X36,'T_0',X37) | ~sP3(X37,'T_0',X36) | ~sP2('T_0',X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f415,f653])).
fof(f1780,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl8),
  inference(cnf_transformation,[],[f1780_D])).
fof(f1780_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1871,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1870])).
fof(f1870,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1577,f701])).
fof(f1577,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f388,f653])).
fof(f1776,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl4),
  inference(cnf_transformation,[],[f1776_D])).
fof(f1776_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1869,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1868])).
fof(f1868,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f1576,f701])).
fof(f1576,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f385,f653])).
fof(f1834,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1)) )),
  inference(condensation,[],[f1833])).
fof(f1833,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,X111)) | ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X112)))),X111) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f1518,f701])).
fof(f1518,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,X110)) | p(s(bool,X111)) | ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X112)))),X111) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f859,f653])).
fof(f1831,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1)) )),
  inference(condensation,[],[f1830])).
fof(f1830,plain,(
  ( ! [X109,X107,X108,X106] : (p(s(bool,X107)) | ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X108)))),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f1517,f701])).
fof(f1517,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,X106)) | p(s(bool,X107)) | ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X108)))),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f816,f653])).
fof(f1828,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1)) )),
  inference(condensation,[],[f1827])).
fof(f1827,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X103)) | ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X104)))),X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f1516,f701])).
fof(f1516,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,X102)) | p(s(bool,X103)) | ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X104)))),X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f799,f653])).
fof(f1825,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1824])).
fof(f1824,plain,(
  ( ! [X94,X92,X93,X91] : (~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f1513,f701])).
fof(f1513,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f486,f653])).
fof(f1822,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1821])).
fof(f1821,plain,(
  ( ! [X90,X88,X87,X89] : (~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1512,f701])).
fof(f1512,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f485,f653])).
fof(f1819,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1818])).
fof(f1818,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1511,f701])).
fof(f1511,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f482,f653])).
fof(f1816,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1815])).
fof(f1815,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1510,f701])).
fof(f1510,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f481,f653])).
fof(f1813,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1812])).
fof(f1812,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1509,f701])).
fof(f1509,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f480,f653])).
fof(f1810,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f1809])).
fof(f1809,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1508,f701])).
fof(f1508,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f475,f653])).
fof(f1807,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1806])).
fof(f1806,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1507,f701])).
fof(f1507,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f465,f653])).
fof(f1804,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1803])).
fof(f1803,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1506,f701])).
fof(f1506,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f464,f653])).
fof(f1801,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1800])).
fof(f1800,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1505,f701])).
fof(f1505,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f463,f653])).
fof(f1798,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1797])).
fof(f1797,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1504,f701])).
fof(f1504,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f458,f653])).
fof(f1795,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f1794])).
fof(f1794,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1503,f701])).
fof(f1503,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f450,f653])).
fof(f1792,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1791])).
fof(f1791,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1502,f701])).
fof(f1502,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f441,f653])).
fof(f1789,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1788])).
fof(f1788,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1501,f701])).
fof(f1501,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f422,f653])).
fof(f1786,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1785])).
fof(f1785,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1500,f701])).
fof(f1500,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f418,f653])).
fof(f1783,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f1782,f515])).
fof(f1782,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f1781])).
fof(f1781,plain,(
  ( ! [X37,X35,X38,X36] : (sP5(X36,X37,'T_0') | ~sP4(X36,'T_0',X37) | ~sP3(X37,'T_0',X36) | ~sP2('T_0',X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1499,f701])).
fof(f1499,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | sP5(X36,X37,'T_0') | ~sP4(X36,'T_0',X37) | ~sP3(X37,'T_0',X36) | ~sP2('T_0',X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f415,f653])).
fof(f1779,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1778])).
fof(f1778,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1498,f701])).
fof(f1498,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f388,f653])).
fof(f1775,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1774])).
fof(f1774,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f1497,f701])).
fof(f1497,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f385,f653])).
fof(f504,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f503])).
fof(f503,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f498,f499])).
fof(f497,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) | $spl0),
  inference(cnf_transformation,[],[f497_D])).
fof(f497_D,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f499,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) | $spl2),
  inference(cnf_transformation,[],[f499_D])).
fof(f499_D,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f275,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) | s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14)))),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) | s('type.rat.rat',sK12) = s('type.rat.rat',sK13)) & (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) | s('type.rat.rat',sK12) != s('type.rat.rat',sK13))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f174])).
fof(f174,plain,(
  ? [X0,X1,X2] : ((s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) & (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s('type.rat.rat',X0) != s('type.rat.rat',X1)))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ? [X0,X1,X2] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) <~> s('type.rat.rat',X0) = s('type.rat.rat',X1))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ~! [X0,X1,X2] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) <=> s('type.rat.rat',X0) = s('type.rat.rat',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X6),s('type.rat.rat',X8))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))) <=> s('type.rat.rat',X6) = s('type.rat.rat',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X6),s('type.rat.rat',X8))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))) <=> s('type.rat.rat',X6) = s('type.rat.rat',X7))),
  file('thm.rat.RAT_EQ_RADD','thm.rat.RAT_EQ_RADD')).
fof(f498,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) | $spl1),
  inference(cnf_transformation,[],[f498_D])).
fof(f498_D,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f500,plain,(
  s('type.rat.rat',sK12) != s('type.rat.rat',sK13) | $spl3),
  inference(cnf_transformation,[],[f500_D])).
fof(f500_D,plain,(
  s('type.rat.rat',sK12) != s('type.rat.rat',sK13) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f276,plain,(
  s('type.rat.rat',sK12) != s('type.rat.rat',sK13) | s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14)))),
  inference(cnf_transformation,[],[f175])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 349
% Passive clauses: 782
% Generated clauses: 8376
% Final active clauses: 338
% Final passive clauses: 349
% Input formulas: 54
% Initial clauses: 210
% 
% Duplicate literals: 868
% Trivial inequalities: 43
% Fw subsumption resolutions: 521
% Fw demodulations: 237
% Bw demodulations: 29
% Condensations: 1213
% Global subsumptions: 4
% 
% Simple tautologies: 118
% Equational tautologies: 164
% Forward subsumptions: 4126
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 2263
% Factoring: 1
% Forward superposition: 961
% Backward superposition: 1775
% Self superposition: 113
% Equality factoring: 106
% Equality resolution: 4
% 
% Split clauses: 74
% Split components: 153
% SAT solver clauses: 1202
% SAT solver unit clauses: 358
% SAT solver binary clauses: 269
% 
% Sat splits: 74
% Sat splitting refutations: 266
% 
% Memory used [KB]: 5117
% Time elapsed: 0.507 s
% ------------------------------
% ------------------------------
% Success in time 1.814 s
