Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real
% (67499)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 145
% Passive clauses: 2870
% Generated clauses: 23681
% Final active clauses: 141
% Final passive clauses: 2398
% Input formulas: 33
% Initial clauses: 72
% 
% Duplicate literals: 1069
% Trivial inequalities: 679
% Fw subsumption resolutions: 3193
% Fw demodulations: 2924
% 
% Simple tautologies: 613
% Equational tautologies: 48
% Forward subsumptions: 11199
% 
% Binary resolution: 142
% Factoring: 2
% Forward superposition: 7132
% Backward superposition: 8017
% Self superposition: 190
% Equality factoring: 106
% Equality resolution: 49
% 
% Split clauses: 555
% Split components: 1152
% SAT solver clauses: 1077
% SAT solver unit clauses: 48
% SAT solver binary clauses: 170
% 
% Sat splits: 555
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1055
% TWLsolver calls for satisfiability: 113
% 
% Memory used [KB]: 10618
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% (67537)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 125
% Passive clauses: 810
% Generated clauses: 14180
% Final active clauses: 105
% Final passive clauses: 519
% Input formulas: 33
% Initial clauses: 72
% 
% Duplicate literals: 687
% Trivial inequalities: 227
% Fw subsumption resolutions: 1021
% Fw demodulations: 1122
% Bw demodulations: 145
% Condensations: 1472
% Global subsumptions: 9
% 
% Simple tautologies: 268
% Equational tautologies: 39
% Forward subsumptions: 7066
% Bw demodulations to eq. taut.: 23
% 
% Binary resolution: 175
% Factoring: 1
% Forward superposition: 4414
% Backward superposition: 3579
% Self superposition: 128
% Equality factoring: 63
% Equality resolution: 35
% 
% Split clauses: 122
% Split components: 274
% SAT solver clauses: 1498
% SAT solver unit clauses: 74
% SAT solver binary clauses: 47
% SAT solver learnt clauses: 18
% 
% Sat splits: 122
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 1462
% TWLsolver calls for satisfiability: 1044
% 
% Memory used [KB]: 10106
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f10434,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f121,f194,f120,f195,f119,f196,f118,f197,f193,f117,f198,f116,f199,f191,f115,f200,f189,f114,f201,f192,f190,f188,f1509,f1513,f1510,f1554,f1555,f1556,f1558,f1561,f1562,f1564,f1565,f2028,f2031,f2029,f2034,f2035,f2030,f2042,f2044,f2120,f2122,f2126,f2128,f2130,f2179,f2180,f2182,f2185,f2186,f2188,f2189,f2190,f2191,f2192,f2193,f2194,f2195,f2196,f2200,f2204,f2208,f2212,f2215,f2219,f2223,f2227,f2231,f2235,f2238,f2242,f2246,f2249,f2252,f2255,f2257,f2260,f2264,f2267,f2270,f2273,f2275,f2278,f2282,f2285,f2288,f2291,f2293,f2296,f2300,f2303,f2306,f2309,f2311,f2314,f2318,f2322,f2326,f2330,f2333,f2337,f2341,f2345,f2349,f2353,f2356,f2360,f2364,f2367,f2370,f2373,f2375,f2378,f2382,f2385,f2388,f2391,f2393,f2396,f2400,f2403,f2406,f2409,f2411,f2414,f2418,f2421,f2424,f2427,f2429,f2432,f2433,f5516,f5529,f5530,f5756,f6713,f6717,f6721,f6718,f6727,f6728,f6729,f6730,f6737,f6745,f6746,f6797,f6798,f7840,f7853,f7903,f7916,f7969,f7973,f7970,f8041,f8046,f8051,f8047,f1511,f8109,f10361,f10376,f10423,f10432,f10433])).
fof(f10433,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706,f5405,f3707,f3709,f3724,f5470,f185,f5545,f5544,f5758,f5694,f5703,f5705,f5707,f5710,f5711,f5712,f5714,f5717,f5718,f5719,f5722,f5723,f5729,f5732,f5734,f5737,f190,f6800,f6802,f6804,f6806,f6808,f6810,f6812,f6814,f6816,f6818,f6822,f6824,f6826,f6828,f6832,f6834,f6836,f6838,f10122,f10139,f10140,f10211,f10213,f10220,f10233,f10234,f10235,f10278,f10280,f10282,f10291,f10292,f10293,f10294,f10295,f10296,f10297,f10299,f10301,f10355,f10307,f10309,f10356,f10311,f10312,f10318,f10321,f10323,f10326,f188,f10383,f10385,f10387,f10391,f10393,f10395,f10397,f10401,f10403,f10405,f10407,f10409,f10411,f10413,f10415,f10417,f10419,f10421,f10422,f193])).
fof(f10422,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f190,f188])).
fof(f10421,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s('type.realax.real',sK0) = s('type.realax.real',X90) | s('type.realax.real',sK0) = s('type.realax.real',X91) | s('type.realax.real',X93) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X89),s('type.realax.real',X91))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X92),s('type.realax.real',X90))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X93),s('type.realax.real',X94))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X89),s('type.realax.real',X90))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X92))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X90))))),s('type.realax.real',X94)))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10420,f188])).
fof(f10420,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s('type.realax.real',sK0) = s('type.realax.real',X91) | s('type.realax.real',sK2) = s('type.realax.real',X90) | s('type.realax.real',X93) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X89),s('type.realax.real',X91))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X92),s('type.realax.real',X90))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X93),s('type.realax.real',X94))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X89),s('type.realax.real',X90))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X92))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X90))))),s('type.realax.real',X94)))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6800,f188])).
fof(f10419,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s('type.realax.real',sK0) = s('type.realax.real',X72) | s('type.realax.real',sK0) = s('type.realax.real',X73) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X75),s('type.realax.real',X76))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X71),s('type.realax.real',X73))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X74),s('type.realax.real',X72))))),s('type.realax.real',X76))) | s('type.realax.real',X75) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X71),s('type.realax.real',X72))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X74))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X72)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10418,f188])).
fof(f10418,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s('type.realax.real',sK0) = s('type.realax.real',X73) | s('type.realax.real',sK2) = s('type.realax.real',X72) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X75),s('type.realax.real',X76))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X71),s('type.realax.real',X73))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X74),s('type.realax.real',X72))))),s('type.realax.real',X76))) | s('type.realax.real',X75) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X71),s('type.realax.real',X72))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X74))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X72)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6802,f188])).
fof(f10417,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s('type.realax.real',sK0) = s('type.realax.real',X60) | s('type.realax.real',sK0) = s('type.realax.real',X61) | s('type.realax.real',X63) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X59),s('type.realax.real',X61))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X62),s('type.realax.real',X60))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real',X63))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X59),s('type.realax.real',X60))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X62))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X60)))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10416,f188])).
fof(f10416,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s('type.realax.real',sK0) = s('type.realax.real',X61) | s('type.realax.real',sK2) = s('type.realax.real',X60) | s('type.realax.real',X63) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X59),s('type.realax.real',X61))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X62),s('type.realax.real',X60))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real',X63))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X59),s('type.realax.real',X60))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X62))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X60)))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6804,f188])).
fof(f10415,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.realax.real',sK0) = s('type.realax.real',X42) | s('type.realax.real',sK0) = s('type.realax.real',X43) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real',X46))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X41),s('type.realax.real',X43))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X44),s('type.realax.real',X42))))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X41),s('type.realax.real',X42))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X44))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X42)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10414,f188])).
fof(f10414,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.realax.real',sK0) = s('type.realax.real',X43) | s('type.realax.real',sK2) = s('type.realax.real',X42) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real',X46))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X41),s('type.realax.real',X43))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X44),s('type.realax.real',X42))))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X41),s('type.realax.real',X42))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X44))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X42)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6806,f188])).
fof(f10413,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.realax.real',sK0) = s('type.realax.real',X9) | s('type.realax.real',sK0) = s('type.realax.real',X10) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X11))))),s('type.realax.real',X12))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X10))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X11),s('type.realax.real',X9))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X12),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9)))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10412,f188])).
fof(f10412,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.realax.real',sK0) = s('type.realax.real',X10) | s('type.realax.real',sK2) = s('type.realax.real',X9) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X11))))),s('type.realax.real',X12))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X10))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X11),s('type.realax.real',X9))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X12),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9)))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6808,f188])).
fof(f10411,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.realax.real',sK0) = s('type.realax.real',X4) | s('type.realax.real',sK0) = s('type.realax.real',X5) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X4)))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10410,f188])).
fof(f10410,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.realax.real',sK0) = s('type.realax.real',X5) | s('type.realax.real',sK2) = s('type.realax.real',X4) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X4)))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6810,f188])).
fof(f10409,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10408,f188])).
fof(f10408,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real',sK2) = s('type.realax.real',X1) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6812,f188])).
fof(f10407,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33))) | s('type.realax.real',sK0) = s('type.realax.real',X35) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X35))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X36),s('type.realax.real',X33))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X36))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X33)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10406,f188])).
fof(f10406,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real',sK0) = s('type.realax.real',X35) | s('type.realax.real',sK2) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X35))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X36),s('type.realax.real',X33))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X36))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X33)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6814,f188])).
fof(f10405,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28))) | s('type.realax.real',sK0) = s('type.realax.real',X30) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X30))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X27))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X31))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X27)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10404,f188])).
fof(f10404,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real',sK0) = s('type.realax.real',X30) | s('type.realax.real',sK2) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X30))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X27))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X31))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X27)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6816,f188])).
fof(f10403,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real',sK0) = s('type.realax.real',X16) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real',X16))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X17),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real',X17))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13)))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10402,f188])).
fof(f10402,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real',sK0) = s('type.realax.real',X16) | s('type.realax.real',sK2) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real',X16))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X17),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real',X17))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13)))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6818,f188])).
fof(f10401,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',sK0) = s('type.realax.real',X5) | s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1)))))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10400,f188])).
fof(f10400,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X5) | s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real',sK2) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1)))))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10399,f188])).
fof(f10399,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real',sK2) = s('type.realax.real',X5) | s('type.realax.real',sK2) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1)))))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10398,f188])).
fof(f10398,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real',sK2) = s('type.realax.real',X1) | s('type.realax.real',sK2) = s('type.realax.real',X5) | s('type.realax.real',sK2) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1)))))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6822,f188])).
fof(f10397,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real',sK0) = s('type.realax.real',X36) | s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X35),s('type.realax.real',X36))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X35))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X36)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10396,f188])).
fof(f10396,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33))) | s('type.realax.real',sK2) = s('type.realax.real',X36) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X35),s('type.realax.real',X36))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X35))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X36)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6824,f188])).
fof(f10395,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real',sK0) = s('type.realax.real',X31) | s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X30),s('type.realax.real',X31))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X31))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X30))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X31)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10394,f188])).
fof(f10394,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28))) | s('type.realax.real',sK2) = s('type.realax.real',X31) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X30),s('type.realax.real',X31))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X31))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X30))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X31)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6826,f188])).
fof(f10393,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real',sK0) = s('type.realax.real',X17) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X17))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X17))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X16))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X17)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10392,f188])).
fof(f10392,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real',sK2) = s('type.realax.real',X17) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X17))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X17))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X16))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X17)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6828,f188])).
fof(f10391,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X6) | s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10390,f188])).
fof(f10390,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real',sK2) = s('type.realax.real',X6) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10389,f188])).
fof(f10389,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real',sK2) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',sK2) = s('type.realax.real',X6) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10388,f188])).
fof(f10388,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real',sK2) = s('type.realax.real',X1) | s('type.realax.real',sK2) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',sK2) = s('type.realax.real',X6) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6832,f188])).
fof(f10387,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X3) | s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10386,f188])).
fof(f10386,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real',sK2) = s('type.realax.real',X3) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6834,f188])).
fof(f10385,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10384,f188])).
fof(f10384,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real',sK2) = s('type.realax.real',X1) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6836,f188])).
fof(f10383,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X3) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3)))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10382,f188])).
fof(f10382,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X3) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) | s('type.realax.real',sK2) = s('type.realax.real',X1)) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f6838,f188])).
fof(f10326,plain,(
  ( ! [X204,X202,X200,X203,X201] : (s('type.realax.real',X202) != s('type.realax.real',X203) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X203),s('type.realax.real',X204))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.bool.COND_3'(s(bool,X200),s('type.realax.real',X201),s('type.realax.real',X202))),s('type.realax.real',X204))) | s(bool,'T_0') = s(bool,X200)) )),
  inference(superposition,[],[f178,f308])).
fof(f10323,plain,(
  ( ! [X189,X187,X185,X188,X186] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X187),s('type.realax.real',X189))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X188),s('type.realax.real',X189))) | s('type.realax.real',X188) = s('type.realax.real','const.bool.COND_3'(s(bool,X185),s('type.realax.real',X186),s('type.realax.real',X187))) | s(bool,'T_0') = s(bool,X185)) )),
  inference(superposition,[],[f177,f308])).
fof(f10321,plain,(
  ( ! [X177,X175,X178,X176,X179] : (s('type.realax.real',X177) != s('type.realax.real',X178) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X179),s('type.realax.real',X178))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X179),s('type.realax.real','const.bool.COND_3'(s(bool,X175),s('type.realax.real',X176),s('type.realax.real',X177))))) | s(bool,'T_0') = s(bool,X175)) )),
  inference(superposition,[],[f176,f308])).
fof(f10318,plain,(
  ( ! [X163,X161,X164,X162,X160] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X163),s('type.realax.real',X162))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X163),s('type.realax.real',X164))) | s('type.realax.real',X164) = s('type.realax.real','const.bool.COND_3'(s(bool,X160),s('type.realax.real',X161),s('type.realax.real',X162))) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f175,f308])).
fof(f10312,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,'F_0') != s(bool,X135) | s(bool,X136) = s(bool,'const.bool.COND_3'(s(bool,X133),s(bool,X134),s(bool,X135))) | s(bool,'T_0') = s(bool,X136) | s(bool,'T_0') = s(bool,X133)) )),
  inference(superposition,[],[f877,f308])).
fof(f10311,plain,(
  ( ! [X132,X130,X131,X129] : (~p(s(bool,X131)) | s(bool,X132) = s(bool,'const.bool.COND_3'(s(bool,X129),s(bool,X130),s(bool,X131))) | s(bool,'F_0') = s(bool,X132) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f370,f308])).
fof(f10356,plain,(
  ( ! [X127,X128,X126] : (s(bool,'F_0') != s(bool,X128) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X126),s(bool,X127),s(bool,X128))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(forward_demodulation,[],[f10310,f157])).
fof(f10310,plain,(
  ( ! [X127,X128,X126] : (s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X128))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X126),s(bool,X127),s(bool,X128))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f368,f308])).
fof(f10309,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X124),s(bool,X125))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X122),s(bool,X123),s(bool,X124))) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f287,f308])).
fof(f10307,plain,(
  ( ! [X116,X114,X117,X115] : (p(s(bool,X116)) | s(bool,X117) = s(bool,'const.bool.COND_3'(s(bool,X114),s(bool,X115),s(bool,X116))) | s(bool,'T_0') = s(bool,X117) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f263,f308])).
fof(f10355,plain,(
  ( ! [X111,X112,X113] : (s(bool,'T_0') = s(bool,X113) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X112),s(bool,X113))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(forward_demodulation,[],[f10306,f157])).
fof(f10306,plain,(
  ( ! [X111,X112,X113] : (s(bool,'T_0') = s(bool,'const.marker.unint_1'(s(bool,X113))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X112),s(bool,X113))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f233,f308])).
fof(f10301,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X93)) | ~p(s(bool,X94)) | s(bool,X94) = s(bool,'const.bool.COND_3'(s(bool,X91),s(bool,X92),s(bool,X93))) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f172,f308])).
fof(f10299,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X85)) | p(s(bool,X86)) | s(bool,X86) = s(bool,'const.bool.COND_3'(s(bool,X83),s(bool,X84),s(bool,X85))) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f169,f308])).
fof(f10297,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X65,X73] : (s(bool,X67) != s(bool,X68) | s(X69,X70) != s(X69,X71) | s(X69,X72) != s(X69,X73) | s(X69,'const.bool.COND_3'(s(bool,X68),s(X69,X72),s(X69,X70))) = s(X69,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X65),s(bool,X66),s(bool,X67))),s(X69,X73),s(X69,X71))) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f167,f308])).
fof(f10296,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X63] : (p(s(bool,X58)) | s(X59,X60) != s(X59,X61) | s(X59,'const.bool.COND_3'(s(bool,X62),s(X59,X63),s(X59,X60))) = s(X59,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X56),s(bool,X57),s(bool,X58))),s(X59,X64),s(X59,X61))) | s(bool,X62) != s(bool,'const.bool.COND_3'(s(bool,X56),s(bool,X57),s(bool,X58))) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f166,f308])).
fof(f10295,plain,(
  ( ! [X47,X54,X52,X50,X48,X55,X53,X51,X49] : (~p(s(bool,X49)) | s(X50,'const.bool.COND_3'(s(bool,X51),s(X50,X52),s(X50,X53))) = s(X50,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X47),s(bool,X48),s(bool,X49))),s(X50,X54),s(X50,X55))) | s(X50,X52) != s(X50,X54) | s(bool,X51) != s(bool,'const.bool.COND_3'(s(bool,X47),s(bool,X48),s(bool,X49))) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f165,f308])).
fof(f10294,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X46)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X44),s(bool,X45),s(bool,X46))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f137,f308])).
fof(f10293,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,X43)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X41),s(bool,X42),s(bool,X43))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f136,f308])).
fof(f10292,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,X40)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X38),s(bool,X39),s(bool,X40))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f133,f308])).
fof(f10291,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X37)) | s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X35),s(bool,X36),s(bool,X37))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f132,f308])).
fof(f10282,plain,(
  ( ! [X30,X28,X29] : (s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X28),s(bool,X29),s(bool,X30)))) )),
  inference(superposition,[],[f308,f130])).
fof(f10280,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X20),s(bool,X21),s(bool,X22))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f308,f231])).
fof(f10278,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X12) | s(bool,X15) = s(bool,'const.bool.COND_3'(s(bool,X12),s(bool,X13),s(bool,X14))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f308,f231])).
fof(f10235,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s(X62,X63) = s(X62,'const.bool.COND_3'(s(bool,X61),s(X62,X64),s(X62,X63))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X59),s(bool,X60),s(bool,X61))) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f308,f308])).
fof(f10234,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(X56,X57) = s(X56,'const.bool.COND_3'(s(bool,X55),s(X56,X58),s(X56,X57))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X54),s(bool,X55)))) )),
  inference(superposition,[],[f308,f161])).
fof(f10233,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(X51,X52) = s(X51,'const.bool.COND_3'(s(bool,X49),s(X51,X53),s(X51,X52))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X49),s(bool,X50)))) )),
  inference(superposition,[],[f308,f160])).
fof(f10220,plain,(
  ( ! [X70,X68,X69] : (s(bool,X70) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X68),s(bool,X69))) | s(bool,'F_0') = s(bool,X70) | s(bool,'F_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f10175,f122])).
fof(f10175,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,'T_0')) | s(bool,X70) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X68),s(bool,X69))) | s(bool,'F_0') = s(bool,X70) | s(bool,'F_0') = s(bool,X68)) )),
  inference(superposition,[],[f370,f287])).
fof(f10213,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X12,'const.bool.COND_3'(s(bool,X10),s(X12,X16),s(X12,X17))) = s(X12,'const.bool.COND_3'(s(bool,X13),s(X12,X14),s(X12,X15))) | s(X12,X14) != s(X12,X16) | s(bool,X13) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(forward_demodulation,[],[f10212,f160])).
fof(f10212,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X12,'const.bool.COND_3'(s(bool,X13),s(X12,X14),s(X12,X15))) = s(X12,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))),s(X12,X16),s(X12,X17))) | s(X12,X14) != s(X12,X16) | s(bool,X13) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f10163,f122])).
fof(f10163,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,'T_0')) | s(X12,'const.bool.COND_3'(s(bool,X13),s(X12,X14),s(X12,X15))) = s(X12,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))),s(X12,X16),s(X12,X17))) | s(X12,X14) != s(X12,X16) | s(bool,X13) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f165,f287])).
fof(f10211,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f10161,f122])).
fof(f10161,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f136,f287])).
fof(f10140,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X33),s(bool,X34))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X33)))) )),
  inference(superposition,[],[f287,f161])).
fof(f10139,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X29),s(bool,X31))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X29),s(bool,X30)))) )),
  inference(superposition,[],[f287,f160])).
fof(f10122,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f287,f231])).
fof(f6838,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) | s('type.realax.real',sK2) = s('type.realax.real',X3) | s('type.realax.real',sK2) = s('type.realax.real',X1)) ) | $spl2),
  inference(forward_demodulation,[],[f6837,f190])).
fof(f6837,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X3) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f185,f190])).
fof(f6836,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X2) | s('type.realax.real',sK2) = s('type.realax.real',X1) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) ) | $spl2),
  inference(forward_demodulation,[],[f6835,f190])).
fof(f6835,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X1) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5694,f190])).
fof(f6834,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X0) | s('type.realax.real',sK2) = s('type.realax.real',X3) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3)))))) ) | $spl2),
  inference(forward_demodulation,[],[f6833,f190])).
fof(f6833,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X3) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5703,f190])).
fof(f6832,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X2) | s('type.realax.real',sK2) = s('type.realax.real',X1) | s('type.realax.real',sK2) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',sK2) = s('type.realax.real',X6) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6)))))) ) | $spl2),
  inference(forward_demodulation,[],[f6831,f190])).
fof(f6831,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X1) | s('type.realax.real',sK2) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',sK2) = s('type.realax.real',X6) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f6830,f190])).
fof(f6830,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',sK2) = s('type.realax.real',X6) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f6829,f190])).
fof(f6829,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X6) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5705,f190])).
fof(f6828,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real',sK2) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real',sK2) = s('type.realax.real',X17) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X17))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X17))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X16))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X17)))))) ) | $spl2),
  inference(forward_demodulation,[],[f6827,f190])).
fof(f6827,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real',sK2) = s('type.realax.real',X17) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X17))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X17))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X16))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X17))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14)))) ) | $spl2),
  inference(forward_demodulation,[],[f5707,f190])).
fof(f6826,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real',sK2) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28))) | s('type.realax.real',sK2) = s('type.realax.real',X31) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X30),s('type.realax.real',X31))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X31))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X30))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X31)))))) ) | $spl2),
  inference(forward_demodulation,[],[f6825,f190])).
fof(f6825,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real',sK2) = s('type.realax.real',X31) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X30),s('type.realax.real',X31))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X31))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X30))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X31))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28)))) ) | $spl2),
  inference(forward_demodulation,[],[f5710,f190])).
fof(f6824,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real',sK2) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33))) | s('type.realax.real',sK2) = s('type.realax.real',X36) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X35),s('type.realax.real',X36))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X35))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X36)))))) ) | $spl2),
  inference(forward_demodulation,[],[f6823,f190])).
fof(f6823,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real',sK2) = s('type.realax.real',X36) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X35),s('type.realax.real',X36))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X35))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X36))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33)))) ) | $spl2),
  inference(forward_demodulation,[],[f5711,f190])).
fof(f6822,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X2) | s('type.realax.real',sK2) = s('type.realax.real',X1) | s('type.realax.real',sK2) = s('type.realax.real',X5) | s('type.realax.real',sK2) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1)))))))))) ) | $spl2),
  inference(forward_demodulation,[],[f6821,f190])).
fof(f6821,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X1) | s('type.realax.real',sK2) = s('type.realax.real',X5) | s('type.realax.real',sK2) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f6820,f190])).
fof(f6820,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X5) | s('type.realax.real',sK2) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f6819,f190])).
fof(f6819,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5712,f190])).
fof(f6818,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real',sK2) = s('type.realax.real',X16) | s('type.realax.real',sK2) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real',X16))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X17),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real',X17))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f6817,f190])).
fof(f6817,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real',sK2) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real',X16))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X17),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real',X17))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))))) | s('type.realax.real',X16) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5714,f190])).
fof(f6816,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real',sK2) = s('type.realax.real',X30) | s('type.realax.real',sK2) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X30))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X27))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X31))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X27)))))) ) | $spl2),
  inference(forward_demodulation,[],[f6815,f190])).
fof(f6815,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real',sK2) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X30))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X27))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X31))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X27))))) | s('type.realax.real',X30) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5717,f190])).
fof(f6814,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real',sK2) = s('type.realax.real',X35) | s('type.realax.real',sK2) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X35))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X36),s('type.realax.real',X33))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X36))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X33)))))) ) | $spl2),
  inference(forward_demodulation,[],[f6813,f190])).
fof(f6813,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real',sK2) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X35))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X36),s('type.realax.real',X33))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X36))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X33))))) | s('type.realax.real',X35) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5718,f190])).
fof(f6812,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X0) | s('type.realax.real',sK2) = s('type.realax.real',X1) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))) ) | $spl2),
  inference(forward_demodulation,[],[f6811,f190])).
fof(f6811,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK2) = s('type.realax.real',X1) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5719,f190])).
fof(f6810,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.realax.real',sK2) = s('type.realax.real',X5) | s('type.realax.real',sK2) = s('type.realax.real',X4) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X4)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f6809,f190])).
fof(f6809,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.realax.real',sK2) = s('type.realax.real',X4) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X4))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5722,f190])).
fof(f6808,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.realax.real',sK2) = s('type.realax.real',X10) | s('type.realax.real',sK2) = s('type.realax.real',X9) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X11))))),s('type.realax.real',X12))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X10))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X11),s('type.realax.real',X9))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X12),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f6807,f190])).
fof(f6807,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.realax.real',sK2) = s('type.realax.real',X9) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X11))))),s('type.realax.real',X12))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X10))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X11),s('type.realax.real',X9))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X12),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9))))))) | s('type.realax.real',X10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5723,f190])).
fof(f6806,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.realax.real',sK2) = s('type.realax.real',X43) | s('type.realax.real',sK2) = s('type.realax.real',X42) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real',X46))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X41),s('type.realax.real',X43))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X44),s('type.realax.real',X42))))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X41),s('type.realax.real',X42))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X44))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X42)))))) ) | $spl2),
  inference(forward_demodulation,[],[f6805,f190])).
fof(f6805,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.realax.real',sK2) = s('type.realax.real',X42) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real',X46))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X41),s('type.realax.real',X43))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X44),s('type.realax.real',X42))))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X41),s('type.realax.real',X42))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X44))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X42))))) | s('type.realax.real',X43) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5729,f190])).
fof(f6804,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s('type.realax.real',sK2) = s('type.realax.real',X61) | s('type.realax.real',sK2) = s('type.realax.real',X60) | s('type.realax.real',X63) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X59),s('type.realax.real',X61))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X62),s('type.realax.real',X60))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real',X63))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X59),s('type.realax.real',X60))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X62))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X60)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f6803,f190])).
fof(f6803,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s('type.realax.real',sK2) = s('type.realax.real',X60) | s('type.realax.real',X63) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X59),s('type.realax.real',X61))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X62),s('type.realax.real',X60))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real',X63))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X59),s('type.realax.real',X60))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X62))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X60))))))) | s('type.realax.real',X61) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5732,f190])).
fof(f6802,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s('type.realax.real',sK2) = s('type.realax.real',X73) | s('type.realax.real',sK2) = s('type.realax.real',X72) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X75),s('type.realax.real',X76))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X71),s('type.realax.real',X73))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X74),s('type.realax.real',X72))))),s('type.realax.real',X76))) | s('type.realax.real',X75) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X71),s('type.realax.real',X72))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X74))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X72)))))) ) | $spl2),
  inference(forward_demodulation,[],[f6801,f190])).
fof(f6801,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s('type.realax.real',sK2) = s('type.realax.real',X72) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X75),s('type.realax.real',X76))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X71),s('type.realax.real',X73))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X74),s('type.realax.real',X72))))),s('type.realax.real',X76))) | s('type.realax.real',X75) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X71),s('type.realax.real',X72))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X74))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X72))))) | s('type.realax.real',X73) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5734,f190])).
fof(f6800,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s('type.realax.real',sK2) = s('type.realax.real',X91) | s('type.realax.real',sK2) = s('type.realax.real',X90) | s('type.realax.real',X93) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X89),s('type.realax.real',X91))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X92),s('type.realax.real',X90))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X93),s('type.realax.real',X94))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X89),s('type.realax.real',X90))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X92))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X90))))),s('type.realax.real',X94)))) ) | $spl2),
  inference(forward_demodulation,[],[f6799,f190])).
fof(f6799,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s('type.realax.real',sK2) = s('type.realax.real',X90) | s('type.realax.real',X93) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X89),s('type.realax.real',X91))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X92),s('type.realax.real',X90))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X93),s('type.realax.real',X94))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X89),s('type.realax.real',X90))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X92))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X90))))),s('type.realax.real',X94))) | s('type.realax.real',X91) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f5737,f190])).
fof(f5737,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s('type.realax.real',X93) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X89),s('type.realax.real',X91))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X92),s('type.realax.real',X90))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X93),s('type.realax.real',X94))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X89),s('type.realax.real',X90))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X92))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X90))))),s('type.realax.real',X94))) | s('type.realax.real',X90) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X91) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f178,f185])).
fof(f5734,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X75),s('type.realax.real',X76))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X71),s('type.realax.real',X73))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X74),s('type.realax.real',X72))))),s('type.realax.real',X76))) | s('type.realax.real',X75) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X71),s('type.realax.real',X72))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X74))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X72))))) | s('type.realax.real',X72) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X73) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f177,f185])).
fof(f5732,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s('type.realax.real',X63) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X59),s('type.realax.real',X61))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X62),s('type.realax.real',X60))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real',X63))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X59),s('type.realax.real',X60))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X62))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X60))))))) | s('type.realax.real',X60) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X61) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f176,f185])).
fof(f5729,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real',X46))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X41),s('type.realax.real',X43))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X44),s('type.realax.real',X42))))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X41),s('type.realax.real',X42))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X44))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X42))))) | s('type.realax.real',X42) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X43) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f175,f185])).
fof(f5723,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X11))))),s('type.realax.real',X12))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X10))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X11),s('type.realax.real',X9))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X12),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9))))))) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f174,f185])).
fof(f5722,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X4))))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f174,f185])).
fof(f5719,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f185,f168])).
fof(f5718,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X35))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X36),s('type.realax.real',X33))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X36))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X33))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33))) | s('type.realax.real',X35) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f185,f161])).
fof(f5717,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X30))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X27))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X31))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X27))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28))) | s('type.realax.real',X30) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f185,f160])).
fof(f5714,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real',X16))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X17),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real',X17))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real',X16) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f185,f168])).
fof(f5712,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f185,f185])).
fof(f5711,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X35),s('type.realax.real',X36))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X35))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X36))))) | s('type.realax.real',X36) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33)))) )),
  inference(superposition,[],[f185,f161])).
fof(f5710,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X30),s('type.realax.real',X31))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X31))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X30))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X31))))) | s('type.realax.real',X31) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28)))) )),
  inference(superposition,[],[f185,f160])).
fof(f5707,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X17))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X17))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X16))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X17))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14)))) )),
  inference(superposition,[],[f185,f168])).
fof(f5705,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f185,f185])).
fof(f5703,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f185,f168])).
fof(f5694,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f185,f168])).
fof(f5758,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',sK0))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0)))))) ) | $spl5),
  inference(forward_demodulation,[],[f5683,f174])).
fof(f5683,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',sK0))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0)))))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f193,f193,f185])).
fof(f5544,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X0))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f193,f177])).
fof(f5545,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK0))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f193,f175])).
fof(f185,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3] : (s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3))))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3] : ((s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3))))))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : ((s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real',X3) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) => s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3] : ((~s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & ~s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) => s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3))))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X23,X24,X25,X26] : ((~s('type.realax.real',X24) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & ~s('type.realax.real',X26) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) => s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X23),s('type.realax.real',X24))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X25),s('type.realax.real',X26))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X23),s('type.realax.real',X26))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X24),s('type.realax.real',X25))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X24),s('type.realax.real',X26))))))),
  file('thm.real.add_rat','thm.real.REAL_ADD_RAT')).
fof(f5470,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20] : (s(X21,'const.bool.COND_3'(s(bool,X19),s(X21,X25),s(X21,X23))) = s(X21,'const.bool.COND_3'(s(bool,X20),s(X21,X24),s(X21,X22))) | s(bool,'T_0') != s(bool,X20) | s(X21,X22) != s(X21,X23) | s(X21,X24) != s(X21,X25) | s(bool,'F_0') = s(bool,X19)) )),
  inference(forward_demodulation,[],[f5436,f157])).
fof(f5436,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20] : (s(bool,'T_0') != s(bool,X20) | s(X21,X22) != s(X21,X23) | s(X21,X24) != s(X21,X25) | s(X21,'const.bool.COND_3'(s(bool,X20),s(X21,X24),s(X21,X22))) = s(X21,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X19))),s(X21,X25),s(X21,X23))) | s(bool,'F_0') = s(bool,X19)) )),
  inference(superposition,[],[f167,f233])).
fof(f3724,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X1) != s(X0,X2) | s(X0,X3) != s(X0,X4) | s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X1))) = s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X4),s(X0,X2)))) )),
  inference(equality_resolution,[],[f167])).
fof(f3709,plain,(
  ( ! [X61,X59,X57,X60,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | s(X57,X58) != s(X57,X59) | s(X57,X60) != s(X57,X61) | s(X57,'const.bool.COND_3'(s(bool,'F_0'),s(X57,X60),s(X57,X58))) = s(X57,'const.bool.COND_3'(s(bool,X56),s(X57,X61),s(X57,X59))) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f167,f130])).
fof(f3707,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s(bool,X39) != s(bool,X40) | s(X41,X42) != s(X41,X43) | s(X41,X44) != s(X41,X45) | s(X41,'const.bool.COND_3'(s(bool,'T_0'),s(X41,X44),s(X41,X42))) = s(X41,'const.bool.COND_3'(s(bool,X40),s(X41,X45),s(X41,X43))) | s(bool,X39) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f167,f231])).
fof(f5405,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'T_0') = s(bool,X32) | s(X34,'const.bool.COND_3'(s(bool,X32),s(X34,X37),s(X34,X35))) = s(X34,'const.bool.COND_3'(s(bool,X33),s(X34,X38),s(X34,X36))) | s(X34,X35) != s(X34,X36) | s(bool,'F_0') != s(bool,X33)) )),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706])).
fof(f3706,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'F_0') != s(bool,X33) | s(X34,X35) != s(X34,X36) | s(X34,X37) != s(X34,X38) | s(X34,'const.bool.COND_3'(s(bool,X32),s(X34,X37),s(X34,X35))) = s(X34,'const.bool.COND_3'(s(bool,X33),s(X34,X38),s(X34,X36))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f167,f130])).
fof(f3704,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X17) != s(bool,X18) | s(X19,X20) != s(X19,X21) | s(X19,X22) != s(X19,X23) | s(X19,'const.bool.COND_3'(s(bool,X16),s(X19,X22),s(X19,X20))) = s(X19,'const.bool.COND_3'(s(bool,X18),s(X19,X23),s(X19,X21))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f167,f231])).
fof(f3702,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(X2,X3) != s(X2,X4) | s(X2,X5) != s(X2,X6) | s(X2,'const.bool.COND_3'(s(bool,X0),s(X2,X5),s(X2,X3))) = s(X2,'const.bool.COND_3'(s(bool,X1),s(X2,X6),s(X2,X4))) | s(bool,X0) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f167,f231])).
fof(f3829,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X3)))))) )),
  inference(forward_demodulation,[],[f3649,f161])).
fof(f3649,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X3)))))) )),
  inference(unit_resulting_resolution,[],[f168,f168,f161,f167])).
fof(f3822,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X4)))) )),
  inference(forward_demodulation,[],[f3821,f161])).
fof(f3821,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X3),s('type.realax.real',X4)))))) )),
  inference(forward_demodulation,[],[f3653,f161])).
fof(f3653,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X3),s('type.realax.real',X4))))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X4)))) )),
  inference(unit_resulting_resolution,[],[f161,f168,f161,f167])).
fof(f3734,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X2)))))) )),
  inference(forward_demodulation,[],[f3733,f161])).
fof(f3733,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X2)))))) )),
  inference(forward_demodulation,[],[f3697,f161])).
fof(f3697,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X5),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X2)))))) )),
  inference(unit_resulting_resolution,[],[f168,f161,f161,f167])).
fof(f136,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.real.add_rat','thm.bool.EQ_CLAUSES|split|2')).
fof(f137,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f94])).
fof(f161,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X16] : ! [X13,X14] : s(X16,X14) = s(X16,'const.bool.COND_3'(s(bool,'F_0'),s(X16,X13),s(X16,X14)))),
  file('thm.real.add_rat','thm.bool.COND_CLAUSES|split|1')).
fof(f202,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f125,f122])).
fof(f125,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f17])).
fof(f17,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.real.add_rat','thm.bool.NOT_CLAUSES|split|1')).
fof(f233,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.marker.unint_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f130,f157])).
fof(f287,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X3),s(bool,X4))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f160,f130])).
fof(f322,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X4)))) )),
  inference(superposition,[],[f161,f130])).
fof(f366,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(forward_demodulation,[],[f348,f157])).
fof(f348,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'const.marker.unint_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(superposition,[],[f233,f160])).
fof(f367,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7)))) )),
  inference(forward_demodulation,[],[f349,f157])).
fof(f349,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.marker.unint_1'(s(bool,X7))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7)))) )),
  inference(superposition,[],[f233,f161])).
fof(f368,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f355,f122])).
fof(f355,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f136,f233])).
fof(f370,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f369,f157])).
fof(f369,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.marker.unint_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f359,f122])).
fof(f359,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.marker.unint_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f172,f233])).
fof(f409,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(forward_demodulation,[],[f399,f157])).
fof(f399,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(superposition,[],[f368,f160])).
fof(f410,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7)))) )),
  inference(forward_demodulation,[],[f400,f157])).
fof(f400,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X7))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7)))) )),
  inference(superposition,[],[f368,f161])).
fof(f431,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X19),s('type.realax.real',X18))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X19),s('type.realax.real',X20))) | s('type.realax.real',X20) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X17),s('type.realax.real',X18)))) )),
  inference(superposition,[],[f175,f161])).
fof(f516,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'const.marker.unint_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f511,f122])).
fof(f511,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,X5) = s(bool,'const.marker.unint_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f370,f233])).
fof(f534,plain,(
  ( ! [X28,X29,X27] : (s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X28)))) )),
  inference(superposition,[],[f231,f161])).
fof(f559,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X20),s(bool,X21))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f231,f161])).
fof(f873,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X30))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f640,f211])).
fof(f640,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f368,f231])).
fof(f910,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X19))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X20),s('type.realax.real',X19))) | s('type.realax.real',X20) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X17),s('type.realax.real',X18)))) )),
  inference(superposition,[],[f177,f161])).
fof(f1029,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK4(X16,X17,X18,bool))) )),
  inference(superposition,[],[f173,f130])).
fof(f1034,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK4(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK4(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f173,f130])).
fof(f1052,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X21) = s(bool,'const.marker.unint_1'(s(bool,X20))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f877,f157])).
fof(f1053,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X24) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f877,f160])).
fof(f1054,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') != s(bool,X26) | s(bool,X27) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f877,f161])).
fof(f1516,plain,(
  ( ! [X68,X66,X64,X62,X67,X65,X63] : (s(bool,X62) != s(bool,X64) | s(X63,'const.bool.COND_3'(s(bool,X62),s(X63,X67),s(X63,X68))) = s(X63,'const.bool.COND_3'(s(bool,X64),s(X63,X65),s(X63,X66))) | s(X63,X65) != s(X63,X67) | s(bool,'F_0') = s(bool,X62)) )),
  inference(forward_demodulation,[],[f1515,f157])).
fof(f1515,plain,(
  ( ! [X68,X66,X64,X62,X67,X65,X63] : (s(X63,'const.bool.COND_3'(s(bool,X62),s(X63,X67),s(X63,X68))) = s(X63,'const.bool.COND_3'(s(bool,X64),s(X63,X65),s(X63,X66))) | s(X63,X65) != s(X63,X67) | s(bool,X64) != s(bool,'const.marker.unint_1'(s(bool,X62))) | s(bool,'F_0') = s(bool,X62)) )),
  inference(forward_demodulation,[],[f1514,f157])).
fof(f1514,plain,(
  ( ! [X68,X66,X64,X62,X67,X65,X63] : (s(X63,'const.bool.COND_3'(s(bool,X64),s(X63,X65),s(X63,X66))) = s(X63,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X62))),s(X63,X67),s(X63,X68))) | s(X63,X65) != s(X63,X67) | s(bool,X64) != s(bool,'const.marker.unint_1'(s(bool,X62))) | s(bool,'F_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1401,f122])).
fof(f1401,plain,(
  ( ! [X68,X66,X64,X62,X67,X65,X63] : (~p(s(bool,'T_0')) | s(X63,'const.bool.COND_3'(s(bool,X64),s(X63,X65),s(X63,X66))) = s(X63,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X62))),s(X63,X67),s(X63,X68))) | s(X63,X65) != s(X63,X67) | s(bool,X64) != s(bool,'const.marker.unint_1'(s(bool,X62))) | s(bool,'F_0') = s(bool,X62)) )),
  inference(superposition,[],[f165,f233])).
fof(f1939,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) )),
  inference(forward_demodulation,[],[f1938,f161])).
fof(f1938,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) )),
  inference(forward_demodulation,[],[f1909,f161])).
fof(f1909,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0'))),s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) )),
  inference(unit_resulting_resolution,[],[f202,f168,f161,f166])).
fof(f1032,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK4(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK4(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f173,f231])).
fof(f1030,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK4(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK4(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f173,f231])).
fof(f2032,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(X33,X34) != s(X33,X35) | s(X33,'const.bool.COND_3'(s(bool,X32),s(X33,X38),s(X33,X35))) = s(X33,'const.bool.COND_3'(s(bool,X36),s(X33,X37),s(X33,X34))) | s(bool,X32) != s(bool,X36) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f1918,f202])).
fof(f1918,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (p(s(bool,'F_0')) | s(X33,X34) != s(X33,X35) | s(X33,'const.bool.COND_3'(s(bool,X32),s(X33,X38),s(X33,X35))) = s(X33,'const.bool.COND_3'(s(bool,X36),s(X33,X37),s(X33,X34))) | s(bool,X32) != s(bool,X36) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f166,f130])).
fof(f1507,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X5),s(X1,X6))) = s(X1,'const.bool.COND_3'(s(bool,X2),s(X1,X3),s(X1,X4))) | s(X1,X3) != s(X1,X5) | s(bool,X0) != s(bool,X2) | s(bool,X0) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1393,f122])).
fof(f1393,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X5),s(X1,X6))) = s(X1,'const.bool.COND_3'(s(bool,X2),s(X1,X3),s(X1,X4))) | s(X1,X3) != s(X1,X5) | s(bool,X0) != s(bool,X2) | s(bool,X0) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f165,f231])).
fof(f1418,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f1417,f160])).
fof(f1417,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real',X4)))) )),
  inference(forward_demodulation,[],[f1388,f161])).
fof(f1388,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'T_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real',X4)))) )),
  inference(unit_resulting_resolution,[],[f122,f168,f161,f165])).
fof(f1258,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X52),s('type.realax.real',X53))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X49),s('type.realax.real',X51))),s('type.realax.real',X50))),s('type.realax.real',X53))) | s('type.realax.real',X52) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X49),s('type.realax.real',X50))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X51),s('type.realax.real',X50)))))) )),
  inference(superposition,[],[f177,f174])).
fof(f1253,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X27),s('type.realax.real',X28))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X27),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X24),s('type.realax.real',X26))),s('type.realax.real',X25))))) | s('type.realax.real',X28) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X24),s('type.realax.real',X25))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X26),s('type.realax.real',X25)))))) )),
  inference(superposition,[],[f175,f174])).
fof(f1249,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real',X10))),s('type.realax.real',X9))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real',X11))) | s('type.realax.real',X11) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X10),s('type.realax.real',X9)))) )),
  inference(superposition,[],[f175,f174])).
fof(f1247,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(superposition,[],[f177,f174])).
fof(f1039,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f877,f130])).
fof(f1027,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK4(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f173,f231])).
fof(f1025,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK4(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f173,f231])).
fof(f969,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.realax.real',X4) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X2))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X5))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X5)))) )),
  inference(superposition,[],[f178,f168])).
fof(f960,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real',X2)))) )),
  inference(unit_resulting_resolution,[],[f168,f178])).
fof(f909,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X15))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X16),s('type.realax.real',X15))) | s('type.realax.real',X16) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X13),s('type.realax.real',X14)))) )),
  inference(superposition,[],[f177,f160])).
fof(f906,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X5),s('type.realax.real',X4))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X2))),s('type.realax.real',X4))) | s('type.realax.real',X5) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) )),
  inference(superposition,[],[f177,f168])).
fof(f885,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) = s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f658,f122])).
fof(f658,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f370,f231])).
fof(f878,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f654,f122])).
fof(f654,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f172,f231])).
fof(f877,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f650,f122])).
fof(f650,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f136,f231])).
fof(f629,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X2),s(X1,X3))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f160,f231])).
fof(f558,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X17) | s(bool,X19) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X17),s(bool,X18))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f231,f160])).
fof(f557,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.marker.unint_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f231,f157])).
fof(f547,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f231,f231])).
fof(f533,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X24),s(bool,X25)))) )),
  inference(superposition,[],[f231,f160])).
fof(f532,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'const.marker.unint_1'(s(bool,X22)))) )),
  inference(superposition,[],[f231,f157])).
fof(f486,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.realax.real',X4) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X2))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X5),s('type.realax.real',X4))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) )),
  inference(superposition,[],[f176,f168])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) )),
  inference(unit_resulting_resolution,[],[f168,f176])).
fof(f430,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X15),s('type.realax.real',X13))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X15),s('type.realax.real',X16))) | s('type.realax.real',X16) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X13),s('type.realax.real',X14)))) )),
  inference(superposition,[],[f175,f160])).
fof(f427,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X5))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X2))))) | s('type.realax.real',X5) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) )),
  inference(superposition,[],[f175,f168])).
fof(f308,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f161,f130])).
fof(f263,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f258,f202])).
fof(f258,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f169,f130])).
fof(f231,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f130,f130])).
fof(f211,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f202,f132])).
fof(f177,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))) | s('type.realax.real',X0) = s('type.realax.real',X1)) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.realax.real',X0) != s('type.realax.real',X1) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))) <=> s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X7,X6,X22] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X22))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X22))) <=> s('type.realax.real',X6) = s('type.realax.real',X7))),
  file('thm.real.add_rat','thm.real.REAL_EQ_RADD')).
fof(f178,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f111])).
fof(f175,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | s('type.realax.real',X1) = s('type.realax.real',X2)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | s('type.realax.real',X1) = s('type.realax.real',X2)) & (s('type.realax.real',X1) != s('type.realax.real',X2) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) <=> s('type.realax.real',X1) = s('type.realax.real',X2))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X7,X6,X22] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X6))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X22))) <=> s('type.realax.real',X6) = s('type.realax.real',X22))),
  file('thm.real.add_rat','thm.real.REAL_EQ_LADD')).
fof(f176,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X1) != s('type.realax.real',X2) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f174,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X7,X6,X22] : s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X22))),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X7))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X22),s('type.realax.real',X7)))))),
  file('thm.real.add_rat','thm.real.REAL_DIV_ADD')).
fof(f173,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.real.add_rat',aHL_EXT)).
fof(f169,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.real.add_rat','thm.bool.IMP_ANTISYM_AX')).
fof(f172,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f82])).
fof(f168,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X7,X6] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real',X6)))),
  file('thm.real.add_rat','thm.realax.REAL_MUL_SYM')).
fof(f165,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,X2)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(X0,X3) != s(X0,X4) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X16] : ! [X18,X19,X7,X20,X6,X21] : ((s(bool,X18) = s(bool,X19) & (p(s(bool,X19)) => s(X16,X7) = s(X16,X20)) & (~p(s(bool,X19)) => s(X16,X6) = s(X16,X21))) => s(X16,'const.bool.COND_3'(s(bool,X18),s(X16,X7),s(X16,X6))) = s(X16,'const.bool.COND_3'(s(bool,X19),s(X16,X20),s(X16,X21))))),
  file('thm.real.add_rat','thm.bool.COND_CONG')).
fof(f166,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X2)) | s(X0,X5) != s(X0,X6) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f80])).
fof(f167,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,X2) | s(X0,X5) != s(X0,X6) | s(X0,X3) != s(X0,X4) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f80])).
fof(f160,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X16] : ! [X13,X14] : s(X16,X13) = s(X16,'const.bool.COND_3'(s(bool,'T_0'),s(X16,X13),s(X16,X14)))),
  file('thm.real.add_rat','thm.bool.COND_CLAUSES|split|0')).
fof(f157,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.marker.unint_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.marker.unint_1'(s(X0,X1)))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.marker.unint_1'(s(X0,X1)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X16] : ! [X7] : s(X16,X7) = s(X16,'const.marker.unint_1'(s(X16,X7)))),
  file('thm.real.add_rat','thm.marker.unint_def')).
fof(f132,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.real.add_rat','thm.bool.EQ_CLAUSES|split|1')).
fof(f133,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f92])).
fof(f130,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.real.add_rat',aHL_BOOL_CASES)).
fof(f122,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real.add_rat','thm.bool.TRUTH')).
fof(f10432,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl6 | $spl14)),
  inference(global_subsumption,[],[f10431,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706,f5405,f3707,f3709,f3724,f5470,f185,f5545,f5544,f5758,f5694,f5703,f5705,f5707,f5710,f5711,f5712,f5714,f5717,f5718,f5719,f5722,f5723,f5729,f5732,f5734,f5737,f190,f6800,f6802,f6804,f6806,f6808,f6810,f6812,f6814,f6816,f6818,f6822,f6824,f6826,f6828,f6832,f6834,f6836,f6838,f193,f10122,f10139,f10140,f10211,f10213,f10220,f10233,f10234,f10235,f10278,f10280,f10282,f10291,f10292,f10293,f10294,f10295,f10296,f10297,f10299,f10301,f10355,f10307,f10309,f10356,f10311,f10312,f10318,f10321,f10323,f10326,f188,f10383,f10385,f10387,f10391,f10393,f10395,f10397,f10401,f10403,f10405,f10407,f10409,f10411,f10413,f10415,f10417,f10419,f10421,f10422])).
fof(f10431,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl0 | $spl6 | $spl14)),
  inference(subsumption_resolution,[],[f10430,f10379])).
fof(f10379,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f10378,f174])).
fof(f10378,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f10377,f960])).
fof(f10377,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f8059,f188])).
fof(f8059,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl6),
  inference(forward_demodulation,[],[f8058,f157])).
fof(f8058,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl6),
  inference(forward_demodulation,[],[f8057,f157])).
fof(f8057,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl6),
  inference(forward_demodulation,[],[f194,f477])).
fof(f10430,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl14)),
  inference(forward_demodulation,[],[f10429,f174])).
fof(f10429,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl14)),
  inference(forward_demodulation,[],[f10428,f157])).
fof(f10428,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl14)),
  inference(forward_demodulation,[],[f10427,f174])).
fof(f10427,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl14)),
  inference(forward_demodulation,[],[f10426,f157])).
fof(f10426,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl14)),
  inference(forward_demodulation,[],[f10425,f960])).
fof(f10425,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl14)),
  inference(forward_demodulation,[],[f10424,f477])).
fof(f10424,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl14)),
  inference(forward_demodulation,[],[f198,f188])).
fof(f10423,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706,f5405,f3707,f3709,f3724,f5470,f185,f5545,f5544,f5758,f5694,f5703,f5705,f5707,f5710,f5711,f5712,f5714,f5717,f5718,f5719,f5722,f5723,f5729,f5732,f5734,f5737,f190,f6800,f6802,f6804,f6806,f6808,f6810,f6812,f6814,f6816,f6818,f6822,f6824,f6826,f6828,f6832,f6834,f6836,f6838,f193,f10122,f10139,f10140,f10211,f10213,f10220,f10233,f10234,f10235,f10278,f10280,f10282,f10291,f10292,f10293,f10294,f10295,f10296,f10297,f10299,f10301,f10355,f10307,f10309,f10356,f10311,f10312,f10318,f10321,f10323,f10326,f188,f10383,f10385,f10387,f10391,f10393,f10395,f10397,f10401,f10403,f10405,f10407,f10409,f10411,f10413,f10415,f10417,f10419,f10421,f10422])).
fof(f10376,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f10357,f160])).
fof(f10357,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ($spl2 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f122,f122,f122,f8116])).
fof(f8116,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X6)) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0)))))))))))) ) | ($spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f8115,f174])).
fof(f8115,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | ($spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f8084,f960])).
fof(f8084,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | ($spl2 | $spl4 | $spl6)),
  inference(backward_demodulation,[],[f8060,f8059])).
fof(f8060,plain,(
  s('type.realax.real',sK0) = s('type.realax.real',sK2) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f192,f190])).
fof(f10361,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f122,f122,f122,f160,f8116])).
fof(f8109,plain,(
  $false | ($spl2 | $spl4 | $spl18)),
  inference(subsumption_resolution,[],[f8080,f174])).
fof(f8080,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) | ($spl2 | $spl4 | $spl18)),
  inference(backward_demodulation,[],[f8060,f7965])).
fof(f7965,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) | $spl18),
  inference(forward_demodulation,[],[f7964,f161])).
fof(f7964,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))))) | $spl18),
  inference(forward_demodulation,[],[f7963,f160])).
fof(f7963,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))) | $spl18),
  inference(forward_demodulation,[],[f7949,f161])).
fof(f7949,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f202,f202,f122,f6791])).
fof(f6791,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X5)) | p(s(bool,X6))) ) | $spl18),
  inference(forward_demodulation,[],[f6790,f157])).
fof(f6790,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | $spl18),
  inference(forward_demodulation,[],[f5763,f157])).
fof(f5763,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | $spl18),
  inference(forward_demodulation,[],[f200,f477])).
fof(f1511,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl24),
  inference(cnf_transformation,[],[f1511_D])).
fof(f1511_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f8047,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2)))))))) ) | $spl36),
  inference(cnf_transformation,[],[f8047_D])).
fof(f8047_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2)))))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f8051,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl16),
  inference(condensation,[],[f8050])).
fof(f8050,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(forward_demodulation,[],[f8049,f160])).
fof(f8049,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X2))) ) | $spl16),
  inference(condensation,[],[f8048])).
fof(f8048,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X25),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X24),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X25)) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) ) | $spl16),
  inference(subsumption_resolution,[],[f8034,f263])).
fof(f8034,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,X23)) | p(s(bool,X24)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X25),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X24),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X25)) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) ) | $spl16),
  inference(superposition,[],[f6841,f231])).
fof(f6841,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X5)) | p(s(bool,X4)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X6))) ) | $spl16),
  inference(forward_demodulation,[],[f6840,f157])).
fof(f6840,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) | $spl16),
  inference(forward_demodulation,[],[f6839,f157])).
fof(f6839,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) | $spl16),
  inference(forward_demodulation,[],[f199,f477])).
fof(f8046,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl16),
  inference(condensation,[],[f8045])).
fof(f8045,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(forward_demodulation,[],[f8044,f160])).
fof(f8044,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X2))) ) | $spl16),
  inference(condensation,[],[f8043])).
fof(f8043,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X21),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X20),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X21)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | $spl16),
  inference(subsumption_resolution,[],[f8033,f263])).
fof(f8033,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,X19)) | p(s(bool,X20)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X21),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X20),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X21)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | $spl16),
  inference(superposition,[],[f6841,f231])).
fof(f8041,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f8040,f161])).
fof(f8040,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))) | $spl16),
  inference(forward_demodulation,[],[f8026,f160])).
fof(f8026,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | $spl16),
  inference(unit_resulting_resolution,[],[f202,f122,f202,f6841])).
fof(f7970,plain,(
  ( ! [X2,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))))))) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl34),
  inference(cnf_transformation,[],[f7970_D])).
fof(f7970_D,plain,(
  ( ! [X2,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))))))) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f7973,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))))))) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl18),
  inference(forward_demodulation,[],[f7972,f160])).
fof(f7972,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl18),
  inference(condensation,[],[f7971])).
fof(f7971,plain,(
  ( ! [X26,X24,X23,X25] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X24),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X25),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X25)) | p(s(bool,X24)) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) ) | $spl18),
  inference(subsumption_resolution,[],[f7957,f263])).
fof(f7957,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,X23)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X24),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X25),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X25)) | p(s(bool,X24)) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) ) | $spl18),
  inference(superposition,[],[f6791,f231])).
fof(f7969,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))))))) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl18),
  inference(forward_demodulation,[],[f7968,f160])).
fof(f7968,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl18),
  inference(condensation,[],[f7967])).
fof(f7967,plain,(
  ( ! [X21,X19,X22,X20] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X20),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X21),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X21)) | p(s(bool,X20)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | $spl18),
  inference(subsumption_resolution,[],[f7956,f263])).
fof(f7956,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,X19)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X20),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X21),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X21)) | p(s(bool,X20)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | $spl18),
  inference(superposition,[],[f6791,f231])).
fof(f7916,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f7901,f7902])).
fof(f7902,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f202,f160,f6842])).
fof(f6842,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X5)) | p(s(bool,X4)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2)))))))))))) ) | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f6841,f6795])).
fof(f6795,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl8),
  inference(forward_demodulation,[],[f6794,f157])).
fof(f6794,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl8),
  inference(forward_demodulation,[],[f6793,f157])).
fof(f6793,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl8),
  inference(forward_demodulation,[],[f195,f477])).
fof(f7901,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f122,f160,f6842])).
fof(f7903,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f122,f202,f160,f6842])).
fof(f7853,plain,(
  $false | ($spl10 | $spl18)),
  inference(subsumption_resolution,[],[f7838,f7839])).
fof(f7839,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl10 | $spl18)),
  inference(unit_resulting_resolution,[],[f122,f160,f6792])).
fof(f6792,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X5))) ) | ($spl10 | $spl18)),
  inference(subsumption_resolution,[],[f6791,f6789])).
fof(f6789,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl10),
  inference(forward_demodulation,[],[f6788,f157])).
fof(f6788,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl10),
  inference(forward_demodulation,[],[f6787,f157])).
fof(f6787,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl10),
  inference(forward_demodulation,[],[f196,f477])).
fof(f7838,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl10 | $spl18)),
  inference(unit_resulting_resolution,[],[f202,f160,f6792])).
fof(f7840,plain,(
  $false | ($spl10 | $spl18)),
  inference(unit_resulting_resolution,[],[f202,f122,f160,f6792])).
fof(f6798,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f207,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1103,f1112,f1115,f1117,f1119,f1122,f1124,f1126,f1128,f1132,f1135,f1137,f1140,f1229,f1246,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1127,f1120,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706,f5405,f3707,f3709,f3724,f5470,f185,f5537,f5536,f5551,f5550,f5563,f5562,f5694,f5703,f5705,f5707,f5710,f5711,f5712,f5714,f5717,f5718,f5719,f5722,f5723,f5729,f5732,f5734,f5737,f192,f6765,f6766,f6767,f6770,f6771,f6772,f6773,f6776,f6777,f6778,f6779,f6780,f6781,f6782,f6783,f6784,f6785,f6786,f6796,f189])).
fof(f6796,plain,(
  s('type.realax.real',sK0) = s('type.realax.real',sK2) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f190,f192])).
fof(f6786,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s('type.realax.real',sK0) = s('type.realax.real',X91) | s('type.realax.real',sK0) = s('type.realax.real',X90) | s('type.realax.real',X93) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X89),s('type.realax.real',X91))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X92),s('type.realax.real',X90))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X93),s('type.realax.real',X94))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X89),s('type.realax.real',X90))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X92))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X90))))),s('type.realax.real',X94)))) ) | $spl4),
  inference(forward_demodulation,[],[f6764,f192])).
fof(f6764,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s('type.realax.real',sK0) = s('type.realax.real',X90) | s('type.realax.real',X93) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X89),s('type.realax.real',X91))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X92),s('type.realax.real',X90))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X93),s('type.realax.real',X94))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X89),s('type.realax.real',X90))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X92))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X91),s('type.realax.real',X90))))),s('type.realax.real',X94))) | s('type.realax.real',X91) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5737])).
fof(f6785,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s('type.realax.real',sK0) = s('type.realax.real',X73) | s('type.realax.real',sK0) = s('type.realax.real',X72) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X75),s('type.realax.real',X76))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X71),s('type.realax.real',X73))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X74),s('type.realax.real',X72))))),s('type.realax.real',X76))) | s('type.realax.real',X75) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X71),s('type.realax.real',X72))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X74))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X72)))))) ) | $spl4),
  inference(forward_demodulation,[],[f6763,f192])).
fof(f6763,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s('type.realax.real',sK0) = s('type.realax.real',X72) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X75),s('type.realax.real',X76))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X71),s('type.realax.real',X73))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X74),s('type.realax.real',X72))))),s('type.realax.real',X76))) | s('type.realax.real',X75) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X71),s('type.realax.real',X72))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X74))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X73),s('type.realax.real',X72))))) | s('type.realax.real',X73) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5734])).
fof(f6784,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s('type.realax.real',sK0) = s('type.realax.real',X61) | s('type.realax.real',sK0) = s('type.realax.real',X60) | s('type.realax.real',X63) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X59),s('type.realax.real',X61))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X62),s('type.realax.real',X60))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real',X63))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X59),s('type.realax.real',X60))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X62))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X60)))))))) ) | $spl4),
  inference(forward_demodulation,[],[f6762,f192])).
fof(f6762,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s('type.realax.real',sK0) = s('type.realax.real',X60) | s('type.realax.real',X63) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X59),s('type.realax.real',X61))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X62),s('type.realax.real',X60))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real',X63))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X59),s('type.realax.real',X60))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X62))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X61),s('type.realax.real',X60))))))) | s('type.realax.real',X61) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5732])).
fof(f6783,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.realax.real',sK0) = s('type.realax.real',X43) | s('type.realax.real',sK0) = s('type.realax.real',X42) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real',X46))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X41),s('type.realax.real',X43))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X44),s('type.realax.real',X42))))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X41),s('type.realax.real',X42))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X44))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X42)))))) ) | $spl4),
  inference(forward_demodulation,[],[f6761,f192])).
fof(f6761,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.realax.real',sK0) = s('type.realax.real',X42) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real',X46))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X41),s('type.realax.real',X43))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X44),s('type.realax.real',X42))))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X41),s('type.realax.real',X42))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X44))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X42))))) | s('type.realax.real',X43) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5729])).
fof(f6782,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.realax.real',sK0) = s('type.realax.real',X9) | s('type.realax.real',sK0) = s('type.realax.real',X10) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X11))))),s('type.realax.real',X12))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X10))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X11),s('type.realax.real',X9))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X12),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9)))))))) ) | $spl4),
  inference(forward_demodulation,[],[f6760,f192])).
fof(f6760,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.realax.real',sK0) = s('type.realax.real',X10) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X11))))),s('type.realax.real',X12))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X10))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X11),s('type.realax.real',X9))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X12),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X9))))))) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5723])).
fof(f6781,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.realax.real',sK0) = s('type.realax.real',X5) | s('type.realax.real',sK0) = s('type.realax.real',X4) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X4)))))))) ) | $spl4),
  inference(forward_demodulation,[],[f6759,f192])).
fof(f6759,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.realax.real',sK0) = s('type.realax.real',X4) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X4))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5722])).
fof(f6780,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))) ) | $spl4),
  inference(forward_demodulation,[],[f6758,f192])).
fof(f6758,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5719])).
fof(f6779,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real',sK0) = s('type.realax.real',X35) | s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X35))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X36),s('type.realax.real',X33))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X36))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X33)))))) ) | $spl4),
  inference(forward_demodulation,[],[f6757,f192])).
fof(f6757,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X35))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X36),s('type.realax.real',X33))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X36))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X35),s('type.realax.real',X33))))) | s('type.realax.real',X35) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5718])).
fof(f6778,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real',sK0) = s('type.realax.real',X30) | s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X30))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X27))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X31))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X27)))))) ) | $spl4),
  inference(forward_demodulation,[],[f6756,f192])).
fof(f6756,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X30))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X27))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X31))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X27))))) | s('type.realax.real',X30) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5717])).
fof(f6777,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real',sK0) = s('type.realax.real',X16) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real',X16))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X17),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real',X17))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13)))))))) ) | $spl4),
  inference(forward_demodulation,[],[f6755,f192])).
fof(f6755,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real',X16))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X17),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real',X17))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))))) | s('type.realax.real',X16) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5714])).
fof(f6776,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X5) | s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1)))))))))) ) | $spl4),
  inference(forward_demodulation,[],[f6775,f192])).
fof(f6775,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X5) | s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(forward_demodulation,[],[f6774,f192])).
fof(f6774,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X5) | s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(forward_demodulation,[],[f6754,f192])).
fof(f6754,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5712])).
fof(f6773,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real',sK0) = s('type.realax.real',X36) | s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X35),s('type.realax.real',X36))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X35))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X36)))))) ) | $spl4),
  inference(forward_demodulation,[],[f6753,f192])).
fof(f6753,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X32),s('type.realax.real',X33))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X33))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X35),s('type.realax.real',X36))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X35))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X36))))) | s('type.realax.real',X36) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5711])).
fof(f6772,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real',sK0) = s('type.realax.real',X31) | s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X30),s('type.realax.real',X31))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X31))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X30))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X31)))))) ) | $spl4),
  inference(forward_demodulation,[],[f6752,f192])).
fof(f6752,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X27),s('type.realax.real',X28))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X27))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X30),s('type.realax.real',X31))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X31))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X30))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X31))))) | s('type.realax.real',X31) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5710])).
fof(f6771,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real',sK0) = s('type.realax.real',X17) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X17))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X17))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X16))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X17)))))) ) | $spl4),
  inference(forward_demodulation,[],[f6751,f192])).
fof(f6751,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X17))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X17))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X16))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X13))),s('type.realax.real',X17))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5707])).
fof(f6770,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X6) | s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6)))))) ) | $spl4),
  inference(forward_demodulation,[],[f6769,f192])).
fof(f6769,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X6) | s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(forward_demodulation,[],[f6768,f192])).
fof(f6768,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X6) | s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(forward_demodulation,[],[f6750,f192])).
fof(f6750,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X5))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real',X6))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5705])).
fof(f6767,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X3) | s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3)))))) ) | $spl4),
  inference(forward_demodulation,[],[f6749,f192])).
fof(f6749,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5703])).
fof(f6766,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) ) | $spl4),
  inference(forward_demodulation,[],[f6748,f192])).
fof(f6748,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f5694])).
fof(f6765,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X3) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3)))))) ) | $spl4),
  inference(forward_demodulation,[],[f6747,f192])).
fof(f6747,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X3) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f185])).
fof(f5562,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK0))),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK2))),s('type.realax.real',X1)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f5537,f177])).
fof(f5563,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK0))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK2)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f5537,f175])).
fof(f5550,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X0))),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X0))),s('type.realax.real',X1)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f5536,f177])).
fof(f5551,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X1))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X1)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f5536,f175])).
fof(f5536,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X0))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X0)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f189,f177])).
fof(f5537,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK0))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK2)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f189,f175])).
fof(f1120,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X30))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X32),s('type.realax.real',X33))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X31),s('type.realax.real',X33))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X32))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X33))))) | s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X29),s('type.realax.real',X30))) | s('type.realax.real',sK0) = s('type.realax.real',X33)) ) | $spl4),
  inference(superposition,[],[f207,f161])).
fof(f1127,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X32))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X33),s('type.realax.real',X30))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X31),s('type.realax.real',X30))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X30))))) | s('type.realax.real',sK0) = s('type.realax.real',X32) | s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X29),s('type.realax.real',X30)))) ) | $spl4),
  inference(superposition,[],[f207,f161])).
fof(f1246,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))))) | s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real',sK0) = s('type.realax.real',X1)) ) | $spl4),
  inference(superposition,[],[f174,f207])).
fof(f1229,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))))) | s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real',sK0) = s('type.realax.real',X1)) ) | $spl4),
  inference(superposition,[],[f174,f207])).
fof(f1140,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (s('type.realax.real',X62) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X58),s('type.realax.real',X60))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X61),s('type.realax.real',X59))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X62),s('type.realax.real',X63))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X58),s('type.realax.real',X59))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X60),s('type.realax.real',X61))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X60),s('type.realax.real',X59))))),s('type.realax.real',X63))) | s('type.realax.real',sK0) = s('type.realax.real',X60) | s('type.realax.real',sK0) = s('type.realax.real',X59)) ) | $spl4),
  inference(superposition,[],[f178,f207])).
fof(f1137,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X44),s('type.realax.real',X45))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X40),s('type.realax.real',X42))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X43),s('type.realax.real',X41))))),s('type.realax.real',X45))) | s('type.realax.real',X44) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X40),s('type.realax.real',X41))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X42),s('type.realax.real',X43))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X42),s('type.realax.real',X41))))) | s('type.realax.real',sK0) = s('type.realax.real',X42) | s('type.realax.real',sK0) = s('type.realax.real',X41)) ) | $spl4),
  inference(superposition,[],[f177,f207])).
fof(f1135,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s('type.realax.real',X32) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X28),s('type.realax.real',X30))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X29))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X33),s('type.realax.real',X32))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X33),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X28),s('type.realax.real',X29))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X31))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X29))))))) | s('type.realax.real',sK0) = s('type.realax.real',X30) | s('type.realax.real',sK0) = s('type.realax.real',X29)) ) | $spl4),
  inference(superposition,[],[f176,f207])).
fof(f1132,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X14),s('type.realax.real',X15))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X14),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X10),s('type.realax.real',X12))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X13),s('type.realax.real',X11))))))) | s('type.realax.real',X15) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X11))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X12),s('type.realax.real',X13))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X12),s('type.realax.real',X11))))) | s('type.realax.real',sK0) = s('type.realax.real',X12) | s('type.realax.real',sK0) = s('type.realax.real',X11)) ) | $spl4),
  inference(superposition,[],[f175,f207])).
fof(f1128,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) | s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real',sK0) = s('type.realax.real',X1)) ) | $spl4),
  inference(superposition,[],[f207,f168])).
fof(f1126,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X26),s('type.realax.real',X27))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X28),s('type.realax.real',X24))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X26),s('type.realax.real',X24))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X28))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X24))))) | s('type.realax.real',sK0) = s('type.realax.real',X27) | s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X24),s('type.realax.real',X25)))) ) | $spl4),
  inference(superposition,[],[f207,f160])).
fof(f1124,plain,(
  ( ! [X14,X19,X17,X15,X13,X18,X16] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X17),s('type.realax.real',X18))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X19),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X13),s('type.realax.real',X15))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X14))))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X17),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X13),s('type.realax.real',X15))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X14))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X18),s('type.realax.real',X19))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X18),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X13),s('type.realax.real',X15))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X14))))))))) | s('type.realax.real',sK0) = s('type.realax.real',X18) | s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X16))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X14))))) | s('type.realax.real',sK0) = s('type.realax.real',X15) | s('type.realax.real',sK0) = s('type.realax.real',X14)) ) | $spl4),
  inference(superposition,[],[f207,f207])).
fof(f1122,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X6))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X3))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X3))))))) | s('type.realax.real',sK0) = s('type.realax.real',X6) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) ) | $spl4),
  inference(superposition,[],[f207,f168])).
fof(f1119,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X26),s('type.realax.real',X24))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X27),s('type.realax.real',X28))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X26),s('type.realax.real',X28))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X24),s('type.realax.real',X27))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X24),s('type.realax.real',X28))))) | s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X24),s('type.realax.real',X25))) | s('type.realax.real',sK0) = s('type.realax.real',X28)) ) | $spl4),
  inference(superposition,[],[f207,f160])).
fof(f1117,plain,(
  ( ! [X14,X19,X17,X15,X13,X18,X16] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X17),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X13),s('type.realax.real',X15))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X14))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X18),s('type.realax.real',X19))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X17),s('type.realax.real',X19))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X13),s('type.realax.real',X15))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X14))))),s('type.realax.real',X18))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X13),s('type.realax.real',X15))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X16),s('type.realax.real',X14))))),s('type.realax.real',X19))))) | s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X16))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X14))))) | s('type.realax.real',sK0) = s('type.realax.real',X19) | s('type.realax.real',sK0) = s('type.realax.real',X15) | s('type.realax.real',sK0) = s('type.realax.real',X14)) ) | $spl4),
  inference(superposition,[],[f207,f207])).
fof(f1115,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X3))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X7))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X7))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X7))))) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))) | s('type.realax.real',sK0) = s('type.realax.real',X7)) ) | $spl4),
  inference(superposition,[],[f207,f168])).
fof(f1112,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))))) | s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real',sK0) = s('type.realax.real',X3)) ) | $spl4),
  inference(superposition,[],[f207,f168])).
fof(f1103,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X1))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real',sK0) = s('type.realax.real',X1)) ) | $spl4),
  inference(superposition,[],[f207,f168])).
fof(f207,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) | s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real',sK0) = s('type.realax.real',X3)) ) | $spl4),
  inference(forward_demodulation,[],[f206,f192])).
fof(f206,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X3))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f192,f185])).
fof(f6797,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f207,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1103,f1112,f1115,f1117,f1119,f1122,f1124,f1126,f1128,f1132,f1135,f1137,f1140,f1229,f1246,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1127,f1120,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706,f5405,f3707,f3709,f3724,f5470,f189,f185,f5537,f5536,f5551,f5550,f5563,f5562,f5694,f5703,f5705,f5707,f5710,f5711,f5712,f5714,f5717,f5718,f5719,f5722,f5723,f5729,f5732,f5734,f5737,f192,f6765,f6766,f6767,f6770,f6771,f6772,f6773,f6776,f6777,f6778,f6779,f6780,f6781,f6782,f6783,f6784,f6785,f6786,f6796])).
fof(f6746,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f207,f208,f211,f231,f263,f308,f425,f427,f430,f450,f458,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f904,f903,f902,f906,f909,f931,f930,f941,f940,f960,f969,f994,f993,f1012,f1011,f1025,f1027,f1039,f1076,f1075,f1160,f1103,f1112,f1115,f1117,f1119,f1122,f1124,f1126,f1128,f1132,f1135,f1137,f1140,f1168,f1167,f1195,f1194,f1229,f1246,f1247,f1249,f1253,f1258,f1302,f1301,f1332,f1331,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1127,f1120,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706,f5405,f3707,f3709,f3724,f5470,f191,f185,f5541,f5540,f5664,f5663,f5676,f5675,f5748,f5694,f5703,f5705,f5707,f5710,f5711,f5712,f5714,f5717,f5718,f5719,f5722,f5723,f5729,f5732,f5734,f5737,f188,f5770,f5771,f5772,f5773,f5775,f5769,f5768,f5767,f192])).
fof(f5767,plain,(
  s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f188,f191])).
fof(f5768,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X0))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) ) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f188,f5540])).
fof(f5769,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK0))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f188,f5541])).
fof(f5775,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',sK0))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0)))))) ) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f188,f5748])).
fof(f5773,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK0))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f188,f5676])).
fof(f5772,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK0))),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f188,f5675])).
fof(f5771,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X1))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))))) ) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f188,f5664])).
fof(f5770,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X0))),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.realax.real',X1)))) ) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f188,f5663])).
fof(f5748,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',sK2))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',sK2)))))) ) | $spl3),
  inference(forward_demodulation,[],[f5686,f174])).
fof(f5686,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',sK2))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',sK2)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f191,f191,f185])).
fof(f5675,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK2))),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5541,f177])).
fof(f5676,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK2))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5541,f175])).
fof(f5663,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X0))),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.realax.real',X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5540,f177])).
fof(f5664,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X1))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5540,f175])).
fof(f5540,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X0))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f191,f177])).
fof(f5541,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK2))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f191,f175])).
fof(f1331,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real',X2))),s('type.realax.real',X3))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real',X2))),s('type.realax.real',X3)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f993,f177])).
fof(f1332,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X1))),s('type.realax.real',X2))),s('type.realax.real',X3))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X1))),s('type.realax.real',X2))),s('type.realax.real',X3)))))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f993,f175])).
fof(f1301,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK0))),s('type.realax.real',X2))))),s('type.realax.real',X3))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK2))),s('type.realax.real',X2))))),s('type.realax.real',X3)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f941,f177])).
fof(f1302,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',sK0))),s('type.realax.real',X3))))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',sK2))),s('type.realax.real',X3)))))))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f941,f175])).
fof(f1194,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK0))),s('type.realax.real',X1))),s('type.realax.real',X2))),s('type.realax.real',X3))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK2))),s('type.realax.real',X1))),s('type.realax.real',X2))),s('type.realax.real',X3)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f940,f177])).
fof(f1195,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK0))),s('type.realax.real',X2))),s('type.realax.real',X3))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK2))),s('type.realax.real',X2))),s('type.realax.real',X3)))))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f940,f175])).
fof(f1167,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK0))))),s('type.realax.real',X2))),s('type.realax.real',X3))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK2))))),s('type.realax.real',X2))),s('type.realax.real',X3)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f904,f177])).
fof(f1168,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',sK0))))),s('type.realax.real',X3))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',sK2))))),s('type.realax.real',X3)))))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f904,f175])).
fof(f1160,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',sK2))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',sK2)))))) ) | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f1095,f174])).
fof(f1095,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',sK2))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',sK2)))))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f208,f208,f207])).
fof(f1075,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',sK0))))))),s('type.realax.real',X3))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',sK2))))))),s('type.realax.real',X3)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f458,f177])).
fof(f1076,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',sK0))))))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',sK2)))))))))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f458,f175])).
fof(f1011,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X1))))),s('type.realax.real',X2))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X1))))),s('type.realax.real',X2)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f931,f177])).
fof(f1012,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X2))))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X2)))))))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f931,f175])).
fof(f993,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real',X2))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real',X2)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f930,f177])).
fof(f994,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X1))),s('type.realax.real',X2))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X1))),s('type.realax.real',X2)))))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f930,f175])).
fof(f940,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK0))),s('type.realax.real',X1))),s('type.realax.real',X2))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK2))),s('type.realax.real',X1))),s('type.realax.real',X2)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f903,f177])).
fof(f941,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK0))),s('type.realax.real',X2))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK2))),s('type.realax.real',X2)))))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f903,f175])).
fof(f930,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X0))),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X0))),s('type.realax.real',X1)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f902,f177])).
fof(f931,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X1))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X1)))))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f902,f175])).
fof(f902,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK0),s('type.realax.real',X0))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK2),s('type.realax.real',X0)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f208,f177])).
fof(f903,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK0))),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK2))),s('type.realax.real',X1)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f425,f177])).
fof(f904,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK0))))),s('type.realax.real',X2))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK2))))),s('type.realax.real',X2)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f450,f177])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',sK0))))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',sK2)))))))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f450,f175])).
fof(f450,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK0))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',sK2)))))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f425,f175])).
fof(f425,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK0))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK2)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f208,f175])).
fof(f208,plain,(
  s('type.realax.real',sK0) != s('type.realax.real',sK2) | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f191,f192])).
fof(f6745,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl10)),
  inference(global_subsumption,[],[f6744,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f192,f207,f208,f211,f231,f263,f308,f425,f427,f430,f450,f458,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f904,f903,f902,f906,f909,f931,f930,f941,f940,f960,f969,f994,f993,f1012,f1011,f1025,f1027,f1039,f1076,f1075,f1160,f1103,f1112,f1115,f1117,f1119,f1122,f1124,f1126,f1128,f1132,f1135,f1137,f1140,f1168,f1167,f1195,f1194,f1229,f1246,f1247,f1249,f1253,f1258,f1302,f1301,f1332,f1331,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1127,f1120,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706,f5405,f3707,f3709,f3724,f5470,f191,f185,f5541,f5540,f5664,f5663,f5676,f5675,f5748,f5694,f5703,f5705,f5707,f5710,f5711,f5712,f5714,f5717,f5718,f5719,f5722,f5723,f5729,f5732,f5734,f5737,f188,f5770,f5771,f5772,f5773,f5775,f5769,f5768,f5767])).
fof(f6744,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f6743,f174])).
fof(f6743,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f6742,f157])).
fof(f6742,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f6741,f174])).
fof(f6741,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f6740,f157])).
fof(f6740,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f6739,f960])).
fof(f6739,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f6738,f477])).
fof(f6738,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f196,f188])).
fof(f6737,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl18)),
  inference(global_subsumption,[],[f6736,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f192,f207,f208,f211,f231,f263,f308,f425,f427,f430,f450,f458,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f904,f903,f902,f906,f909,f931,f930,f941,f940,f960,f969,f994,f993,f1012,f1011,f1025,f1027,f1039,f1076,f1075,f1160,f1103,f1112,f1115,f1117,f1119,f1122,f1124,f1126,f1128,f1132,f1135,f1137,f1140,f1168,f1167,f1195,f1194,f1229,f1246,f1247,f1249,f1253,f1258,f1302,f1301,f1332,f1331,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1127,f1120,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706,f5405,f3707,f3709,f3724,f5470,f191,f185,f5541,f5540,f5664,f5663,f5676,f5675,f5748,f5694,f5703,f5705,f5707,f5710,f5711,f5712,f5714,f5717,f5718,f5719,f5722,f5723,f5729,f5732,f5734,f5737,f188,f5770,f5771,f5772,f5773,f5775,f5769,f5768,f5767])).
fof(f6736,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl18)),
  inference(forward_demodulation,[],[f6735,f174])).
fof(f6735,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl18)),
  inference(forward_demodulation,[],[f6734,f157])).
fof(f6734,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl18)),
  inference(forward_demodulation,[],[f6733,f174])).
fof(f6733,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl18)),
  inference(forward_demodulation,[],[f6732,f157])).
fof(f6732,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl18)),
  inference(forward_demodulation,[],[f6731,f960])).
fof(f6731,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl18)),
  inference(forward_demodulation,[],[f5763,f188])).
fof(f6730,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f192,f207,f208,f211,f231,f263,f308,f425,f427,f430,f450,f458,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f904,f903,f902,f906,f909,f931,f930,f941,f940,f960,f969,f994,f993,f1012,f1011,f1025,f1027,f1039,f1076,f1075,f1160,f1103,f1112,f1115,f1117,f1119,f1122,f1124,f1126,f1128,f1132,f1135,f1137,f1140,f1168,f1167,f1195,f1194,f1229,f1246,f1247,f1249,f1253,f1258,f1302,f1301,f1332,f1331,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1127,f1120,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706,f5405,f3707,f3709,f3724,f5470,f191,f185,f5541,f5540,f5664,f5663,f5676,f5675,f5748,f5694,f5703,f5705,f5707,f5710,f5711,f5712,f5714,f5717,f5718,f5719,f5722,f5723,f5729,f5732,f5734,f5737,f188,f5770,f5771,f5772,f5773,f5775,f5769,f5768,f5767])).
fof(f6729,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f192,f207,f208,f211,f231,f263,f308,f425,f427,f430,f450,f458,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f904,f903,f902,f906,f909,f931,f930,f941,f940,f960,f969,f994,f993,f1012,f1011,f1025,f1027,f1039,f1076,f1075,f1160,f1103,f1112,f1115,f1117,f1119,f1122,f1124,f1126,f1128,f1132,f1135,f1137,f1140,f1168,f1167,f1195,f1194,f1229,f1246,f1247,f1249,f1253,f1258,f1302,f1301,f1332,f1331,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1127,f1120,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706,f5405,f3707,f3709,f3724,f5470,f191,f185,f5541,f5540,f5664,f5663,f5676,f5675,f5748,f5694,f5703,f5705,f5707,f5710,f5711,f5712,f5714,f5717,f5718,f5719,f5722,f5723,f5729,f5732,f5734,f5737,f188,f5770,f5771,f5772,f5773,f5775,f5769,f5768])).
fof(f6728,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f192,f207,f208,f211,f231,f263,f308,f425,f427,f430,f450,f458,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f904,f903,f902,f906,f909,f931,f930,f941,f940,f960,f969,f994,f993,f1012,f1011,f1025,f1027,f1039,f1076,f1075,f1160,f1103,f1112,f1115,f1117,f1119,f1122,f1124,f1126,f1128,f1132,f1135,f1137,f1140,f1168,f1167,f1195,f1194,f1229,f1246,f1247,f1249,f1253,f1258,f1302,f1301,f1332,f1331,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1127,f1120,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706,f5405,f3707,f3709,f3724,f5470,f191,f185,f5541,f5540,f5664,f5663,f5676,f5675,f5748,f5694,f5703,f5705,f5707,f5710,f5711,f5712,f5714,f5717,f5718,f5719,f5722,f5723,f5729,f5732,f5734,f5737,f188,f5770,f5771,f5772,f5773,f5775,f5769])).
fof(f6727,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f192,f207,f208,f211,f231,f263,f308,f425,f427,f430,f450,f458,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f904,f903,f902,f906,f909,f931,f930,f941,f940,f960,f969,f994,f993,f1012,f1011,f1025,f1027,f1039,f1076,f1075,f1160,f1103,f1112,f1115,f1117,f1119,f1122,f1124,f1126,f1128,f1132,f1135,f1137,f1140,f1168,f1167,f1195,f1194,f1229,f1246,f1247,f1249,f1253,f1258,f1302,f1301,f1332,f1331,f1418,f1507,f1509,f1513,f2028,f2031,f2032,f2034,f1030,f1032,f1939,f1516,f1127,f1120,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f3734,f3822,f3829,f3702,f3704,f3706,f5405,f3707,f3709,f3724,f5470,f191,f185,f5541,f5540,f5664,f5663,f5676,f5675,f5748,f5694,f5703,f5705,f5707,f5710,f5711,f5712,f5714,f5717,f5718,f5719,f5722,f5723,f5729,f5732,f5734,f5737,f188,f5770,f5771,f5772,f5773,f5775])).
fof(f6718,plain,(
  ( ! [X2,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl32),
  inference(cnf_transformation,[],[f6718_D])).
fof(f6718_D,plain,(
  ( ! [X2,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f6721,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(forward_demodulation,[],[f6720,f160])).
fof(f6720,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))))) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(condensation,[],[f6719])).
fof(f6719,plain,(
  ( ! [X26,X24,X23,X25] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X24),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X25),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))))) | p(s(bool,X25)) | p(s(bool,X24)) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) ) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(subsumption_resolution,[],[f6705,f263])).
fof(f6705,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,X23)) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X24),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X25),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))))) | p(s(bool,X25)) | p(s(bool,X24)) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) ) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(superposition,[],[f5780,f231])).
fof(f5780,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))))) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(forward_demodulation,[],[f5777,f174])).
fof(f5777,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(backward_demodulation,[],[f188,f5766])).
fof(f5766,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl3 | $spl5 | $spl18)),
  inference(forward_demodulation,[],[f5765,f157])).
fof(f5765,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl3 | $spl5 | $spl18)),
  inference(forward_demodulation,[],[f5764,f157])).
fof(f5764,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl3 | $spl5 | $spl18)),
  inference(forward_demodulation,[],[f5763,f5685])).
fof(f5685,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',sK2))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2)))))) ) | ($spl3 | $spl5)),
  inference(unit_resulting_resolution,[],[f193,f191,f185])).
fof(f6717,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(forward_demodulation,[],[f6716,f160])).
fof(f6716,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))))) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(condensation,[],[f6715])).
fof(f6715,plain,(
  ( ! [X21,X19,X22,X20] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X20),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X21),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))))) | p(s(bool,X21)) | p(s(bool,X20)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(subsumption_resolution,[],[f6704,f263])).
fof(f6704,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,X19)) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X20),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,X21),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))))) | p(s(bool,X21)) | p(s(bool,X20)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(superposition,[],[f5780,f231])).
fof(f6713,plain,(
  $false | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(subsumption_resolution,[],[f6712,f161])).
fof(f6712,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(forward_demodulation,[],[f6711,f160])).
fof(f6711,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(forward_demodulation,[],[f6697,f161])).
fof(f6697,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK0))))))))) | ($spl0 | $spl3 | $spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f202,f202,f122,f5780])).
fof(f5756,plain,(
  $false | ($spl3 | $spl5 | $spl20)),
  inference(subsumption_resolution,[],[f5750,f161])).
fof(f5750,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))) | ($spl3 | $spl5 | $spl20)),
  inference(backward_demodulation,[],[f5685,f5619])).
fof(f5619,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))) | $spl20),
  inference(forward_demodulation,[],[f5618,f161])).
fof(f5618,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))) | $spl20),
  inference(forward_demodulation,[],[f5570,f161])).
fof(f5570,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | $spl20),
  inference(unit_resulting_resolution,[],[f202,f202,f202,f5533])).
fof(f5533,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | $spl20),
  inference(forward_demodulation,[],[f5532,f157])).
fof(f5532,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | $spl20),
  inference(forward_demodulation,[],[f5531,f157])).
fof(f5531,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | $spl20),
  inference(forward_demodulation,[],[f201,f477])).
fof(f5530,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5512,f160])).
fof(f5512,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | $spl12),
  inference(unit_resulting_resolution,[],[f122,f202,f202,f205])).
fof(f205,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X6)) | p(s(bool,X4)) | p(s(bool,X5)) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2)))))))))))) ) | $spl12),
  inference(forward_demodulation,[],[f204,f157])).
fof(f204,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl12),
  inference(forward_demodulation,[],[f203,f157])).
fof(f203,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK3))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl12),
  inference(forward_demodulation,[],[f197,f168])).
fof(f5529,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5513,f122])).
fof(f5513,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(unit_resulting_resolution,[],[f202,f202,f160,f205])).
fof(f5516,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f202,f202,f122,f160,f205])).
fof(f2433,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399,f2402,f2405,f2408,f2410,f2417,f2420,f2423,f2426,f2428,f2043])).
fof(f2043,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl28),
  inference(unit_resulting_resolution,[],[f2030,f136])).
fof(f2428,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2046,f157])).
fof(f2046,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X4)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f168,f157,f2030,f165])).
fof(f2426,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2425,f157])).
fof(f2425,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2047,f174])).
fof(f2047,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real',X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f174,f157,f2030,f165])).
fof(f2423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2422,f157])).
fof(f2422,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2048,f157])).
fof(f2048,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X4)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f157,f157,f2030,f165])).
fof(f2420,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2419,f157])).
fof(f2419,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2049,f160])).
fof(f2049,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X4))),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f160,f157,f2030,f165])).
fof(f2417,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X4),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2416,f157])).
fof(f2416,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X4),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2415,f160])).
fof(f2415,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X4),s(X0,X2))),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2050,f2129])).
fof(f2050,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X2))),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f161,f157,f2030,f165])).
fof(f2410,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2052,f157])).
fof(f2052,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X4)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f168,f157,f2030,f165])).
fof(f2408,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2407,f157])).
fof(f2407,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X5))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2053,f174])).
fof(f2053,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real',X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f174,f157,f2030,f165])).
fof(f2405,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2404,f157])).
fof(f2404,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2054,f157])).
fof(f2054,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X4)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f157,f157,f2030,f165])).
fof(f2402,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2401,f157])).
fof(f2401,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2055,f160])).
fof(f2055,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X4))),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f160,f157,f2030,f165])).
fof(f2399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X4),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2398,f157])).
fof(f2398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X4),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2397,f160])).
fof(f2397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X4),s(X0,X2))),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2056,f2129])).
fof(f2056,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X2))),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f161,f157,f2030,f165])).
fof(f2392,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2058,f160])).
fof(f2058,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f168,f160,f2030,f165])).
fof(f2390,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2389,f160])).
fof(f2389,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2059,f174])).
fof(f2059,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X5))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real',X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f174,f160,f2030,f165])).
fof(f2387,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2386,f160])).
fof(f2386,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2060,f157])).
fof(f2060,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f157,f160,f2030,f165])).
fof(f2384,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2383,f160])).
fof(f2383,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,X2),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2061,f160])).
fof(f2061,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X5))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f160,f160,f2030,f165])).
fof(f2381,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X5),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2380,f160])).
fof(f2380,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,X5),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2379,f160])).
fof(f2379,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X5),s(X0,X2))),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2062,f2129])).
fof(f2062,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X5),s(X0,X2))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f161,f160,f2030,f165])).
fof(f2374,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2064,f160])).
fof(f2064,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f168,f160,f2030,f165])).
fof(f2372,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2371,f174])).
fof(f2371,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2065,f160])).
fof(f2065,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f174,f160,f2030,f165])).
fof(f2369,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2368,f157])).
fof(f2368,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2066,f160])).
fof(f2066,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f157,f160,f2030,f165])).
fof(f2366,plain,(
  ( ! [X6,X4,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2365,f160])).
fof(f2365,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X6))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2067,f160])).
fof(f2067,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X3),s(X0,X5))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f160,f160,f2030,f165])).
fof(f2363,plain,(
  ( ! [X6,X4,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X5),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2362,f160])).
fof(f2362,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X5),s(X0,X6))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2361,f160])).
fof(f2361,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X5),s(X0,X3))),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2068,f2129])).
fof(f2068,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X5),s(X0,X3))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f161,f160,f2030,f165])).
fof(f2355,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2354,f160])).
fof(f2354,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2070,f2129])).
fof(f2070,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f168,f161,f2030,f165])).
fof(f2352,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2351,f160])).
fof(f2351,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2350,f2129])).
fof(f2350,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2071,f174])).
fof(f2071,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real',X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f174,f161,f2030,f165])).
fof(f2348,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2347,f160])).
fof(f2347,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2346,f2129])).
fof(f2346,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X1))),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2072,f157])).
fof(f2072,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X1))),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f157,f161,f2030,f165])).
fof(f2344,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2343,f160])).
fof(f2343,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X6))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2342,f2129])).
fof(f2342,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X1))),s(X0,X2),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2073,f160])).
fof(f2073,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X1))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X5))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f160,f161,f2030,f165])).
fof(f2340,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X5),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2339,f160])).
fof(f2339,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X1))),s(X0,X5),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2338,f160])).
fof(f2338,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X1))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X5),s(X0,X2))),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2074,f2129])).
fof(f2074,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X1))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X5),s(X0,X2))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f161,f161,f2030,f165])).
fof(f2332,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2331,f160])).
fof(f2331,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2076,f2129])).
fof(f2076,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f168,f161,f2030,f165])).
fof(f2329,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2328,f160])).
fof(f2328,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2327,f174])).
fof(f2327,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2077,f2129])).
fof(f2077,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f174,f161,f2030,f165])).
fof(f2325,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2324,f160])).
fof(f2324,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2323,f157])).
fof(f2323,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2078,f2129])).
fof(f2078,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X1))),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f157,f161,f2030,f165])).
fof(f2321,plain,(
  ( ! [X6,X2,X0,X3] : (s(X0,X3) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X3),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2320,f160])).
fof(f2320,plain,(
  ( ! [X6,X4,X2,X0,X3] : (s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X3),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2319,f2129])).
fof(f2319,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X3),s(X0,X6))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2079,f160])).
fof(f2079,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X3),s(X0,X5))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f160,f161,f2030,f165])).
fof(f2317,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X5),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2316,f160])).
fof(f2316,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X5),s(X0,X6))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2315,f160])).
fof(f2315,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X5),s(X0,X3))),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2080,f2129])).
fof(f2080,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X5),s(X0,X3))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f161,f161,f2030,f165])).
fof(f2310,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2082,f157])).
fof(f2082,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X4)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f168,f157,f2030,f165])).
fof(f2308,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2307,f157])).
fof(f2307,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X5))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2083,f174])).
fof(f2083,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real',X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f174,f157,f2030,f165])).
fof(f2305,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2304,f157])).
fof(f2304,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2084,f157])).
fof(f2084,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X4)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f157,f157,f2030,f165])).
fof(f2302,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2301,f157])).
fof(f2301,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2085,f160])).
fof(f2085,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X4))),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f160,f157,f2030,f165])).
fof(f2299,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X4),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2298,f157])).
fof(f2298,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X4),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2297,f160])).
fof(f2297,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X4),s(X0,X2))),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2086,f2129])).
fof(f2086,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X2))),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f161,f157,f2030,f165])).
fof(f2292,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2088,f157])).
fof(f2088,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X4)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f168,f157,f2030,f165])).
fof(f2290,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2289,f157])).
fof(f2289,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2089,f174])).
fof(f2089,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real',X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f174,f157,f2030,f165])).
fof(f2287,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2286,f157])).
fof(f2286,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2090,f157])).
fof(f2090,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X4)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f157,f157,f2030,f165])).
fof(f2284,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2283,f157])).
fof(f2283,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2091,f160])).
fof(f2091,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X4))),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f160,f157,f2030,f165])).
fof(f2281,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X4),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2280,f157])).
fof(f2280,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,X4),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2279,f160])).
fof(f2279,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X4),s(X0,X2))),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2092,f2129])).
fof(f2092,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X1))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X2))),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f161,f157,f2030,f165])).
fof(f2274,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2094,f160])).
fof(f2094,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f168,f160,f2030,f165])).
fof(f2272,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2271,f174])).
fof(f2271,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2095,f160])).
fof(f2095,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f174,f160,f2030,f165])).
fof(f2269,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2268,f157])).
fof(f2268,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2096,f160])).
fof(f2096,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f157,f160,f2030,f165])).
fof(f2266,plain,(
  ( ! [X6,X4,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2265,f160])).
fof(f2265,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X6))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2097,f160])).
fof(f2097,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X3),s(X0,X5))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f160,f160,f2030,f165])).
fof(f2263,plain,(
  ( ! [X6,X4,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X5),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2262,f160])).
fof(f2262,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X5),s(X0,X6))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2261,f160])).
fof(f2261,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X5),s(X0,X3))),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2098,f2129])).
fof(f2098,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X5),s(X0,X3))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f161,f160,f2030,f165])).
fof(f2256,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2100,f160])).
fof(f2100,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f168,f160,f2030,f165])).
fof(f2254,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2253,f160])).
fof(f2253,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X5))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2101,f174])).
fof(f2101,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X5))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real',X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f174,f160,f2030,f165])).
fof(f2251,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2250,f160])).
fof(f2250,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2102,f157])).
fof(f2102,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f157,f160,f2030,f165])).
fof(f2248,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2247,f160])).
fof(f2247,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,X2),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2103,f160])).
fof(f2103,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X5))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f160,f160,f2030,f165])).
fof(f2245,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X5),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2244,f160])).
fof(f2244,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,X5),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2243,f160])).
fof(f2243,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X5),s(X0,X2))),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2104,f2129])).
fof(f2104,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X4))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X5),s(X0,X2))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f161,f160,f2030,f165])).
fof(f2237,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2236,f160])).
fof(f2236,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2106,f2129])).
fof(f2106,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f168,f161,f2030,f165])).
fof(f2234,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2233,f160])).
fof(f2233,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2232,f174])).
fof(f2232,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2107,f2129])).
fof(f2107,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real',X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f174,f161,f2030,f165])).
fof(f2230,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2229,f160])).
fof(f2229,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2228,f157])).
fof(f2228,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2108,f2129])).
fof(f2108,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X1))),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f157,f161,f2030,f165])).
fof(f2226,plain,(
  ( ! [X6,X2,X0,X3] : (s(X0,X3) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X3),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2225,f160])).
fof(f2225,plain,(
  ( ! [X6,X4,X2,X0,X3] : (s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X3),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2224,f2129])).
fof(f2224,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X3),s(X0,X6))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2109,f160])).
fof(f2109,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X3),s(X0,X5))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f160,f161,f2030,f165])).
fof(f2222,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X5),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2221,f160])).
fof(f2221,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X5),s(X0,X6))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2220,f160])).
fof(f2220,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X5),s(X0,X3))),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2110,f2129])).
fof(f2110,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X2))),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X5),s(X0,X3))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f161,f161,f2030,f165])).
fof(f2214,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2213,f160])).
fof(f2213,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2112,f2129])).
fof(f2112,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))),s('type.realax.real',X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f168,f161,f2030,f165])).
fof(f2211,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2210,f160])).
fof(f2210,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4)))) ) | $spl28),
  inference(forward_demodulation,[],[f2209,f2129])).
fof(f2209,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2113,f174])).
fof(f2113,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real',X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f174,f161,f2030,f165])).
fof(f2207,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2206,f160])).
fof(f2206,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2205,f2129])).
fof(f2205,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X1))),s(X0,X2),s(X0,X5)))) ) | $spl28),
  inference(forward_demodulation,[],[f2114,f157])).
fof(f2114,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X1))),s(X0,'const.marker.unint_1'(s(X0,X2))),s(X0,X5)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f157,f161,f2030,f165])).
fof(f2203,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2202,f160])).
fof(f2202,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X6))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f2201,f2129])).
fof(f2201,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X1))),s(X0,X2),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2115,f160])).
fof(f2115,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X1))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X5))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f160,f161,f2030,f165])).
fof(f2199,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X5),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2198,f160])).
fof(f2198,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X1))),s(X0,X5),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2197,f160])).
fof(f2197,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X1))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X5),s(X0,X2))),s(X0,X6)))) ) | $spl28),
  inference(forward_demodulation,[],[f2116,f2129])).
fof(f2116,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X1))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X5),s(X0,X2))),s(X0,X6)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f161,f161,f2030,f165])).
fof(f2118,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl28),
  inference(unit_resulting_resolution,[],[f2030,f2030,f172])).
fof(f2124,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl28),
  inference(unit_resulting_resolution,[],[f2030,f2030,f172])).
fof(f2174,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) ) | $spl28),
  inference(backward_demodulation,[],[f2129,f1939])).
fof(f2165,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real',X20) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X17),s('type.realax.real',X18))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X19))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X20),s('type.realax.real',X19)))) ) | $spl28),
  inference(backward_demodulation,[],[f2129,f910])).
fof(f2160,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real',X20) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X17),s('type.realax.real',X18))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X19),s('type.realax.real',X18))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X19),s('type.realax.real',X20)))) ) | $spl28),
  inference(backward_demodulation,[],[f2129,f431])).
fof(f2181,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,X3)) ) | $spl28),
  inference(subsumption_resolution,[],[f2156,f2129])).
fof(f2156,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.marker.unint_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) ) | $spl28),
  inference(backward_demodulation,[],[f2129,f368])).
fof(f2149,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | $spl28),
  inference(backward_demodulation,[],[f2129,f161])).
fof(f2129,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(unit_resulting_resolution,[],[f2030,f133])).
fof(f2432,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2431,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399,f2402,f2405,f2408,f2410,f2417,f2420,f2423,f2426,f2428])).
fof(f2431,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X2)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2430,f157])).
fof(f2430,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real',sK0),s('type.realax.real',X2)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2045,f192])).
fof(f2045,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) ) | ($spl4 | $spl28)),
  inference(unit_resulting_resolution,[],[f192,f157,f2030,f165])).
fof(f2429,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399,f2402,f2405,f2408,f2410,f2417,f2420,f2423,f2426,f2428])).
fof(f2427,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399,f2402,f2405,f2408,f2410,f2417,f2420,f2423,f2426])).
fof(f2424,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399,f2402,f2405,f2408,f2410,f2417,f2420,f2423])).
fof(f2421,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399,f2402,f2405,f2408,f2410,f2417,f2420])).
fof(f2418,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399,f2402,f2405,f2408,f2410,f2417])).
fof(f2414,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2413,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399,f2402,f2405,f2408,f2410])).
fof(f2413,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X2)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2412,f192])).
fof(f2412,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X2))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2051,f157])).
fof(f2051,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real',sK0),s('type.realax.real',X2)))) ) | ($spl4 | $spl28)),
  inference(unit_resulting_resolution,[],[f192,f157,f2030,f165])).
fof(f2411,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399,f2402,f2405,f2408,f2410])).
fof(f2409,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399,f2402,f2405,f2408])).
fof(f2406,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399,f2402,f2405])).
fof(f2403,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399,f2402])).
fof(f2400,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392,f2399])).
fof(f2396,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2395,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392])).
fof(f2395,plain,(
  ( ! [X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2394,f160])).
fof(f2394,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X2))),s('type.realax.real',sK0),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2057,f192])).
fof(f2057,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(unit_resulting_resolution,[],[f192,f160,f2030,f165])).
fof(f2393,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390,f2392])).
fof(f2391,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387,f2390])).
fof(f2388,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384,f2387])).
fof(f2385,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381,f2384])).
fof(f2382,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374,f2381])).
fof(f2378,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2377,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374])).
fof(f2377,plain,(
  ( ! [X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2376,f192])).
fof(f2376,plain,(
  ( ! [X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2063,f160])).
fof(f2063,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X2))),s('type.realax.real',sK0),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(unit_resulting_resolution,[],[f192,f160,f2030,f165])).
fof(f2375,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372,f2374])).
fof(f2373,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369,f2372])).
fof(f2370,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366,f2369])).
fof(f2367,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363,f2366])).
fof(f2364,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355,f2363])).
fof(f2360,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2359,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355])).
fof(f2359,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2358,f160])).
fof(f2358,plain,(
  ( ! [X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',sK0),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2357,f2129])).
fof(f2357,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X0))),s('type.realax.real',sK0),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2069,f192])).
fof(f2069,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(unit_resulting_resolution,[],[f192,f161,f2030,f165])).
fof(f2356,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352,f2355])).
fof(f2353,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348,f2352])).
fof(f2349,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344,f2348])).
fof(f2345,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340,f2344])).
fof(f2341,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332,f2340])).
fof(f2337,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2336,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332])).
fof(f2336,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2335,f160])).
fof(f2335,plain,(
  ( ! [X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',sK0),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2334,f192])).
fof(f2334,plain,(
  ( ! [X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',sK0),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2075,f2129])).
fof(f2075,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X0))),s('type.realax.real',sK0),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(unit_resulting_resolution,[],[f192,f161,f2030,f165])).
fof(f2333,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329,f2332])).
fof(f2330,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325,f2329])).
fof(f2326,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321,f2325])).
fof(f2322,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317,f2321])).
fof(f2318,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310,f2317])).
fof(f2314,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2313,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310])).
fof(f2313,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X2)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2312,f192])).
fof(f2312,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X2))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2081,f157])).
fof(f2081,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real',sK0),s('type.realax.real',X2)))) ) | ($spl4 | $spl28)),
  inference(unit_resulting_resolution,[],[f192,f157,f2030,f165])).
fof(f2311,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308,f2310])).
fof(f2309,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305,f2308])).
fof(f2306,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302,f2305])).
fof(f2303,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299,f2302])).
fof(f2300,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292,f2299])).
fof(f2296,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2295,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292])).
fof(f2295,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X2)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2294,f157])).
fof(f2294,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real',sK0),s('type.realax.real',X2)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2087,f192])).
fof(f2087,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.marker.unint_1'(s(bool,X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) ) | ($spl4 | $spl28)),
  inference(unit_resulting_resolution,[],[f192,f157,f2030,f165])).
fof(f2293,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290,f2292])).
fof(f2291,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287,f2290])).
fof(f2288,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284,f2287])).
fof(f2285,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281,f2284])).
fof(f2282,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274,f2281])).
fof(f2278,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2277,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274])).
fof(f2277,plain,(
  ( ! [X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2276,f192])).
fof(f2276,plain,(
  ( ! [X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2093,f160])).
fof(f2093,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X2))),s('type.realax.real',sK0),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(unit_resulting_resolution,[],[f192,f160,f2030,f165])).
fof(f2275,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272,f2274])).
fof(f2273,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269,f2272])).
fof(f2270,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266,f2269])).
fof(f2267,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263,f2266])).
fof(f2264,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256,f2263])).
fof(f2260,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2259,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256])).
fof(f2259,plain,(
  ( ! [X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2258,f160])).
fof(f2258,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X2))),s('type.realax.real',sK0),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2099,f192])).
fof(f2099,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(unit_resulting_resolution,[],[f192,f160,f2030,f165])).
fof(f2257,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254,f2256])).
fof(f2255,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251,f2254])).
fof(f2252,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248,f2251])).
fof(f2249,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245,f2248])).
fof(f2246,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237,f2245])).
fof(f2242,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2241,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237])).
fof(f2241,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2240,f160])).
fof(f2240,plain,(
  ( ! [X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',sK0),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2239,f192])).
fof(f2239,plain,(
  ( ! [X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',sK0),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2105,f2129])).
fof(f2105,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X0))),s('type.realax.real',sK0),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(unit_resulting_resolution,[],[f192,f161,f2030,f165])).
fof(f2238,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234,f2237])).
fof(f2235,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230,f2234])).
fof(f2231,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226,f2230])).
fof(f2227,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222,f2226])).
fof(f2223,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214,f2222])).
fof(f2219,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2218,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214])).
fof(f2218,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2217,f160])).
fof(f2217,plain,(
  ( ! [X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',sK0),s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2216,f2129])).
fof(f2216,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X0))),s('type.realax.real',sK0),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f2111,f192])).
fof(f2111,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real',sK0),s('type.realax.real',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3)))) ) | ($spl4 | $spl28)),
  inference(unit_resulting_resolution,[],[f192,f161,f2030,f165])).
fof(f2215,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211,f2214])).
fof(f2212,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207,f2211])).
fof(f2208,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203,f2207])).
fof(f2204,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199,f2203])).
fof(f2200,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118,f2199])).
fof(f2196,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124,f2118])).
fof(f2195,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f2119,f122])).
fof(f2119,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(unit_resulting_resolution,[],[f211,f2030,f172])).
fof(f2194,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f2121,f2030])).
fof(f2121,plain,(
  ~p(s(bool,'F_0')) | $spl28),
  inference(unit_resulting_resolution,[],[f211,f2030,f172])).
fof(f2193,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174,f2124])).
fof(f2192,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f2125,f122])).
fof(f2125,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(unit_resulting_resolution,[],[f211,f2030,f172])).
fof(f2191,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f2127,f2030])).
fof(f2127,plain,(
  ~p(s(bool,'F_0')) | $spl28),
  inference(unit_resulting_resolution,[],[f211,f2030,f172])).
fof(f2190,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165,f2174])).
fof(f2189,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2172,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165])).
fof(f2172,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X29),s('type.realax.real',X30))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X32))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X33),s('type.realax.real',X30))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X31),s('type.realax.real',X30))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X30))))) | s('type.realax.real',sK0) = s('type.realax.real',X32)) ) | ($spl4 | $spl28)),
  inference(backward_demodulation,[],[f2129,f1127])).
fof(f2188,plain,(
  $false | ($spl4 | $spl28)),
  inference(global_subsumption,[],[f2171,f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165])).
fof(f2171,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X29),s('type.realax.real',X30))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X30))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X32),s('type.realax.real',X33))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X31),s('type.realax.real',X33))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X32))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X33))))) | s('type.realax.real',sK0) = s('type.realax.real',X33)) ) | ($spl4 | $spl28)),
  inference(backward_demodulation,[],[f2129,f1120])).
fof(f2186,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160,f2165])).
fof(f2185,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181,f2160])).
fof(f2182,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f130,f133,f132,f157,f160,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f263,f308,f427,f430,f477,f486,f532,f533,f547,f557,f558,f629,f877,f878,f885,f906,f909,f960,f969,f1025,f1027,f1039,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1030,f1032,f1516,f1054,f1053,f1052,f1034,f1029,f910,f873,f559,f534,f516,f431,f410,f409,f370,f368,f367,f366,f322,f287,f233,f202,f161,f137,f136,f1939,f2028,f2031,f2032,f2034,f2030,f2129,f2149,f2181])).
fof(f2180,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f2150,f122])).
fof(f2150,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(backward_demodulation,[],[f2129,f202])).
fof(f2179,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f2178,f2030])).
fof(f2178,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl28),
  inference(subsumption_resolution,[],[f2147,f2129])).
fof(f2147,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl28),
  inference(backward_demodulation,[],[f2129,f136])).
fof(f2130,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f211,f2030,f133])).
fof(f2128,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f2030,f211,f2030,f172])).
fof(f2126,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f122,f211,f2030,f172])).
fof(f2122,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f2030,f211,f2030,f172])).
fof(f2120,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f122,f211,f2030,f172])).
fof(f2044,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f157,f2030,f136])).
fof(f2042,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f202,f2030])).
fof(f2030,plain,(
  ( ! [X17] : (p(s(bool,X17))) ) | $spl28),
  inference(cnf_transformation,[],[f2030_D])).
fof(f2030_D,plain,(
  ( ! [X17] : (p(s(bool,X17))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2035,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(X56,X58) = s(X56,'const.bool.COND_3'(s(bool,X59),s(X56,X60),s(X56,X57))) | s(bool,'F_0') != s(bool,X59) | s(X56,X57) != s(X56,X58)) ) | $spl30),
  inference(cnf_transformation,[],[f2035_D])).
fof(f2035_D,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(X56,X58) = s(X56,'const.bool.COND_3'(s(bool,X59),s(X56,X60),s(X56,X57))) | s(bool,'F_0') != s(bool,X59) | s(X56,X57) != s(X56,X58)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2034,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (s(X56,X58) = s(X56,'const.bool.COND_3'(s(bool,X59),s(X56,X60),s(X56,X57))) | p(s(bool,X55)) | s(X56,X57) != s(X56,X58) | s(bool,'F_0') != s(bool,X59)) )),
  inference(forward_demodulation,[],[f2033,f161])).
fof(f2033,plain,(
  ( ! [X61,X59,X57,X60,X58,X56,X55] : (p(s(bool,X55)) | s(X56,X57) != s(X56,X58) | s(X56,'const.bool.COND_3'(s(bool,'F_0'),s(X56,X61),s(X56,X58))) = s(X56,'const.bool.COND_3'(s(bool,X59),s(X56,X60),s(X56,X57))) | s(bool,'F_0') != s(bool,X59)) )),
  inference(subsumption_resolution,[],[f1921,f132])).
fof(f1921,plain,(
  ( ! [X61,X59,X57,X60,X58,X56,X55] : (p(s(bool,X55)) | s(X56,X57) != s(X56,X58) | s(X56,'const.bool.COND_3'(s(bool,'F_0'),s(X56,X61),s(X56,X58))) = s(X56,'const.bool.COND_3'(s(bool,X59),s(X56,X60),s(X56,X57))) | s(bool,'F_0') != s(bool,X59) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f166,f130])).
fof(f2029,plain,(
  ( ! [X23,X21,X19,X22,X20,X18,X16] : (s(X18,X19) != s(X18,X20) | s(bool,'T_0') = s(bool,X16) | s(bool,X16) != s(bool,X21) | s(X18,'const.bool.COND_3'(s(bool,X16),s(X18,X23),s(X18,X20))) = s(X18,'const.bool.COND_3'(s(bool,X21),s(X18,X22),s(X18,X19)))) ) | $spl26),
  inference(cnf_transformation,[],[f2029_D])).
fof(f2029_D,plain,(
  ( ! [X23,X21,X19,X22,X20,X18,X16] : (s(X18,X19) != s(X18,X20) | s(bool,'T_0') = s(bool,X16) | s(bool,X16) != s(bool,X21) | s(X18,'const.bool.COND_3'(s(bool,X16),s(X18,X23),s(X18,X20))) = s(X18,'const.bool.COND_3'(s(bool,X21),s(X18,X22),s(X18,X19)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2031,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X25)) | s(X26,X27) != s(X26,X28) | s(X26,'const.bool.COND_3'(s(bool,X24),s(X26,X31),s(X26,X28))) = s(X26,'const.bool.COND_3'(s(bool,X29),s(X26,X30),s(X26,X27))) | s(bool,X24) != s(bool,X29) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1917,f132])).
fof(f1917,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X25)) | s(X26,X27) != s(X26,X28) | s(X26,'const.bool.COND_3'(s(bool,X24),s(X26,X31),s(X26,X28))) = s(X26,'const.bool.COND_3'(s(bool,X29),s(X26,X30),s(X26,X27))) | s(bool,X24) != s(bool,X29) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f166,f231])).
fof(f2028,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,X17)) | s(X18,X19) != s(X18,X20) | s(X18,'const.bool.COND_3'(s(bool,X16),s(X18,X23),s(X18,X20))) = s(X18,'const.bool.COND_3'(s(bool,X21),s(X18,X22),s(X18,X19))) | s(bool,X16) != s(bool,X21) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f1916,f132])).
fof(f1916,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,X17)) | s(X18,X19) != s(X18,X20) | s(X18,'const.bool.COND_3'(s(bool,X16),s(X18,X23),s(X18,X20))) = s(X18,'const.bool.COND_3'(s(bool,X21),s(X18,X22),s(X18,X19))) | s(bool,X16) != s(bool,X21) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f166,f231])).
fof(f1565,plain,(
  $false | ($spl4 | $spl24)),
  inference(global_subsumption,[],[f1548,f121,f120,f119,f118,f117,f116,f115,f114,f122,f202,f130,f133,f132,f137,f136,f157,f160,f161,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f233,f263,f287,f308,f322,f366,f367,f368,f370,f409,f410,f427,f430,f431,f477,f486,f516,f532,f533,f534,f547,f557,f558,f559,f629,f873,f877,f878,f885,f906,f909,f910,f960,f969,f1025,f1027,f1029,f1030,f1032,f1034,f1039,f1052,f1053,f1054,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1516,f1511,f1553,f1525,f1557,f1536,f1541])).
fof(f1541,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real',X20) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X17),s('type.realax.real',X18))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X19))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X20),s('type.realax.real',X19)))) ) | $spl24),
  inference(backward_demodulation,[],[f1511,f910])).
fof(f1536,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real',X20) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X17),s('type.realax.real',X18))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X19),s('type.realax.real',X18))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X19),s('type.realax.real',X20)))) ) | $spl24),
  inference(backward_demodulation,[],[f1511,f431])).
fof(f1557,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,X3)) ) | $spl24),
  inference(subsumption_resolution,[],[f1532,f1511])).
fof(f1532,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.marker.unint_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) ) | $spl24),
  inference(backward_demodulation,[],[f1511,f368])).
fof(f1525,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | $spl24),
  inference(backward_demodulation,[],[f1511,f161])).
fof(f1553,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f1523,f1511])).
fof(f1523,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl24),
  inference(backward_demodulation,[],[f1511,f136])).
fof(f1548,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X29),s('type.realax.real',X30))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X32))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X33),s('type.realax.real',X30))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X31),s('type.realax.real',X30))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X30))))) | s('type.realax.real',sK0) = s('type.realax.real',X32)) ) | ($spl4 | $spl24)),
  inference(backward_demodulation,[],[f1511,f1127])).
fof(f1564,plain,(
  $false | ($spl4 | $spl24)),
  inference(global_subsumption,[],[f1547,f121,f120,f119,f118,f117,f116,f115,f114,f122,f202,f130,f133,f132,f137,f136,f157,f160,f161,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f233,f263,f287,f308,f322,f366,f367,f368,f370,f409,f410,f427,f430,f431,f477,f486,f516,f532,f533,f534,f547,f557,f558,f559,f629,f873,f877,f878,f885,f906,f909,f910,f960,f969,f1025,f1027,f1029,f1030,f1032,f1034,f1039,f1052,f1053,f1054,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1516,f1511,f1553,f1525,f1557,f1536,f1541])).
fof(f1547,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X29),s('type.realax.real',X30))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X30))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X32),s('type.realax.real',X33))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X31),s('type.realax.real',X33))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X32))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X33))))) | s('type.realax.real',sK0) = s('type.realax.real',X33)) ) | ($spl4 | $spl24)),
  inference(backward_demodulation,[],[f1511,f1120])).
fof(f1562,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f202,f130,f133,f132,f137,f136,f157,f160,f161,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f233,f263,f287,f308,f322,f366,f367,f368,f370,f409,f410,f427,f430,f431,f477,f486,f516,f532,f533,f534,f547,f557,f558,f559,f629,f873,f877,f878,f885,f906,f909,f910,f960,f969,f1025,f1027,f1029,f1030,f1032,f1034,f1039,f1052,f1053,f1054,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1516,f1511,f1553,f1525,f1557,f1536,f1541])).
fof(f1561,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f202,f130,f133,f132,f137,f136,f157,f160,f161,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f233,f263,f287,f308,f322,f366,f367,f368,f370,f409,f410,f427,f430,f431,f477,f486,f516,f532,f533,f534,f547,f557,f558,f559,f629,f873,f877,f878,f885,f906,f909,f910,f960,f969,f1025,f1027,f1029,f1030,f1032,f1034,f1039,f1052,f1053,f1054,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1516,f1511,f1553,f1525,f1557,f1536])).
fof(f1558,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f202,f130,f133,f132,f137,f136,f157,f160,f161,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f233,f263,f287,f308,f322,f366,f367,f368,f370,f409,f410,f427,f430,f431,f477,f486,f516,f532,f533,f534,f547,f557,f558,f559,f629,f873,f877,f878,f885,f906,f909,f910,f960,f969,f1025,f1027,f1029,f1030,f1032,f1034,f1039,f1052,f1053,f1054,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1516,f1511,f1553,f1525,f1557])).
fof(f1556,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1526,f122])).
fof(f1526,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f1511,f202])).
fof(f1555,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f202,f130,f133,f132,f137,f136,f157,f160,f161,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f233,f263,f287,f308,f322,f366,f367,f368,f370,f409,f410,f427,f430,f431,f477,f486,f516,f532,f533,f534,f547,f557,f558,f559,f629,f873,f877,f878,f885,f906,f909,f910,f960,f969,f1025,f1027,f1029,f1030,f1032,f1034,f1039,f1052,f1053,f1054,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1516,f1511,f1553,f1525])).
fof(f1554,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f121,f120,f119,f118,f117,f116,f115,f114,f122,f202,f130,f133,f132,f137,f136,f157,f160,f161,f167,f166,f165,f168,f172,f169,f173,f174,f176,f175,f178,f177,f185,f211,f231,f233,f263,f287,f308,f322,f366,f367,f368,f370,f409,f410,f427,f430,f431,f477,f486,f516,f532,f533,f534,f547,f557,f558,f559,f629,f873,f877,f878,f885,f906,f909,f910,f960,f969,f1025,f1027,f1029,f1030,f1032,f1034,f1039,f1052,f1053,f1054,f1247,f1249,f1253,f1258,f1418,f1507,f1509,f1513,f1516,f1511,f1553])).
fof(f1510,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X1,'const.bool.COND_3'(s(bool,'T_0'),s(X1,X2),s(X1,X3))) = s(X1,'const.bool.COND_3'(s(bool,X4),s(X1,X5),s(X1,X6))) | s(bool,'T_0') != s(bool,X4) | s(X1,X2) != s(X1,X5)) ) | $spl22),
  inference(cnf_transformation,[],[f1510_D])).
fof(f1510_D,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X1,'const.bool.COND_3'(s(bool,'T_0'),s(X1,X2),s(X1,X3))) = s(X1,'const.bool.COND_3'(s(bool,X4),s(X1,X5),s(X1,X6))) | s(bool,'T_0') != s(bool,X4) | s(X1,X2) != s(X1,X5)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1513,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,'const.bool.COND_3'(s(bool,'T_0'),s(X1,X2),s(X1,X3))) = s(X1,'const.bool.COND_3'(s(bool,X4),s(X1,X5),s(X1,X6))) | s(X1,X2) != s(X1,X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(condensation,[],[f1512])).
fof(f1512,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (s(X48,'const.bool.COND_3'(s(bool,'T_0'),s(X48,X52),s(X48,X53))) = s(X48,'const.bool.COND_3'(s(bool,X49),s(X48,X50),s(X48,X51))) | s(X48,X50) != s(X48,X52) | s(bool,'T_0') != s(bool,X49) | s(bool,X47) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1399,f263])).
fof(f1399,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (~p(s(bool,X47)) | s(X48,'const.bool.COND_3'(s(bool,'T_0'),s(X48,X52),s(X48,X53))) = s(X48,'const.bool.COND_3'(s(bool,X49),s(X48,X50),s(X48,X51))) | s(X48,X50) != s(X48,X52) | s(bool,'T_0') != s(bool,X49) | s(bool,X47) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f165,f231])).
fof(f1509,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,'const.bool.COND_3'(s(bool,'T_0'),s(X1,X2),s(X1,X3))) = s(X1,'const.bool.COND_3'(s(bool,X4),s(X1,X5),s(X1,X6))) | s(X1,X2) != s(X1,X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(condensation,[],[f1508])).
fof(f1508,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s(X40,'const.bool.COND_3'(s(bool,'T_0'),s(X40,X44),s(X40,X45))) = s(X40,'const.bool.COND_3'(s(bool,X41),s(X40,X42),s(X40,X43))) | s(X40,X42) != s(X40,X44) | s(bool,'T_0') != s(bool,X41) | s(bool,X39) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1398,f263])).
fof(f1398,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (~p(s(bool,X39)) | s(X40,'const.bool.COND_3'(s(bool,'T_0'),s(X40,X44),s(X40,X45))) = s(X40,'const.bool.COND_3'(s(bool,X41),s(X40,X42),s(X40,X43))) | s(X40,X42) != s(X40,X44) | s(bool,'T_0') != s(bool,X41) | s(bool,X39) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f165,f231])).
fof(f188,plain,(
  s('type.realax.real',sK0) = s('type.realax.real',sK2) | $spl0),
  inference(cnf_transformation,[],[f188_D])).
fof(f188_D,plain,(
  s('type.realax.real',sK0) = s('type.realax.real',sK2) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f190,plain,(
  s('type.realax.real',sK2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(cnf_transformation,[],[f190_D])).
fof(f190_D,plain,(
  s('type.realax.real',sK2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f192,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl4),
  inference(cnf_transformation,[],[f192_D])).
fof(f192_D,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f201,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | $spl20),
  inference(cnf_transformation,[],[f201_D])).
fof(f201_D,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f114,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X6)) | s('type.realax.real',sK2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X5)) | s('type.realax.real',sK0) = s('type.realax.real',sK2) | p(s(bool,X4))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X4] : (((p(s(bool,X4)) | s('type.realax.real',sK0) = s('type.realax.real',sK2)) & (~p(s(bool,X4)) | s('type.realax.real',sK0) != s('type.realax.real',sK2))) | ! [X5] : (((p(s(bool,X5)) | s('type.realax.real',sK2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (~p(s(bool,X5)) | s('type.realax.real',sK2) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ! [X6] : (((p(s(bool,X6)) | s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (~p(s(bool,X6)) | s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f88])).
fof(f88,plain,(
  ? [X0,X1,X2,X3] : ! [X4] : (((p(s(bool,X4)) | s('type.realax.real',X0) = s('type.realax.real',X2)) & (~p(s(bool,X4)) | s('type.realax.real',X0) != s('type.realax.real',X2))) | ! [X5] : (((p(s(bool,X5)) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (~p(s(bool,X5)) | s('type.realax.real',X2) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ! [X6] : (((p(s(bool,X6)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (~p(s(bool,X6)) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))))))))))))))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ? [X0,X1,X2,X3] : ! [X4] : ((p(s(bool,X4)) <~> s('type.realax.real',X0) = s('type.realax.real',X2)) | ! [X5] : ((p(s(bool,X5)) <~> s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) | ! [X6] : ((p(s(bool,X6)) <~> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))))))))))))))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ~! [X0,X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s('type.realax.real',X0) = s('type.realax.real',X2)) & ? [X5] : ((p(s(bool,X5)) <=> s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & ? [X6] : ((p(s(bool,X6)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))) = s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X3),s('type.realax.real',X2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X3))),s('type.realax.real',X2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))))))))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ? [X10] : ((p(s(bool,X10)) <=> s('type.realax.real',X6) = s('type.realax.real',X8)) & ? [X11] : ((p(s(bool,X11)) <=> s('type.realax.real',X8) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & ? [X12] : ((p(s(bool,X12)) <=> s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X9),s('type.realax.real',X8))))) = s('type.realax.real','const.bool.COND_3'(s(bool,X12),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X6))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X9),s('type.realax.real',X8))))),s('type.realax.real','const.bool.COND_3'(s(bool,X11),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X9),s('type.realax.real',X8))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X10),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X9))),s('type.realax.real',X8))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X8))))))))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ? [X10] : ((p(s(bool,X10)) <=> s('type.realax.real',X6) = s('type.realax.real',X8)) & ? [X11] : ((p(s(bool,X11)) <=> s('type.realax.real',X8) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & ? [X12] : ((p(s(bool,X12)) <=> s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X9),s('type.realax.real',X8))))) = s('type.realax.real','const.bool.COND_3'(s(bool,X12),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X6))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X9),s('type.realax.real',X8))))),s('type.realax.real','const.bool.COND_3'(s(bool,X11),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X9),s('type.realax.real',X8))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X10),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X9))),s('type.realax.real',X8))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X8))))))))))))))),
  file('thm.real.add_rat','thm.real.add_rat')).
fof(f189,plain,(
  s('type.realax.real',sK0) != s('type.realax.real',sK2) | $spl1),
  inference(cnf_transformation,[],[f189_D])).
fof(f189_D,plain,(
  s('type.realax.real',sK0) != s('type.realax.real',sK2) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f200,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) | $spl18),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X6))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f115,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X6)) | s('type.realax.real',sK2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X5)) | s('type.realax.real',sK0) != s('type.realax.real',sK2) | ~p(s(bool,X4))) )),
  inference(cnf_transformation,[],[f89])).
fof(f191,plain,(
  s('type.realax.real',sK2) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(cnf_transformation,[],[f191_D])).
fof(f191_D,plain,(
  s('type.realax.real',sK2) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f199,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) | $spl16),
  inference(cnf_transformation,[],[f199_D])).
fof(f199_D,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f116,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X6)) | s('type.realax.real',sK2) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X5)) | s('type.realax.real',sK0) = s('type.realax.real',sK2) | p(s(bool,X4))) )),
  inference(cnf_transformation,[],[f89])).
fof(f198,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) | $spl14),
  inference(cnf_transformation,[],[f198_D])).
fof(f198_D,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X6))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f117,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X6)) | s('type.realax.real',sK2) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X5)) | s('type.realax.real',sK0) != s('type.realax.real',sK2) | ~p(s(bool,X4))) )),
  inference(cnf_transformation,[],[f89])).
fof(f193,plain,(
  s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl5),
  inference(cnf_transformation,[],[f193_D])).
fof(f193_D,plain,(
  s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f197,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl12),
  inference(cnf_transformation,[],[f197_D])).
fof(f197_D,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f118,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X6)) | s('type.realax.real',sK2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X5)) | s('type.realax.real',sK0) = s('type.realax.real',sK2) | p(s(bool,X4))) )),
  inference(cnf_transformation,[],[f89])).
fof(f196,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl10),
  inference(cnf_transformation,[],[f196_D])).
fof(f196_D,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f119,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X6)) | s('type.realax.real',sK2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X5)) | s('type.realax.real',sK0) != s('type.realax.real',sK2) | ~p(s(bool,X4))) )),
  inference(cnf_transformation,[],[f89])).
fof(f195,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl8),
  inference(cnf_transformation,[],[f195_D])).
fof(f195_D,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f120,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X6)) | s('type.realax.real',sK2) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X5)) | s('type.realax.real',sK0) = s('type.realax.real',sK2) | p(s(bool,X4))) )),
  inference(cnf_transformation,[],[f89])).
fof(f194,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl6),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | ~p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f121,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))),s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK3),s('type.realax.real',sK2))))))),s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK1),s('type.realax.real',sK3))),s('type.realax.real',sK2))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK2))))))))))) | s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X6)) | s('type.realax.real',sK2) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X5)) | s('type.realax.real',sK0) != s('type.realax.real',sK2) | ~p(s(bool,X4))) )),
  inference(cnf_transformation,[],[f89])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 89
% Passive clauses: 413
% Generated clauses: 10407
% Final active clauses: 31
% Final passive clauses: 110
% Input formulas: 33
% Initial clauses: 72
% 
% Duplicate literals: 97
% Trivial inequalities: 15
% Fw subsumption resolutions: 117
% Fw demodulations: 5073
% Bw demodulations: 121
% Condensations: 177
% Global subsumptions: 101
% 
% Simple tautologies: 50
% Equational tautologies: 46
% Forward subsumptions: 2520
% Fw demodulations to eq. taut.: 1652
% 
% Unit resulting resolution: 2658
% Factoring: 1
% Forward superposition: 1076
% Backward superposition: 665
% Self superposition: 62
% Equality factoring: 26
% Equality resolution: 6
% 
% Split clauses: 19
% Split components: 54
% SAT solver clauses: 681
% SAT solver unit clauses: 110
% SAT solver binary clauses: 204
% 
% Sat splits: 19
% Sat splitting refutations: 127
% 
% Memory used [KB]: 5628
% Time elapsed: 0.327 s
% ------------------------------
% ------------------------------
% Success in time 1.647 s
