Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f497,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f481])).
fof(f481,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0')),
  inference(superposition,[],[f86,f446])).
fof(f446,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(forward_demodulation,[],[f422,f109])).
fof(f109,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X10] : ! [X7,X8] : s(X10,X7) = s(X10,'const.bool.COND_3'(s(bool,'T_0'),s(X10,X7),s(X10,X8)))),
  file('thm.extreal.min_infty|split|2','thm.bool.COND_CLAUSES|split|0')).
fof(f422,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'T_0'),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(superposition,[],[f120,f138])).
fof(f138,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(resolution,[],[f96,f90])).
fof(f90,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6))))),
  file('thm.extreal.min_infty|split|2','thm.extreal.le_infty|split|0')).
fof(f96,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.extreal.min_infty|split|2','thm.bool.EQ_CLAUSES|split|1')).
fof(f120,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X11] : s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X11))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X11))),s('type.extreal.extreal',X6),s('type.extreal.extreal',X11)))),
  file('thm.extreal.min_infty|split|2','thm.extreal.extreal_min_def')).
fof(f86,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK0)))),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f65])).
fof(f65,plain,(
  ? [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ~! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6)))),
  file('thm.extreal.min_infty|split|2','thm.extreal.min_infty|split|2')).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 29
% Passive clauses: 97
% Generated clauses: 401
% Final active clauses: 27
% Final passive clauses: 52
% Input formulas: 27
% Initial clauses: 46
% 
% Duplicate literals: 2
% Trivial inequalities: 5
% Fw subsumption resolutions: 14
% Fw demodulations: 60
% Bw demodulations: 2
% 
% Simple tautologies: 18
% Equational tautologies: 5
% Forward subsumptions: 199
% 
% Binary resolution: 47
% Factoring: 1
% Forward superposition: 110
% Backward superposition: 102
% Self superposition: 12
% Equality factoring: 2
% Equality resolution: 3
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 2
% SAT solver unit clauses: 1
% SAT solver binary clauses: 1
% 
% Sat splits: 1
% 
% TWLsolver clauses: 2
% TWLsolver calls for satisfiability: 2
% 
% Memory used [KB]: 895
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
% Success in time 0.021 s
