Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f17880,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f17826,f3528,f440,f763,f441,f765,f761,f17083,f16886,f442,f766,f764,f762,f544,f767,f769,f453,f770])).
fof(f770,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f770_D])).
fof(f770_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f453,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.real.eq_ints|split|1',aHL_FALSITY)).
fof(f769,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(cnf_transformation,[],[f769_D])).
fof(f769_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f767,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f767_D])).
fof(f767_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f544,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f352])).
fof(f352,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.real.eq_ints|split|1','thm.bool.IMP_CLAUSES|split|2')).
fof(f762,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) | $spl1),
  inference(cnf_transformation,[],[f762_D])).
fof(f762_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f764,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | $spl3),
  inference(cnf_transformation,[],[f764_D])).
fof(f764_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f766,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | $spl5),
  inference(cnf_transformation,[],[f766_D])).
fof(f766_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f442,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14)))))),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) & s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f316])).
fof(f316,plain,(
  ? [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ? [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))))),
  inference(nnf_transformation,[],[f256])).
fof(f256,plain,(
  ? [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) <~> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ~! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.real.eq_ints|split|1','thm.real.eq_ints|split|1')).
fof(f16886,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f16885,f764])).
fof(f16885,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | ($spl0 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f16804,f767])).
fof(f16804,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | ($spl0 | $spl4)),
  inference(superposition,[],[f504,f16786])).
fof(f16786,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f16785,f9167])).
fof(f9167,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(resolution,[],[f9162,f493])).
fof(f493,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.real.eq_ints|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f9162,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) )),
  inference(equality_resolution,[],[f598])).
fof(f598,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f371])).
fof(f371,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.realax.real',X0) != s('type.realax.real',X1) | (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))))),
  inference(flattening,[],[f370])).
fof(f370,plain,(
  ! [X0,X1] : (((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.realax.real',X0) != s('type.realax.real',X1) | (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))))),
  inference(nnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) <=> s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X13,X14] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X14)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X13))))) <=> s('type.realax.real',X13) = s('type.realax.real',X14))),
  file('thm.real.eq_ints|split|1','thm.real.REAL_LE_ANTISYM')).
fof(f16785,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f16703,f569])).
fof(f569,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X7,X6] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6)))))),
  file('thm.real.eq_ints|split|1','thm.real.REAL_LE')).
fof(f16703,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl4)),
  inference(superposition,[],[f484,f3529])).
fof(f3529,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f761,f765])).
fof(f484,plain,(
  ( ! [X0] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X13] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13)))))),
  file('thm.real.eq_ints|split|1','thm.real.REAL_NEG_GE|48|')).
fof(f504,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X6))),
  file('thm.real.eq_ints|split|1','thm.arithmetic.LESS_EQ_|48|')).
fof(f17083,plain,(
  $false | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f17082,f764])).
fof(f17082,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f17001,f767])).
fof(f17001,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | $spl0),
  inference(superposition,[],[f504,f16989])).
fof(f16989,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(forward_demodulation,[],[f16988,f462])).
fof(f462,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.real.eq_ints|split|1','thm.numeral.numeral_distrib|split|26')).
fof(f16988,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(forward_demodulation,[],[f16987,f569])).
fof(f16987,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) | $spl0),
  inference(forward_demodulation,[],[f16940,f569])).
fof(f16940,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(superposition,[],[f484,f761])).
fof(f761,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) | $spl0),
  inference(cnf_transformation,[],[f761_D])).
fof(f761_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f765,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl4),
  inference(cnf_transformation,[],[f765_D])).
fof(f765_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f441,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14)))))),
  inference(cnf_transformation,[],[f317])).
fof(f763,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | $spl2),
  inference(cnf_transformation,[],[f763_D])).
fof(f763_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f440,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14)))))),
  inference(cnf_transformation,[],[f317])).
fof(f3528,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f447,f3279])).
fof(f3279,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f3278,f765])).
fof(f3278,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f762,f763])).
fof(f447,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f100])).
fof(f100,axiom,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.real.eq_ints|split|1','thm.real.REAL_NEG_|48|')).
fof(f17826,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f17825,f766])).
fof(f17825,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | ($spl0 | $spl2 | $spl6)),
  inference(subsumption_resolution,[],[f17744,f767])).
fof(f17744,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | ($spl0 | $spl2)),
  inference(superposition,[],[f504,f17426])).
fof(f17426,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',X1)))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f17425,f462])).
fof(f17425,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',X1)))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f17377,f569])).
fof(f17377,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f569,f17348])).
fof(f17348,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f17138,f447])).
fof(f17138,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f763,f761])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 772
% Passive clauses: 3508
% Generated clauses: 17423
% Final active clauses: 697
% Final passive clauses: 2366
% Input formulas: 117
% Initial clauses: 311
% 
% Duplicate literals: 208
% Trivial inequalities: 55
% Fw subsumption resolutions: 1463
% Fw demodulations: 1267
% Bw demodulations: 66
% 
% Simple tautologies: 201
% Equational tautologies: 28
% Forward subsumptions: 10651
% 
% Binary resolution: 4276
% Factoring: 1147
% Forward superposition: 3973
% Backward superposition: 4510
% Self superposition: 20
% Equality factoring: 2
% Equality resolution: 15
% 
% Split clauses: 207
% Split components: 417
% SAT solver clauses: 370
% SAT solver unit clauses: 15
% SAT solver binary clauses: 59
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 2
% 
% Sat splits: 207
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 367
% TWLsolver calls for satisfiability: 323
% 
% Memory used [KB]: 8315
% Time elapsed: 0.350 s
% ------------------------------
% ------------------------------
% Success in time 0.357 s
