Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (56588)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 188
% Passive clauses: 2191
% Generated clauses: 23390
% Final active clauses: 175
% Final passive clauses: 1788
% Input formulas: 19
% Initial clauses: 49
% 
% Duplicate literals: 1627
% Trivial inequalities: 133
% Fw subsumption resolutions: 3555
% Fw demodulations: 3407
% Bw demodulations: 12
% 
% Simple tautologies: 357
% Equational tautologies: 21
% Forward subsumptions: 9928
% 
% Binary resolution: 663
% Forward superposition: 8710
% Backward superposition: 4847
% Self superposition: 113
% Equality factoring: 34
% Equality resolution: 31
% 
% Split clauses: 1064
% Split components: 2134
% SAT solver clauses: 1914
% SAT solver unit clauses: 13
% SAT solver binary clauses: 97
% 
% Sat splits: 1064
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1913
% TWLsolver calls for satisfiability: 249
% 
% Memory used [KB]: 11897
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (56601)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 190
% Passive clauses: 1484
% Generated clauses: 11514
% Final active clauses: 136
% Final passive clauses: 1018
% Input formulas: 19
% Initial clauses: 49
% 
% Duplicate literals: 709
% Trivial inequalities: 125
% Fw subsumption resolutions: 768
% Fw demodulations: 988
% Bw demodulations: 259
% Condensations: 668
% Global subsumptions: 4
% 
% Simple tautologies: 145
% Equational tautologies: 30
% Forward subsumptions: 4746
% Bw demodulations to eq. taut.: 69
% 
% Binary resolution: 507
% Forward superposition: 3879
% Backward superposition: 1940
% Self superposition: 72
% Equality factoring: 44
% Equality resolution: 9
% 
% Split clauses: 122
% Split components: 249
% SAT solver clauses: 1867
% SAT solver unit clauses: 53
% SAT solver binary clauses: 39
% SAT solver learnt clauses: 8
% 
% Sat splits: 122
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1851
% TWLsolver calls for satisfiability: 412
% 
% Memory used [KB]: 13176
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (56619)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 856
% Passive clauses: 6879
% Generated clauses: 130593
% Final active clauses: 85
% Final passive clauses: 477
% Input formulas: 19
% Initial clauses: 49
% 
% Duplicate literals: 8349
% Trivial inequalities: 556
% Fw subsumption resolutions: 11935
% Fw demodulations: 17011
% Bw demodulations: 5257
% Condensations: 11464
% Global subsumptions: 1422
% 
% Simple tautologies: 762
% Equational tautologies: 1519
% Forward subsumptions: 68755
% Fw demodulations to eq. taut.: 42
% Bw demodulations to eq. taut.: 216
% 
% Unit resulting resolution: 6437
% Forward superposition: 39610
% Backward superposition: 16604
% Self superposition: 423
% Equality factoring: 160
% Equality resolution: 26
% 
% Split clauses: 452
% Split components: 910
% SAT solver clauses: 10763
% SAT solver unit clauses: 23
% SAT solver binary clauses: 772
% 
% Sat splits: 452
% Sat splitting refutations: 1556
% 
% Memory used [KB]: 64092
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 183 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (56919)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1888
% Passive clauses: 43635
% Generated clauses: 101379
% Final active clauses: 440
% Final passive clauses: 27882
% Input formulas: 19
% Initial clauses: 49
% 
% Selected by SInE selection: 19
% SInE iterations: 1
% 
% Duplicate literals: 991
% Trivial inequalities: 2466
% Fw subsumption resolutions: 8221
% Fw demodulations: 8664
% Bw demodulations: 1011
% 
% Simple tautologies: 1532
% Equational tautologies: 114
% Forward subsumptions: 36951
% Bw demodulations to eq. taut.: 117
% 
% Binary resolution: 22822
% Factoring: 20
% Forward superposition: 42944
% Backward superposition: 11453
% Self superposition: 57
% Equality factoring: 1
% Equality resolution: 78
% 
% Split clauses: 617
% Split components: 1241
% SAT solver clauses: 790
% SAT solver unit clauses: 45
% SAT solver binary clauses: 290
% 
% Sat splits: 617
% Sat splitting refutations: 172
% 
% Memory used [KB]: 40297
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (56992)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 82
% Passive clauses: 1608
% Generated clauses: 9753
% Final active clauses: 63
% Final passive clauses: 1172
% Input formulas: 19
% Initial clauses: 49
% 
% Duplicate literals: 449
% Trivial inequalities: 277
% Fw subsumption resolutions: 1313
% Fw demodulations: 146
% Bw demodulations: 314
% Global subsumptions: 16
% 
% Simple tautologies: 476
% Equational tautologies: 73
% Forward subsumptions: 5158
% Bw demodulations to eq. taut.: 30
% 
% Unit resulting resolution: 350
% Forward superposition: 3401
% Backward superposition: 2849
% Self superposition: 82
% Equality factoring: 58
% Equality resolution: 11
% 
% Split clauses: 13
% Split components: 26
% SAT solver clauses: 1715
% SAT solver unit clauses: 85
% SAT solver binary clauses: 62
% 
% Sat splits: 13
% Sat splitting refutations: 76
% 
% Memory used [KB]: 7931
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (57021)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 590
% Passive clauses: 4701
% Generated clauses: 15855
% Final active clauses: 575
% Final passive clauses: 4046
% Input formulas: 19
% Initial clauses: 49
% 
% Duplicate literals: 283
% Trivial inequalities: 681
% Fw subsumption resolutions: 334
% Bw subsumption resolutions: 5
% Fw demodulations: 969
% Bw demodulations: 29
% Condensations: 28
% 
% Simple tautologies: 330
% Equational tautologies: 3
% Forward subsumptions: 8559
% Backward subsumptions: 76
% 
% Binary resolution: 2947
% Factoring: 10
% Forward superposition: 7702
% Backward superposition: 2879
% Equality factoring: 3
% Equality resolution: 6
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 4706
% SAT solver unit clauses: 9
% SAT solver binary clauses: 138
% 
% Sat splits: 3
% 
% Memory used [KB]: 24306
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 120 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% (57156)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 185
% Passive clauses: 1713
% Generated clauses: 4036
% Final active clauses: 165
% Final passive clauses: 1047
% Input formulas: 19
% Initial clauses: 49
% 
% Duplicate literals: 180
% Trivial inequalities: 119
% Fw demodulations: 52
% Bw demodulations: 3
% Global subsumptions: 4
% 
% Simple tautologies: 47
% Equational tautologies: 72
% Forward subsumptions: 1792
% Backward subsumptions: 824
% 
% Binary resolution: 940
% Unit resulting resolution: 29
% Forward superposition: 1099
% Backward superposition: 1537
% Self superposition: 14
% Equality factoring: 4
% Equality resolution: 6
% 
% Split clauses: 65
% Split components: 132
% SAT solver clauses: 1892
% SAT solver unit clauses: 4
% SAT solver binary clauses: 44
% 
% Sat splits: 65
% 
% Memory used [KB]: 22003
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% remaining time: 110 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f54605,plain,(
  $false),
  inference(global_subsumption,[],[f54604,f15523,f14785])).
fof(f14785,plain,(
  s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) != s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0))),
  inference(trivial_inequality_removal,[],[f14747])).
fof(f14747,plain,(
  s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) != s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0))))),
  inference(resolution,[],[f14686,f1398])).
fof(f1398,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),sK3)))) | s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) != s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X0))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0))))) )),
  inference(resolution,[],[f179,f120])).
fof(f120,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | s(X0,X2) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) | (s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,sK9(X4,X3,X2,X1,X0)))) & p(s(bool,'const.bool.IN_2'(s(X1,sK9(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3)))))) & (! [X5] : (s(X0,X2) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) | ? [X6] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))))) & (! [X5] : (s(X0,X2) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3))))))))),
  inference(rectify,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) | ? [X5] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) & p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))))) & (! [X5] : (s(X0,X2) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3))))))))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) <=> ? [X5] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) & p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3))))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) <=> ? [X5] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) & p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6,X7] : ! [X12,X9,X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X10),s(fun(X7,bool),X9)))))) <=> ? [X11] : (s(X6,X12) = s(X6,happ(s(fun(X7,X6),X10),s(X7,X11))) & p(s(bool,'const.bool.IN_2'(s(X7,X11),s(fun(X7,bool),X9))))))),
  file('thm.pred_set.pred_set_|49|76_prove_ltb','thm.pred_set.IN_IMAGE')).
fof(f179,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)))))) | s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) != s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0))),
  inference(resolution,[],[f75,f117])).
fof(f117,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | s(X1,sK7(X4,X3,X2,X1,X0)) != s(X1,sK8(X4,X3,X2,X1,X0)) | ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,sK6(X4,X3,X2,X1,X0)))),s(fun(X0,bool),X4))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X1,X10),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X10))),s(fun(X0,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X9),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X9))) | s(X1,X8) = s(X1,X9)))) & ((p(s(bool,'const.bool.IN_2'(s(X1,sK6(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,sK6(X4,X3,X2,X1,X0)))),s(fun(X0,bool),X4))))) | (p(s(bool,'const.bool.IN_2'(s(X1,sK7(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,sK8(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,sK7(X4,X3,X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK8(X4,X3,X2,X1,X0)))) & s(X1,sK7(X4,X3,X2,X1,X0)) != s(X1,sK8(X4,X3,X2,X1,X0))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8])],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X1,X10),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X10))),s(fun(X0,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X9),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X9))) | s(X1,X8) = s(X1,X9)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6)) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))))),
  inference(rectify,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6)) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & ((? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) | (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) => (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) => s(X1,X5) = s(X1,X6)))))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) => (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) => s(X1,X5) = s(X1,X6)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X7] : ! [X10,X9,X8] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X10),s(fun(X7,bool),X9),s(fun(X6,bool),X8)))) <=> (! [X11] : (p(s(bool,'const.bool.IN_2'(s(X7,X11),s(fun(X7,bool),X9)))) => p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(X7,X6),X10),s(X7,X11))),s(fun(X6,bool),X8))))) & ! [X11,X12] : ((p(s(bool,'const.bool.IN_2'(s(X7,X11),s(fun(X7,bool),X9)))) & p(s(bool,'const.bool.IN_2'(s(X7,X12),s(fun(X7,bool),X9))))) => (s(X6,happ(s(fun(X7,X6),X10),s(X7,X11))) = s(X6,happ(s(fun(X7,X6),X10),s(X7,X12))) => s(X7,X11) = s(X7,X12)))))),
  file('thm.pred_set.pred_set_|49|76_prove_ltb','thm.pred_set.INJ_DEF')).
fof(f75,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3))))))),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3),s(fun(sK0,bool),sK2)))) & ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f44])).
fof(f44,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),X2)))) & ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),X2)))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),X2)))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X4),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X10),s(fun(X7,bool),X9),s(fun(X6,bool),X8)))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X10),s(fun(X7,bool),X9),s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X10),s(fun(X7,bool),X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X10),s(fun(X7,bool),X9),s(fun(X6,bool),X8)))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X10),s(fun(X7,bool),X9),s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X10),s(fun(X7,bool),X9)))))))),
  file('thm.pred_set.pred_set_|49|76_prove_ltb','thm.pred_set.pred_set_|49|76_prove_ltb')).
fof(f14686,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3))))),
  inference(global_subsumption,[],[f14683,f178,f176])).
fof(f176,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3)))) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) = s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0))))),
  inference(resolution,[],[f75,f114])).
fof(f114,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,sK7(X4,X3,X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK8(X4,X3,X2,X1,X0)))) | p(s(bool,'const.bool.IN_2'(s(X1,sK6(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f178,plain,(
  s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) != s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3))))),
  inference(resolution,[],[f75,f116])).
fof(f116,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | s(X1,sK7(X4,X3,X2,X1,X0)) != s(X1,sK8(X4,X3,X2,X1,X0)) | p(s(bool,'const.bool.IN_2'(s(X1,sK6(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f14683,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) | s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) = s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3))))),
  inference(duplicate_literal_removal,[],[f14646])).
fof(f14646,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) | s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) = s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3))))),
  inference(resolution,[],[f881,f174])).
fof(f174,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3))))),
  inference(resolution,[],[f75,f112])).
fof(f112,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X1,sK8(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,sK6(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f881,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),sK3)))) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X0))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) | s(sK1,X0) = s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3))))) )),
  inference(resolution,[],[f172,f124])).
fof(f124,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X2),s(fun(sK1,bool),sK3)))) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X1))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X2))) | s(sK1,X1) = s(sK1,X2) | ~p(s(bool,'const.bool.IN_2'(s(sK1,X1),s(fun(sK1,bool),sK3))))) )),
  inference(resolution,[],[f74,f109])).
fof(f109,plain,(
  ( ! [X4,X2,X0,X8,X3,X1,X9] : (s(X1,X8) = s(X1,X9) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X9))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X9),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f74,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f48])).
fof(f172,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3))))),
  inference(resolution,[],[f75,f110])).
fof(f110,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X1,sK7(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,sK6(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f15523,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3)))),
  inference(resolution,[],[f14787,f81])).
fof(f81,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.pred_set_|49|76_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f14787,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3))))),
  inference(trivial_inequality_removal,[],[f14745])).
fof(f14745,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3)))) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0))))),
  inference(resolution,[],[f14686,f1241])).
fof(f1241,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3)))) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X0))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0))))) )),
  inference(resolution,[],[f173,f120])).
fof(f173,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3))))),
  inference(resolution,[],[f75,f111])).
fof(f111,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X1,sK7(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,sK6(X4,X3,X2,X1,X0)))),s(fun(X0,bool),X4))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f54604,plain,(
  s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) = s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3)))),
  inference(equality_resolution,[],[f18851])).
fof(f18851,plain,(
  ( ! [X3] : (s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X3))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) | s(sK1,X3) = s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,X3),s(fun(sK1,bool),sK3)))) )),
  inference(global_subsumption,[],[f18840,f151])).
fof(f151,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(backward_demodulation,[],[f129,f127])).
fof(f127,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3),s(fun(sK0,bool),sK2)))),
  inference(resolution,[],[f74,f86])).
fof(f86,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.pred_set_|49|76_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f129,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3),s(fun(sK0,bool),sK2)))),
  inference(resolution,[],[f74,f83])).
fof(f83,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.pred_set_|49|76_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f18840,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'F_0') | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X3))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) | s(sK1,X3) = s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,X3),s(fun(sK1,bool),sK3)))) )),
  inference(backward_demodulation,[],[f18677,f18739])).
fof(f18739,plain,(
  ( ! [X3] : (s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X3))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3))))) | s(sK1,X3) = s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,X3),s(fun(sK1,bool),sK3)))) )),
  inference(backward_demodulation,[],[f18646,f7013])).
fof(f7013,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3))))) | s(sK1,X3) = s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X3))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,X3),s(fun(sK1,bool),sK3)))) )),
  inference(resolution,[],[f1290,f852])).
fof(f852,plain,(
  ( ! [X50,X49] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X49),s(fun(sK1,bool),sK3)))) | s(sK1,X49) = s(sK1,X50) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X49))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X50))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,X50),s(fun(sK1,bool),sK3)))) )),
  inference(resolution,[],[f124,f82])).
fof(f82,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f1290,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)))))),
  inference(resolution,[],[f175,f87])).
fof(f87,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f52])).
fof(f175,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)),s(fun(sK1,bool),sK3))))),
  inference(resolution,[],[f75,f113])).
fof(f113,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X1,sK8(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,sK6(X4,X3,X2,X1,X0)))),s(fun(X0,bool),X4))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f18646,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) = s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0))))),
  inference(resolution,[],[f18525,f177])).
fof(f177,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)))))) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK7('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) = s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK8('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0))))),
  inference(resolution,[],[f75,f115])).
fof(f115,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,sK7(X4,X3,X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK8(X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,sK6(X4,X3,X2,X1,X0)))),s(fun(X0,bool),X4))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f18525,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(sK1,X0),X1),s(fun(sK1,bool),sK3))))))) )),
  inference(equality_resolution,[],[f14754])).
fof(f14754,plain,(
  ( ! [X43,X44,X42] : (s(X42,X43) != s(X42,happ(s(fun(sK1,X42),X44),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))) | p(s(bool,'const.bool.IN_2'(s(X42,X43),s(fun(X42,bool),'const.pred_set.IMAGE_2'(s(fun(sK1,X42),X44),s(fun(sK1,bool),sK3))))))) )),
  inference(resolution,[],[f14686,f120])).
fof(f18677,plain,(
  ( ! [X103,X102] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X102,happ(s(fun(sK1,X102),X103),s(sK1,sK6('const.pred_set.IMAGE_2'(s(fun(sK1,sK0),sK4),s(fun(sK1,bool),sK3)),sK3,sK4,sK1,sK0)))),s(fun(X102,bool),'const.pred_set.IMAGE_2'(s(fun(sK1,X102),X103),s(fun(sK1,bool),sK3)))))) )),
  inference(resolution,[],[f18525,f81])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1001
% Passive clauses: 18924
% Generated clauses: 48931
% Final active clauses: 761
% Final passive clauses: 6684
% Input formulas: 19
% Initial clauses: 47
% Discarded non-redundant clauses: 20313
% 
% Duplicate literals: 1439
% Trivial inequalities: 1550
% Fw demodulations: 1098
% Bw demodulations: 816
% Condensations: 323
% Global subsumptions: 75
% 
% Simple tautologies: 656
% Equational tautologies: 142
% Forward subsumptions: 24681
% Bw demodulations to eq. taut.: 310
% 
% Binary resolution: 15146
% Factoring: 15
% Forward superposition: 22666
% Backward superposition: 5817
% Self superposition: 2
% Equality factoring: 17
% Equality resolution: 22
% 
% SAT solver clauses: 19287
% SAT solver unit clauses: 248
% SAT solver binary clauses: 558
% SAT solver learnt clauses: 151
% SAT solver learnt literals: 290
% 
% TWLsolver clauses: 18999
% TWLsolver calls for satisfiability: 847254
% 
% Memory used [KB]: 80979
% Time elapsed: 8.569 s
% ------------------------------
% ------------------------------
% Success in time 27.576 s
