Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (59003)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 944
% Passive clauses: 8882
% Generated clauses: 23907
% Final active clauses: 938
% Final passive clauses: 7664
% Input formulas: 46
% Initial clauses: 185
% 
% Duplicate literals: 331
% Trivial inequalities: 90
% Fw subsumption resolutions: 2950
% Fw demodulations: 1588
% Bw demodulations: 5
% 
% Simple tautologies: 231
% Equational tautologies: 35
% Forward subsumptions: 9077
% 
% Binary resolution: 8918
% Factoring: 615
% Forward superposition: 2840
% Backward superposition: 6131
% Self superposition: 28
% Equality factoring: 26
% Equality resolution: 7
% 
% Split clauses: 705
% Split components: 1416
% SAT solver clauses: 1003
% SAT solver unit clauses: 23
% SAT solver binary clauses: 208
% 
% Sat splits: 705
% 
% TWLsolver clauses: 1003
% TWLsolver calls for satisfiability: 424
% 
% Memory used [KB]: 12409
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (59017)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 440
% Passive clauses: 2601
% Generated clauses: 14105
% Final active clauses: 434
% Final passive clauses: 2158
% Input formulas: 46
% Initial clauses: 185
% 
% Duplicate literals: 263
% Trivial inequalities: 77
% Fw subsumption resolutions: 1160
% Fw demodulations: 1202
% Bw demodulations: 7
% Condensations: 422
% Global subsumptions: 13
% 
% Simple tautologies: 282
% Equational tautologies: 19
% Forward subsumptions: 4824
% 
% Binary resolution: 3290
% Factoring: 164
% Forward superposition: 2461
% Backward superposition: 2118
% Self superposition: 39
% Equality factoring: 30
% Equality resolution: 8
% 
% Split clauses: 71
% Split components: 142
% SAT solver clauses: 3381
% SAT solver unit clauses: 191
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 71
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 3329
% TWLsolver calls for satisfiability: 1809
% 
% Memory used [KB]: 13432
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (59051)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2300
% Passive clauses: 11409
% Generated clauses: 121067
% Final active clauses: 1712
% Final passive clauses: 3592
% Input formulas: 46
% Initial clauses: 185
% Discarded non-redundant clauses: 2313
% 
% Duplicate literals: 5702
% Trivial inequalities: 125
% Fw subsumption resolutions: 5643
% Fw demodulations: 17852
% Bw demodulations: 3637
% Condensations: 9874
% Global subsumptions: 368
% 
% Simple tautologies: 162
% Equational tautologies: 598
% Forward subsumptions: 65935
% Bw demodulations to eq. taut.: 160
% 
% Unit resulting resolution: 20571
% Factoring: 7
% Forward superposition: 28213
% Backward superposition: 26375
% Self superposition: 186
% Equality factoring: 144
% Equality resolution: 4
% 
% Split clauses: 551
% Split components: 1125
% SAT solver clauses: 13919
% SAT solver unit clauses: 1643
% SAT solver binary clauses: 3629
% 
% Sat splits: 551
% Sat splitting refutations: 1037
% 
% Memory used [KB]: 85201
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% (59313)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2092
% Passive clauses: 63228
% Generated clauses: 184416
% Final active clauses: 2050
% Final passive clauses: 59339
% Input formulas: 46
% Initial clauses: 181
% 
% Selected by SInE selection: 42
% SInE iterations: 1
% 
% Duplicate literals: 4754
% Trivial inequalities: 240
% Fw subsumption resolutions: 17144
% Fw demodulations: 2683
% Bw demodulations: 29
% 
% Simple tautologies: 4108
% Equational tautologies: 260
% Forward subsumptions: 84025
% 
% Binary resolution: 142127
% Factoring: 484
% Forward superposition: 13493
% Backward superposition: 2078
% Self superposition: 6
% Equality factoring: 2
% Equality resolution: 21
% 
% Split clauses: 9602
% Split components: 22090
% SAT solver clauses: 10287
% SAT solver unit clauses: 28
% SAT solver binary clauses: 804
% 
% Sat splits: 9602
% Sat splitting refutations: 38
% 
% Memory used [KB]: 63836
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.relation
% (59399)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 547
% Passive clauses: 4849
% Generated clauses: 22247
% Final active clauses: 533
% Final passive clauses: 4102
% Input formulas: 46
% Initial clauses: 185
% 
% Duplicate literals: 358
% Trivial inequalities: 81
% Fw subsumption resolutions: 1179
% Fw demodulations: 2490
% Bw demodulations: 67
% Global subsumptions: 5
% 
% Simple tautologies: 383
% Equational tautologies: 75
% Forward subsumptions: 12224
% 
% Unit resulting resolution: 9694
% Factoring: 2
% Forward superposition: 3754
% Backward superposition: 4333
% Self superposition: 67
% Equality factoring: 6
% Equality resolution: 9
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 5447
% SAT solver unit clauses: 761
% SAT solver binary clauses: 1549
% 
% Sat splits: 15
% Sat splitting refutations: 576
% 
% Memory used [KB]: 13816
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.relation
% (59427)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1406
% Passive clauses: 7108
% Generated clauses: 23230
% Final active clauses: 684
% Final passive clauses: 3262
% Input formulas: 46
% Initial clauses: 185
% 
% Duplicate literals: 994
% Trivial inequalities: 69
% Fw subsumption resolutions: 849
% Bw subsumption resolutions: 1546
% Fw demodulations: 971
% Bw demodulations: 61
% Condensations: 1669
% Global subsumptions: 406
% 
% Simple tautologies: 1474
% Equational tautologies: 48
% Forward subsumptions: 9059
% Backward subsumptions: 1629
% 
% Binary resolution: 7420
% Factoring: 61
% Forward superposition: 2879
% Backward superposition: 3136
% Self superposition: 75
% Equality factoring: 4
% Equality resolution: 13
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 8039
% SAT solver unit clauses: 672
% SAT solver binary clauses: 1569
% 
% Sat splits: 20
% Sat splitting refutations: 502
% 
% Memory used [KB]: 21748
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f4102,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f260,f450,f259,f452,f649,f648,f647,f646,f685,f645,f683,f644,f643,f642,f687,f632,f689,f630,f629,f628,f627,f693,f626,f691,f625,f624,f623,f695,f619,f697,f600,f699,f597,f596,f595,f594,f703,f593,f701,f592,f591,f590,f705,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f713,f571,f570,f569,f568,f709,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f717,f555,f707,f554,f553,f552,f715,f551,f550,f549,f548,f547,f546,f545,f544,f543,f711,f542,f541,f540,f719,f538,f537,f536,f535,f723,f534,f721,f533,f532,f531,f725,f518,f517,f516,f515,f729,f514,f727,f513,f512,f511,f731,f509,f508,f507,f506,f735,f505,f733,f504,f503,f502,f737,f501,f500,f499,f498,f741,f497,f739,f496,f495,f494,f743,f493,f492,f491,f490,f489,f488,f487,f486,f485,f484,f483,f753,f482,f751,f481,f480,f479,f478,f749,f477,f745,f476,f475,f474,f747,f468,f467,f466,f464,f463,f462,f458,f457,f757,f456,f455,f651,f652,f758,f755,f759,f763,f779,f778,f777,f785,f776,f787,f770,f769,f768,f789,f767,f791,f1014,f1013,f1012,f1011,f1010,f1009,f1008,f1007,f1006,f1048,f1002,f1001,f1000,f999,f998,f997,f1052,f996,f995,f994,f993,f992,f1050,f991,f1054,f988,f987,f986,f985,f984,f983,f982,f981,f980,f1056,f966,f965,f964,f963,f962,f961,f1060,f960,f959,f958,f957,f956,f1058,f955,f1062,f954,f953,f952,f951,f1070,f950,f949,f948,f1066,f947,f1074,f946,f1064,f945,f1072,f944,f1068,f943,f1076,f941,f940,f939,f938,f937,f936,f1080,f935,f934,f933,f1078,f932,f931,f930,f1082,f923,f922,f921,f920,f919,f918,f917,f916,f915,f1084,f898,f897,f896,f895,f894,f893,f892,f891,f890,f1086,f877,f1088,f876,f1090,f875,f874,f873,f872,f1098,f871,f870,f1094,f869,f868,f1102,f867,f1092,f866,f1096,f865,f1100,f864,f1104,f861,f860,f859,f858,f857,f856,f855,f854,f853,f1106,f850,f849,f848,f847,f846,f845,f844,f843,f842,f1108,f841,f840,f839,f838,f837,f836,f835,f834,f833,f1110,f832,f831,f830,f829,f828,f827,f826,f825,f824,f823,f822,f821,f820,f1120,f819,f818,f817,f1118,f816,f815,f1116,f814,f1112,f813,f1114,f812,f811,f810,f809,f808,f807,f1124,f806,f1122,f805,f1126,f803,f801,f796,f795,f794,f793,f1130,f1131,f1133,f1022,f1023,f1128,f1203,f1202,f1201,f1206,f1289,f1288,f1287,f1292,f1301,f1293,f1299,f1294,f1297,f1494,f1500,f1502,f1519,f1518,f1517,f1521,f1531,f1522,f1529,f1523,f1527,f1582,f1588,f1593,f1599,f1605,f1610,f1638,f1642,f1648,f1652,f1658,f1662,f1668,f1672,f1758,f1757,f1756,f1844,f1754,f1753,f1752,f1846,f1824,f1853,f1825,f1855,f2059,f2065,f2072,f2092,f2096,f2104,f2108,f2120,f2126,f2144,f2147,f2145,f2148,f2152,f2151,f2155,f2404,f2410,f2412,f2456,f2462,f2464,f2473,f2479,f2481,f2518,f2520,f2525,f2527,f2532,f2534,f2539,f2543,f2548,f2552,f2611,f2617,f2623,f2629,f2639,f2638,f2637,f2722,f2635,f2634,f2633,f2724,f2703,f2732,f2704,f2734,f2705,f2736,f2706,f2738,f2749,f2759,f2744,f2762,f3297,f3309,f3295,f3311,f3483,f3489,f3492,f3511,f3513,f3519,f3521,f3532,f3534,f3890,f3889,f3887,f3886,f3885,f3884,f3926,f3883,f3924,f3882,f3928,f3881,f3880,f3932,f3879,f3930,f3878,f3934,f3872,f3871,f3938,f3870,f3936,f3869,f3940,f3868,f3942,f3867,f3944,f3853,f3852,f3851,f3850,f3849,f3848,f3847,f3846,f3845,f3949,f3839,f3838,f3837,f3836,f3835,f3834,f3954,f3832,f3831,f3830,f3829,f3828,f3952,f3827,f3957,f3822,f3821,f3820,f3819,f3818,f3817,f3816,f3815,f3814,f3961,f3810,f3809,f3808,f3965,f3784,f3783,f3782,f3781,f3780,f3779,f3974,f3777,f3776,f3775,f3774,f3773,f3972,f3772,f3977,f3771,f3770,f3769,f3768,f3985,f3766,f3765,f3764,f3981,f3763,f3990,f3761,f3979,f3760,f3988,f3759,f3983,f3758,f3993,f3754,f3753,f3752,f3751,f3750,f3749,f3999,f3746,f3745,f3744,f3997,f3743,f3742,f3741,f4003,f3738,f3737,f3920,f3734,f3918,f3733,f3922,f3716,f3715,f3714,f3713,f3712,f3711,f3710,f3709,f3708,f4009,f3679,f3678,f3677,f3676,f3675,f3674,f3673,f3672,f3671,f4019,f3650,f4025,f3649,f4027,f3645,f4029,f3644,f4031,f3642,f4033,f3641,f4035,f3640,f4037,f3639,f4039,f3636,f3635,f3634,f3633,f4049,f3630,f3629,f4045,f3628,f3627,f4055,f3624,f4043,f3623,f4047,f3622,f4053,f3621,f4059,f3616,f3615,f3614,f3613,f3612,f3611,f3610,f3609,f3608,f4063,f3603,f3602,f3601,f3600,f3599,f3598,f3597,f3596,f3595,f4067,f3594,f3593,f4069,f3592,f3591,f4071,f3590,f3589,f3588,f3587,f3586,f3585,f3584,f3583,f3582,f4073,f3581,f3580,f3579,f3578,f3577,f3576,f3575,f3574,f3573,f3572,f3571,f3570,f3569,f4083,f3568,f3567,f3566,f4081,f3565,f3564,f4079,f3563,f4075,f3562,f4077,f3560,f3559,f3557,f3556,f3554,f3553,f4088,f3551,f4086,f3550,f4091,f3548,f3546,f3541,f3540,f3539,f3538,f4097,f4098,f4100,f3892,f3893,f4095,f4101])).
fof(f4101,plain,(
  s(bool,'T_0') != s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ~p(s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | p(s(bool,'T_0'))),
  introduced(theory_axiom,[])).
fof(f4095,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl400),
  inference(cnf_transformation,[],[f4095_D])).
fof(f4095_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f3893,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(resolution,[],[f1133,f284])).
fof(f284,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.relation.reflexive_TC','thm.bool.EQ_CLAUSES|split|2')).
fof(f3892,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(resolution,[],[f1133,f282])).
fof(f282,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.relation.reflexive_TC','thm.bool.EQ_CLAUSES|split|3')).
fof(f4100,plain,(
  ~p(s(bool,'T_0')) | $spl405),
  inference(cnf_transformation,[],[f4100_D])).
fof(f4100_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl405),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl405])])).
fof(f4098,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl175)),
  inference(forward_demodulation,[],[f3537,f2794])).
fof(f2794,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X0))),s(sK12,X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f453,f337])).
fof(f337,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.relation.reflexive_TC','thm.bool.IMP_ANTISYM_AX')).
fof(f453,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X0))),s(sK12,X0))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f308])).
fof(f308,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | ~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK14(X1,X0)))),s(X0,sK14(X1,X0))))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X6,fun(X6,bool)),X7)))) <=> ! [X11] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))),s(X6,X11)))))),
  file('thm.relation.reflexive_TC','thm.relation.reflexive_def')).
fof(f261,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.reflexive_TC','thm.bool.TRUTH')).
fof(f3537,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f333])).
fof(f333,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X7,X11,X12] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))),s(X6,X12)))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X11),s(X6,X12)))))),
  file('thm.relation.reflexive_TC','thm.relation.TC_RULES|split|0')).
fof(f4097,plain,(
  s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl403),
  inference(cnf_transformation,[],[f4097_D])).
fof(f4097_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl403),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl403])])).
fof(f3538,plain,(
  s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f277])).
fof(f277,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.relation.reflexive_TC','thm.bool.EQ_CLAUSES|split|1')).
fof(f3539,plain,(
  s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f279])).
fof(f279,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.relation.reflexive_TC','thm.bool.EQ_CLAUSES|split|0')).
fof(f3540,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f282])).
fof(f3541,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f284])).
fof(f3546,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f1133,f334])).
fof(f334,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f138])).
fof(f267,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.reflexive_TC',aHL_FALSITY)).
fof(f3548,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f1133,f334])).
fof(f4091,plain,(
  ~sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl399),
  inference(cnf_transformation,[],[f4091_D])).
fof(f4091_D,plain,(
  ~sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') <=> ~$spl399),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl399])])).
fof(f3550,plain,(
  ~sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f267,f1133,f342])).
fof(f342,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f4086,plain,(
  ~sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl395),
  inference(cnf_transformation,[],[f4086_D])).
fof(f4086_D,plain,(
  ~sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl395),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl395])])).
fof(f3551,plain,(
  ~sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f267,f1133,f342])).
fof(f4088,plain,(
  ~sP0('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl397),
  inference(cnf_transformation,[],[f4088_D])).
fof(f4088_D,plain,(
  ~sP0('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl397),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl397])])).
fof(f3553,plain,(
  ~sP0('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f267,f1133,f342])).
fof(f3554,plain,(
  ~sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f267,f1133,f342])).
fof(f3556,plain,(
  ~sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f267,f1133,f346])).
fof(f346,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f140,f145])).
fof(f140,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X15,X6] : ((~(p(s(bool,X6)) | p(s(bool,X15))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X15)) => p(s(bool,'F_0')))))),
  file('thm.relation.reflexive_TC','thm.sat.sat_9_prove_ltb')).
fof(f3557,plain,(
  ~sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f267,f1133,f346])).
fof(f3559,plain,(
  ~sP0('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f267,f1133,f346])).
fof(f3560,plain,(
  ~sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f267,f1133,f346])).
fof(f4077,plain,(
  ~sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl387),
  inference(cnf_transformation,[],[f4077_D])).
fof(f4077_D,plain,(
  ~sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') <=> ~$spl387),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl387])])).
fof(f3562,plain,(
  ~sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f1133,f357])).
fof(f357,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f217])).
fof(f217,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f4075,plain,(
  ~sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl385),
  inference(cnf_transformation,[],[f4075_D])).
fof(f4075_D,plain,(
  ~sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl385),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl385])])).
fof(f3563,plain,(
  ~sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f1133,f357])).
fof(f4079,plain,(
  ~sP1('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl389),
  inference(cnf_transformation,[],[f4079_D])).
fof(f4079_D,plain,(
  ~sP1('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl389),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl389])])).
fof(f3564,plain,(
  ~sP1('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f1133,f357])).
fof(f3565,plain,(
  ~sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f1133,f357])).
fof(f4081,plain,(
  sP1('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl390),
  inference(cnf_transformation,[],[f4081_D])).
fof(f4081_D,plain,(
  sP1('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f3566,plain,(
  sP1('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f1133,f359])).
fof(f359,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f445,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f306])).
fof(f306,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X11] : (s(X6,X11) = s(X6,X11) <=> p(s(bool,'T_0')))),
  file('thm.relation.reflexive_TC','thm.bool.REFL_CLAUSE')).
fof(f3567,plain,(
  sP1('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f1133,f359])).
fof(f262,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.reflexive_TC',aHL_TRUTH)).
fof(f3568,plain,(
  sP1('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f1133,f359])).
fof(f4083,plain,(
  sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl392),
  inference(cnf_transformation,[],[f4083_D])).
fof(f4083_D,plain,(
  sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f3569,plain,(
  sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f1133,f362])).
fof(f362,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f3570,plain,(
  sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f1133,f362])).
fof(f3571,plain,(
  sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f1133,f362])).
fof(f3572,plain,(
  sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f1133,f363])).
fof(f363,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X1,X0) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP1(X1,X0)) & (~sP1(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP1(X1,X0))),
  inference(definition_folding,[],[f109,f147])).
fof(f109,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X18,X19] : ((p(s(bool,X19)) <=> ~p(s(bool,X18))) <=> ((p(s(bool,X19)) | p(s(bool,X18))) & (~p(s(bool,X18)) | ~p(s(bool,X19)))))),
  file('thm.relation.reflexive_TC','thm.sat.sat_29_prove_ltb')).
fof(f3573,plain,(
  sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f1133,f363])).
fof(f3574,plain,(
  sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f1133,f363])).
fof(f3575,plain,(
  sP1('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f1133,f364])).
fof(f364,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f219])).
fof(f3576,plain,(
  sP1('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f1133,f364])).
fof(f3577,plain,(
  sP1('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f1133,f364])).
fof(f3578,plain,(
  ~sP1('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f1133,f366])).
fof(f366,plain,(
  ( ! [X0,X1] : (~sP1(X1,X0) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f219])).
fof(f3579,plain,(
  ~sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f1133,f366])).
fof(f3580,plain,(
  ~sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f1133,f366])).
fof(f3581,plain,(
  ~sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f1133,f366])).
fof(f4073,plain,(
  ~sP5('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl383),
  inference(cnf_transformation,[],[f4073_D])).
fof(f4073_D,plain,(
  ~sP5('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl383),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl383])])).
fof(f3582,plain,(
  ~sP5('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f445,f1133,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f225])).
fof(f225,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f3583,plain,(
  ~sP5('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f445,f1133,f373])).
fof(f3584,plain,(
  ~sP5('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f445,f1133,f373])).
fof(f3585,plain,(
  ~sP5('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f262,f1133,f373])).
fof(f3586,plain,(
  ~sP5('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f262,f1133,f373])).
fof(f3587,plain,(
  ~sP5('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f262,f1133,f373])).
fof(f3588,plain,(
  ~sP5('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f261,f1133,f373])).
fof(f3589,plain,(
  ~sP5('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f261,f1133,f373])).
fof(f3590,plain,(
  ~sP5('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f261,f1133,f373])).
fof(f4071,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl380),
  inference(cnf_transformation,[],[f4071_D])).
fof(f4071_D,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f3591,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl2 | $spl175)),
  inference(unit_resulting_resolution,[],[f520,f862,f851,f1133,f375])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f226])).
fof(f851,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f379])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f228])).
fof(f228,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f862,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f383])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f231])).
fof(f231,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f520,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f386])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f3592,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl2 | $spl175)),
  inference(unit_resulting_resolution,[],[f520,f519,f851,f1133,f375])).
fof(f519,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f232])).
fof(f4069,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl378),
  inference(cnf_transformation,[],[f4069_D])).
fof(f4069_D,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f3593,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl2 | $spl175)),
  inference(unit_resulting_resolution,[],[f521,f862,f510,f1133,f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f226])).
fof(f510,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f378])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f229])).
fof(f521,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f387])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f234])).
fof(f3594,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl2 | $spl175)),
  inference(unit_resulting_resolution,[],[f521,f862,f851,f1133,f376])).
fof(f4067,plain,(
  ~sP4('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl377),
  inference(cnf_transformation,[],[f4067_D])).
fof(f4067_D,plain,(
  ~sP4('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl377),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl377])])).
fof(f3595,plain,(
  ~sP4('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f445,f1133,f377])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f3596,plain,(
  ~sP4('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f445,f1133,f377])).
fof(f3597,plain,(
  ~sP4('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f445,f1133,f377])).
fof(f3598,plain,(
  ~sP4('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f262,f1133,f377])).
fof(f3599,plain,(
  ~sP4('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f262,f1133,f377])).
fof(f3600,plain,(
  ~sP4('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f262,f1133,f377])).
fof(f3601,plain,(
  ~sP4('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f261,f1133,f377])).
fof(f3602,plain,(
  ~sP4('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f261,f1133,f377])).
fof(f3603,plain,(
  ~sP4('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f261,f1133,f377])).
fof(f4063,plain,(
  ~sP3('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl375),
  inference(cnf_transformation,[],[f4063_D])).
fof(f4063_D,plain,(
  ~sP3('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl375),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl375])])).
fof(f3608,plain,(
  ~sP3('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f445,f1133,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f232])).
fof(f3609,plain,(
  ~sP3('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f445,f1133,f381])).
fof(f3610,plain,(
  ~sP3('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f445,f1133,f381])).
fof(f3611,plain,(
  ~sP3('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f262,f1133,f381])).
fof(f3612,plain,(
  ~sP3('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f262,f1133,f381])).
fof(f3613,plain,(
  ~sP3('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f262,f1133,f381])).
fof(f3614,plain,(
  ~sP3('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f261,f1133,f381])).
fof(f3615,plain,(
  ~sP3('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f261,f1133,f381])).
fof(f3616,plain,(
  ~sP3('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f261,f1133,f381])).
fof(f4059,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','F_0') | $spl373),
  inference(cnf_transformation,[],[f4059_D])).
fof(f4059_D,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','F_0') <=> ~$spl373),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl373])])).
fof(f3621,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f267,f1133,f385])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f234])).
fof(f4053,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl369),
  inference(cnf_transformation,[],[f4053_D])).
fof(f4053_D,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') <=> ~$spl369),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl369])])).
fof(f3622,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f267,f1133,f385])).
fof(f4047,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl365),
  inference(cnf_transformation,[],[f4047_D])).
fof(f4047_D,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl365),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl365])])).
fof(f3623,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f1133,f1133,f385])).
fof(f4043,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl361),
  inference(cnf_transformation,[],[f4043_D])).
fof(f4043_D,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl361),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl361])])).
fof(f3624,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f1133,f1133,f385])).
fof(f4055,plain,(
  ~sP2('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl371),
  inference(cnf_transformation,[],[f4055_D])).
fof(f4055_D,plain,(
  ~sP2('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') <=> ~$spl371),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl371])])).
fof(f3627,plain,(
  ~sP2('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f267,f1133,f385])).
fof(f3628,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f267,f1133,f385])).
fof(f4045,plain,(
  ~sP2('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl363),
  inference(cnf_transformation,[],[f4045_D])).
fof(f4045_D,plain,(
  ~sP2('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl363),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl363])])).
fof(f3629,plain,(
  ~sP2('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f1133,f1133,f385])).
fof(f3630,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f1133,f1133,f385])).
fof(f4049,plain,(
  ~sP2('F_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl367),
  inference(cnf_transformation,[],[f4049_D])).
fof(f4049_D,plain,(
  ~sP2('F_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl367),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl367])])).
fof(f3633,plain,(
  ~sP2('F_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f267,f1133,f385])).
fof(f3634,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f267,f1133,f385])).
fof(f3635,plain,(
  ~sP2('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f1133,f1133,f385])).
fof(f3636,plain,(
  ~sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f1133,f1133,f385])).
fof(f4039,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','F_0') | $spl358),
  inference(cnf_transformation,[],[f4039_D])).
fof(f4039_D,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','F_0') <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f3639,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','F_0') | ($spl81 | $spl175)),
  inference(unit_resulting_resolution,[],[f763,f1133,f389])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f115,f152,f151,f150,f149])).
fof(f115,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X17,X18,X19] : ((p(s(bool,X19)) <=> s(bool,X17) = s(bool,X18)) <=> ((p(s(bool,X19)) | p(s(bool,X18)) | p(s(bool,X17))) & (p(s(bool,X19)) | ~p(s(bool,X17)) | ~p(s(bool,X18))) & (p(s(bool,X18)) | ~p(s(bool,X17)) | ~p(s(bool,X19))) & (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)))))),
  file('thm.relation.reflexive_TC','thm.sat.sat_2|49|_prove_ltb')).
fof(f4037,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl356),
  inference(cnf_transformation,[],[f4037_D])).
fof(f4037_D,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f3640,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl173 | $spl175)),
  inference(unit_resulting_resolution,[],[f1130,f1133,f389])).
fof(f4035,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','T_0') | $spl354),
  inference(cnf_transformation,[],[f4035_D])).
fof(f4035_D,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','T_0') <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f3641,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','T_0') | ($spl81 | $spl175)),
  inference(unit_resulting_resolution,[],[f763,f1133,f389])).
fof(f4033,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl352),
  inference(cnf_transformation,[],[f4033_D])).
fof(f4033_D,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f3642,plain,(
  sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | ($spl173 | $spl175)),
  inference(unit_resulting_resolution,[],[f1130,f1133,f389])).
fof(f4031,plain,(
  ~sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl351),
  inference(cnf_transformation,[],[f4031_D])).
fof(f4031_D,plain,(
  ~sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') <=> ~$spl351),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl351])])).
fof(f3644,plain,(
  ~sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | ($spl76 | $spl175)),
  inference(unit_resulting_resolution,[],[f755,f1133,f392])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~sP5(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f235])).
fof(f4029,plain,(
  ~sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl349),
  inference(cnf_transformation,[],[f4029_D])).
fof(f4029_D,plain,(
  ~sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl349),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl349])])).
fof(f3645,plain,(
  ~sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | ($spl170 | $spl175)),
  inference(unit_resulting_resolution,[],[f1128,f1133,f392])).
fof(f4027,plain,(
  ~sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl347),
  inference(cnf_transformation,[],[f4027_D])).
fof(f4027_D,plain,(
  ~sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl347),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl347])])).
fof(f3649,plain,(
  ~sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ($spl76 | $spl175)),
  inference(unit_resulting_resolution,[],[f755,f1133,f392])).
fof(f4025,plain,(
  ~sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl345),
  inference(cnf_transformation,[],[f4025_D])).
fof(f4025_D,plain,(
  ~sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl345),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl345])])).
fof(f3650,plain,(
  ~sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl170 | $spl175)),
  inference(unit_resulting_resolution,[],[f1128,f1133,f392])).
fof(f4019,plain,(
  ~sP6('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl343),
  inference(cnf_transformation,[],[f4019_D])).
fof(f4019_D,plain,(
  ~sP6('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl343),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl343])])).
fof(f3671,plain,(
  ~sP6('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f445,f1133,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f3672,plain,(
  ~sP6('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f445,f1133,f400])).
fof(f3673,plain,(
  ~sP6('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f445,f1133,f400])).
fof(f3674,plain,(
  ~sP6('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f262,f1133,f400])).
fof(f3675,plain,(
  ~sP6('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f262,f1133,f400])).
fof(f3676,plain,(
  ~sP6('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f262,f1133,f400])).
fof(f3677,plain,(
  ~sP6('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f261,f1133,f400])).
fof(f3678,plain,(
  ~sP6('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f261,f1133,f400])).
fof(f3679,plain,(
  ~sP6('T_0','T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f261,f1133,f400])).
fof(f4009,plain,(
  ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','T_0') | $spl341),
  inference(cnf_transformation,[],[f4009_D])).
fof(f4009_D,plain,(
  ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','T_0') <=> ~$spl341),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl341])])).
fof(f3708,plain,(
  ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f445,f1133,f409])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f117,f155,f154])).
fof(f155,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f117,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X17,X18,X19] : ((p(s(bool,X19)) <=> (p(s(bool,X18)) & p(s(bool,X17)))) <=> ((p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,X17))) & (p(s(bool,X18)) | ~p(s(bool,X19))) & (p(s(bool,X17)) | ~p(s(bool,X19)))))),
  file('thm.relation.reflexive_TC','thm.sat.sat_23_prove_ltb')).
fof(f3709,plain,(
  ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f445,f1133,f409])).
fof(f3710,plain,(
  ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f445,f1133,f409])).
fof(f3711,plain,(
  ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f262,f1133,f409])).
fof(f3712,plain,(
  ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f262,f1133,f409])).
fof(f3713,plain,(
  ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f262,f1133,f409])).
fof(f3714,plain,(
  ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f261,f1133,f409])).
fof(f3715,plain,(
  ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f261,f1133,f409])).
fof(f3716,plain,(
  ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0','T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f261,f1133,f409])).
fof(f3922,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl278),
  inference(cnf_transformation,[],[f3922_D])).
fof(f3922_D,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f3733,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f267,f942,f1133,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f244])).
fof(f244,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f942,plain,(
  ( ! [X0,X1] : (sP8('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f247])).
fof(f247,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f3918,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl274),
  inference(cnf_transformation,[],[f3918_D])).
fof(f3918_D,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f3734,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f1133,f942,f1133,f416])).
fof(f3920,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl276),
  inference(cnf_transformation,[],[f3920_D])).
fof(f3920_D,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f3737,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f267,f942,f1133,f416])).
fof(f3738,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f1133,f942,f1133,f416])).
fof(f4003,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl339),
  inference(cnf_transformation,[],[f4003_D])).
fof(f4003_D,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') <=> ~$spl339),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl339])])).
fof(f3741,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f267,f1133,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f248])).
fof(f3742,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f267,f1133,f417])).
fof(f3743,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f267,f1133,f417])).
fof(f3997,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl335),
  inference(cnf_transformation,[],[f3997_D])).
fof(f3997_D,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl335),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl335])])).
fof(f3744,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f1133,f1133,f417])).
fof(f3745,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f1133,f1133,f417])).
fof(f3746,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f1133,f1133,f417])).
fof(f3999,plain,(
  ~sP8('T_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl337),
  inference(cnf_transformation,[],[f3999_D])).
fof(f3999_D,plain,(
  ~sP8('T_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl337),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl337])])).
fof(f3749,plain,(
  ~sP8('T_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f267,f1133,f417])).
fof(f3750,plain,(
  ~sP8('T_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f267,f1133,f417])).
fof(f3751,plain,(
  ~sP8('T_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f267,f1133,f417])).
fof(f3752,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f445,f1133,f1133,f417])).
fof(f3753,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f262,f1133,f1133,f417])).
fof(f3754,plain,(
  ~sP8('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f261,f1133,f1133,f417])).
fof(f3993,plain,(
  sP9('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl332),
  inference(cnf_transformation,[],[f3993_D])).
fof(f3993_D,plain,(
  sP9('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f3758,plain,(
  sP9('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f267,f1133,f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f119,f158,f157])).
fof(f119,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X17,X18,X19] : ((p(s(bool,X19)) <=> (p(s(bool,X18)) | p(s(bool,X17)))) <=> ((p(s(bool,X19)) | ~p(s(bool,X18))) & (p(s(bool,X19)) | ~p(s(bool,X17))) & (p(s(bool,X18)) | p(s(bool,X17)) | ~p(s(bool,X19)))))),
  file('thm.relation.reflexive_TC','thm.sat.sat_25_prove_ltb')).
fof(f3983,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl324),
  inference(cnf_transformation,[],[f3983_D])).
fof(f3983_D,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f3759,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f267,f1133,f421])).
fof(f3988,plain,(
  sP9('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl328),
  inference(cnf_transformation,[],[f3988_D])).
fof(f3988_D,plain,(
  sP9('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f3760,plain,(
  sP9('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f1133,f1133,f421])).
fof(f3979,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl320),
  inference(cnf_transformation,[],[f3979_D])).
fof(f3979_D,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f3761,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f1133,f1133,f421])).
fof(f3990,plain,(
  sP9('F_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl330),
  inference(cnf_transformation,[],[f3990_D])).
fof(f3990_D,plain,(
  sP9('F_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f3763,plain,(
  sP9('F_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f267,f1133,f421])).
fof(f3981,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl322),
  inference(cnf_transformation,[],[f3981_D])).
fof(f3981_D,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f3764,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f267,f1133,f421])).
fof(f3765,plain,(
  sP9('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f1133,f1133,f421])).
fof(f3766,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f1133,f1133,f421])).
fof(f3985,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','F_0') | $spl326),
  inference(cnf_transformation,[],[f3985_D])).
fof(f3985_D,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','F_0') <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f3768,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f267,f1133,f421])).
fof(f3769,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f267,f1133,f421])).
fof(f3770,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f1133,f1133,f421])).
fof(f3771,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f1133,f1133,f421])).
fof(f3977,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl319),
  inference(cnf_transformation,[],[f3977_D])).
fof(f3977_D,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') <=> ~$spl319),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl319])])).
fof(f3772,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f445,f1133,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f250])).
fof(f3972,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl315),
  inference(cnf_transformation,[],[f3972_D])).
fof(f3972_D,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl315),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl315])])).
fof(f3773,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f445,f1133,f424])).
fof(f3774,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f262,f1133,f424])).
fof(f3775,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f262,f1133,f424])).
fof(f3776,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f261,f1133,f424])).
fof(f3777,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f261,f1133,f424])).
fof(f3974,plain,(
  ~sP9('T_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl317),
  inference(cnf_transformation,[],[f3974_D])).
fof(f3974_D,plain,(
  ~sP9('T_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl317),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl317])])).
fof(f3779,plain,(
  ~sP9('T_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f445,f1133,f424])).
fof(f3780,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f445,f1133,f424])).
fof(f3781,plain,(
  ~sP9('T_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f262,f1133,f424])).
fof(f3782,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f262,f1133,f424])).
fof(f3783,plain,(
  ~sP9('T_0','F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f267,f261,f1133,f424])).
fof(f3784,plain,(
  ~sP9('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f261,f1133,f424])).
fof(f3965,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl312),
  inference(cnf_transformation,[],[f3965_D])).
fof(f3965_D,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f3808,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f445,f990,f1133,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f252])).
fof(f252,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f990,plain,(
  ( ! [X0,X1] : (sP10('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f255])).
fof(f255,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f3809,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f262,f990,f1133,f433])).
fof(f3810,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f261,f990,f1133,f433])).
fof(f3961,plain,(
  ~sP10('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl311),
  inference(cnf_transformation,[],[f3961_D])).
fof(f3961_D,plain,(
  ~sP10('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') <=> ~$spl311),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl311])])).
fof(f3814,plain,(
  ~sP10('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f445,f1133,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f256])).
fof(f3815,plain,(
  ~sP10('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f445,f1133,f434])).
fof(f3816,plain,(
  ~sP10('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f445,f1133,f434])).
fof(f3817,plain,(
  ~sP10('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f262,f1133,f434])).
fof(f3818,plain,(
  ~sP10('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f262,f1133,f434])).
fof(f3819,plain,(
  ~sP10('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f262,f1133,f434])).
fof(f3820,plain,(
  ~sP10('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f261,f1133,f434])).
fof(f3821,plain,(
  ~sP10('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f261,f1133,f434])).
fof(f3822,plain,(
  ~sP10('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f261,f1133,f434])).
fof(f3957,plain,(
  sP11('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl308),
  inference(cnf_transformation,[],[f3957_D])).
fof(f3957_D,plain,(
  sP11('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f3827,plain,(
  sP11('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f445,f1133,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP11(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f144,f161,f160])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X17,X18,X19] : ((p(s(bool,X19)) <=> (p(s(bool,X18)) => p(s(bool,X17)))) <=> ((p(s(bool,X19)) | p(s(bool,X18))) & (p(s(bool,X19)) | ~p(s(bool,X17))) & (~p(s(bool,X18)) | p(s(bool,X17)) | ~p(s(bool,X19)))))),
  file('thm.relation.reflexive_TC','thm.sat.sat_27_prove_ltb')).
fof(f3952,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl304),
  inference(cnf_transformation,[],[f3952_D])).
fof(f3952_D,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f3828,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f445,f1133,f438])).
fof(f3829,plain,(
  sP11('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f262,f1133,f438])).
fof(f3830,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f262,f1133,f438])).
fof(f3831,plain,(
  sP11('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f267,f261,f1133,f438])).
fof(f3832,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f1133,f261,f1133,f438])).
fof(f3954,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','T_0') | $spl306),
  inference(cnf_transformation,[],[f3954_D])).
fof(f3954_D,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','T_0') <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f3834,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f267,f1133,f438])).
fof(f3835,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f267,f1133,f438])).
fof(f3836,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f267,f1133,f438])).
fof(f3837,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f1133,f1133,f438])).
fof(f3838,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f1133,f1133,f438])).
fof(f3839,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f1133,f1133,f438])).
fof(f3949,plain,(
  ~sP11('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl303),
  inference(cnf_transformation,[],[f3949_D])).
fof(f3949_D,plain,(
  ~sP11('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') <=> ~$spl303),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl303])])).
fof(f3845,plain,(
  ~sP11('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f445,f1133,f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f258])).
fof(f3846,plain,(
  ~sP11('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f445,f1133,f441])).
fof(f3847,plain,(
  ~sP11('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f445,f1133,f441])).
fof(f3848,plain,(
  ~sP11('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f262,f1133,f441])).
fof(f3849,plain,(
  ~sP11('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f262,f1133,f441])).
fof(f3850,plain,(
  ~sP11('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f262,f1133,f441])).
fof(f3851,plain,(
  ~sP11('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f445,f261,f1133,f441])).
fof(f3852,plain,(
  ~sP11('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f262,f261,f1133,f441])).
fof(f3853,plain,(
  ~sP11('T_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'T_0') | $spl175),
  inference(unit_resulting_resolution,[],[f261,f261,f1133,f441])).
fof(f3944,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl300),
  inference(cnf_transformation,[],[f3944_D])).
fof(f3944_D,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f3867,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | ($spl2 | $spl175)),
  inference(unit_resulting_resolution,[],[f1133,f661])).
fof(f661,plain,(
  ( ! [X6] : (sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X6) | p(s(bool,X6))) ) | $spl2),
  inference(resolution,[],[f452,f359])).
fof(f3942,plain,(
  sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl298),
  inference(cnf_transformation,[],[f3942_D])).
fof(f3942_D,plain,(
  sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f3868,plain,(
  sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ($spl2 | $spl175)),
  inference(unit_resulting_resolution,[],[f1133,f662])).
fof(f662,plain,(
  ( ! [X7] : (sP1(X7,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | p(s(bool,X7))) ) | $spl2),
  inference(resolution,[],[f452,f362])).
fof(f3940,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl296),
  inference(cnf_transformation,[],[f3940_D])).
fof(f3940_D,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f3869,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ($spl2 | $spl175)),
  inference(unit_resulting_resolution,[],[f267,f1133,f678])).
fof(f678,plain,(
  ( ! [X35,X36] : (sP11(X36,X35,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | p(s(bool,X35)) | p(s(bool,X36))) ) | $spl2),
  inference(resolution,[],[f452,f438])).
fof(f3936,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl292),
  inference(cnf_transformation,[],[f3936_D])).
fof(f3936_D,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f3870,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ($spl2 | $spl175)),
  inference(unit_resulting_resolution,[],[f1133,f1133,f678])).
fof(f3938,plain,(
  sP11('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl294),
  inference(cnf_transformation,[],[f3938_D])).
fof(f3938_D,plain,(
  sP11('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f3871,plain,(
  sP11('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ($spl2 | $spl175)),
  inference(unit_resulting_resolution,[],[f267,f1133,f678])).
fof(f3872,plain,(
  sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ($spl2 | $spl175)),
  inference(unit_resulting_resolution,[],[f1133,f1133,f678])).
fof(f3934,plain,(
  sP9('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl290),
  inference(cnf_transformation,[],[f3934_D])).
fof(f3934_D,plain,(
  sP9('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f3878,plain,(
  sP9('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f267,f1133,f1041])).
fof(f1041,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X27) | p(s(bool,X27)) | p(s(bool,X26))) ) | $spl1),
  inference(resolution,[],[f450,f421])).
fof(f3930,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | $spl286),
  inference(cnf_transformation,[],[f3930_D])).
fof(f3930_D,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f3879,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f1133,f1133,f1041])).
fof(f3932,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl288),
  inference(cnf_transformation,[],[f3932_D])).
fof(f3932_D,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f3880,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f267,f1133,f1041])).
fof(f3881,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f1133,f1133,f1041])).
fof(f3928,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl284),
  inference(cnf_transformation,[],[f3928_D])).
fof(f3928_D,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f3882,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f267,f1133,f1042])).
fof(f1042,plain,(
  ( ! [X28,X29] : (sP9(X28,X29,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl1),
  inference(resolution,[],[f450,f421])).
fof(f3924,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl280),
  inference(cnf_transformation,[],[f3924_D])).
fof(f3924_D,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f3883,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f1133,f1133,f1042])).
fof(f3926,plain,(
  sP9('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl282),
  inference(cnf_transformation,[],[f3926_D])).
fof(f3926_D,plain,(
  sP9('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f3884,plain,(
  sP9('F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f267,f1133,f1042])).
fof(f3885,plain,(
  sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f1133,f1133,f1042])).
fof(f3886,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'F_0') | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f267,f1133,f1043])).
fof(f1043,plain,(
  ( ! [X30,X31] : (sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X30,X31) | p(s(bool,X30)) | p(s(bool,X31))) ) | $spl1),
  inference(resolution,[],[f450,f421])).
fof(f3887,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f1133,f1133,f1043])).
fof(f3889,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f267,f1133,f1043])).
fof(f3890,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))),'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))) | ($spl1 | $spl175)),
  inference(unit_resulting_resolution,[],[f1133,f1133,f1043])).
fof(f3534,plain,(
  ~sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl273),
  inference(cnf_transformation,[],[f3534_D])).
fof(f3534_D,plain,(
  ~sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl273),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl273])])).
fof(f3532,plain,(
  ~sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl143),
  inference(unit_resulting_resolution,[],[f1100,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f226])).
fof(f3521,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl271),
  inference(cnf_transformation,[],[f3521_D])).
fof(f3521_D,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl271),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl271])])).
fof(f3519,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl139),
  inference(unit_resulting_resolution,[],[f1096,f370])).
fof(f3513,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl269),
  inference(cnf_transformation,[],[f3513_D])).
fof(f3513_D,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl269),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl269])])).
fof(f3511,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl137),
  inference(unit_resulting_resolution,[],[f1094,f370])).
fof(f3492,plain,(
  sP8('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl266),
  inference(cnf_transformation,[],[f3492_D])).
fof(f3492_D,plain,(
  sP8('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f3489,plain,(
  sP8('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl114),
  inference(resolution,[],[f1072,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f245])).
fof(f3483,plain,(
  sP8('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl114),
  inference(unit_resulting_resolution,[],[f1072,f412])).
fof(f3311,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl264),
  inference(cnf_transformation,[],[f3311_D])).
fof(f3311_D,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f3295,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl173)),
  inference(unit_resulting_resolution,[],[f1130,f1033])).
fof(f1033,plain,(
  ( ! [X10,X11] : (sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X10,X11) | s(bool,X10) = s(bool,X11)) ) | $spl1),
  inference(resolution,[],[f450,f389])).
fof(f3309,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl262),
  inference(cnf_transformation,[],[f3309_D])).
fof(f3309_D,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f3297,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | ($spl1 | $spl173)),
  inference(unit_resulting_resolution,[],[f1130,f1033])).
fof(f2762,plain,(
  sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl260),
  inference(cnf_transformation,[],[f2762_D])).
fof(f2762_D,plain,(
  sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f2744,plain,(
  sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl173),
  inference(unit_resulting_resolution,[],[f267,f1130,f389])).
fof(f2759,plain,(
  sP5('F_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl258),
  inference(cnf_transformation,[],[f2759_D])).
fof(f2759_D,plain,(
  sP5('F_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f2749,plain,(
  sP5('F_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl173),
  inference(unit_resulting_resolution,[],[f267,f1130,f389])).
fof(f2738,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,'F_0'))) | $spl257),
  inference(cnf_transformation,[],[f2738_D])).
fof(f2738_D,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,'F_0'))) <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f2706,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,'F_0'))) | ($spl1 | $spl170)),
  inference(superposition,[],[f882,f1128])).
fof(f882,plain,(
  ( ! [X0] : (~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X0,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f304,f450,f392])).
fof(f304,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X11] : s(X6,X11) = s(X6,'const.combin.I_1'(s(X6,X11)))),
  file('thm.relation.reflexive_TC','thm.combin.I_THM')).
fof(f2736,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,'F_0')),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl255),
  inference(cnf_transformation,[],[f2736_D])).
fof(f2736_D,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,'F_0')),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f2705,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,'F_0')),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl170)),
  inference(superposition,[],[f879,f1128])).
fof(f879,plain,(
  ( ! [X0] : (~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,X0)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f304,f450,f392])).
fof(f2734,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,'F_0')),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl252),
  inference(cnf_transformation,[],[f2734_D])).
fof(f2734_D,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,'F_0')),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f2704,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,'F_0')),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl2 | $spl170)),
  inference(superposition,[],[f527,f1128])).
fof(f527,plain,(
  ( ! [X0] : (sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,X0)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f452,f390])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP5(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f235])).
fof(f2732,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,'F_0'))) | $spl250),
  inference(cnf_transformation,[],[f2732_D])).
fof(f2732_D,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,'F_0'))) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f2703,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,'F_0'))) | ($spl2 | $spl170)),
  inference(superposition,[],[f524,f1128])).
fof(f524,plain,(
  ( ! [X0] : (sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0,'const.combin.I_1'(s(bool,X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f452,f390])).
fof(f2724,plain,(
  sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl248),
  inference(cnf_transformation,[],[f2724_D])).
fof(f2724_D,plain,(
  sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f2633,plain,(
  sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f445,f1128,f390])).
fof(f2634,plain,(
  sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f262,f1128,f390])).
fof(f2635,plain,(
  sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f261,f1128,f390])).
fof(f2722,plain,(
  sP5('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl246),
  inference(cnf_transformation,[],[f2722_D])).
fof(f2722_D,plain,(
  sP5('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f2637,plain,(
  sP5('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl170),
  inference(unit_resulting_resolution,[],[f445,f1128,f390])).
fof(f2638,plain,(
  sP5('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl170),
  inference(unit_resulting_resolution,[],[f262,f1128,f390])).
fof(f2639,plain,(
  sP5('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl170),
  inference(unit_resulting_resolution,[],[f261,f1128,f390])).
fof(f2629,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,'T_0'))) | $spl245),
  inference(cnf_transformation,[],[f2629_D])).
fof(f2629_D,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,'T_0'))) <=> ~$spl245),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl245])])).
fof(f2623,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,'T_0'))) | ($spl1 | $spl76)),
  inference(superposition,[],[f882,f755])).
fof(f2617,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,'T_0')),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl243),
  inference(cnf_transformation,[],[f2617_D])).
fof(f2617_D,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,'T_0')),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f2611,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.combin.I_1'(s(bool,'T_0')),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ($spl1 | $spl76)),
  inference(superposition,[],[f879,f755])).
fof(f2552,plain,(
  ~sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl241),
  inference(cnf_transformation,[],[f2552_D])).
fof(f2552_D,plain,(
  ~sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') <=> ~$spl241),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl241])])).
fof(f2548,plain,(
  ~sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl151),
  inference(unit_resulting_resolution,[],[f1108,f372])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f226])).
fof(f2543,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') | $spl239),
  inference(cnf_transformation,[],[f2543_D])).
fof(f2543_D,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') <=> ~$spl239),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl239])])).
fof(f2539,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') | $spl149),
  inference(unit_resulting_resolution,[],[f1106,f371])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f226])).
fof(f2534,plain,(
  ~sP5('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl237),
  inference(cnf_transformation,[],[f2534_D])).
fof(f2534_D,plain,(
  ~sP5('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl237),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl237])])).
fof(f2532,plain,(
  ~sP5('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl147),
  inference(unit_resulting_resolution,[],[f1104,f370])).
fof(f2527,plain,(
  ~sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl235),
  inference(cnf_transformation,[],[f2527_D])).
fof(f2527_D,plain,(
  ~sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f2525,plain,(
  ~sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl145),
  inference(unit_resulting_resolution,[],[f1102,f370])).
fof(f2520,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','F_0') | $spl233),
  inference(cnf_transformation,[],[f2520_D])).
fof(f2520_D,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','F_0') <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f2518,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','F_0') | $spl141),
  inference(unit_resulting_resolution,[],[f1098,f370])).
fof(f2481,plain,(
  sP8('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl230),
  inference(cnf_transformation,[],[f2481_D])).
fof(f2481_D,plain,(
  sP8('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f2479,plain,(
  sP8('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl118),
  inference(resolution,[],[f1076,f412])).
fof(f2473,plain,(
  sP8('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f1076,f412])).
fof(f2464,plain,(
  sP8('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl228),
  inference(cnf_transformation,[],[f2464_D])).
fof(f2464_D,plain,(
  sP8('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f2462,plain,(
  sP8('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl116),
  inference(resolution,[],[f1074,f412])).
fof(f2456,plain,(
  sP8('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl116),
  inference(unit_resulting_resolution,[],[f1074,f412])).
fof(f2412,plain,(
  sP10('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl226),
  inference(cnf_transformation,[],[f2412_D])).
fof(f2412_D,plain,(
  sP10('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f2410,plain,(
  sP10('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl96),
  inference(resolution,[],[f1054,f429])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f253])).
fof(f2404,plain,(
  sP10('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f1054,f429])).
fof(f2155,plain,(
  sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl224),
  inference(cnf_transformation,[],[f2155_D])).
fof(f2155_D,plain,(
  sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f2151,plain,(
  sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f267,f521,f851,f862,f376])).
fof(f2152,plain,(
  sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f267,f521,f510,f862,f376])).
fof(f2148,plain,(
  s(bool,'F_0') != s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ~p(s(bool,'F_0')) | p(s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))))),
  introduced(theory_axiom,[])).
fof(f2145,plain,(
  ( ! [X6,X7] : (sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X6,X7) | ~sP2(X7,X6,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ~sP3(X6,X7,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))))) ) | $spl220),
  inference(cnf_transformation,[],[f2145_D])).
fof(f2145_D,plain,(
  ( ! [X6,X7] : (sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X6,X7) | ~sP2(X7,X6,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ~sP3(X6,X7,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f2147,plain,(
  p(s(bool,'F_0')) | $spl222),
  inference(cnf_transformation,[],[f2147_D])).
fof(f2147_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f2144,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X6,X7) | ~sP3(X6,X7,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ~sP2(X7,X6,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))))) ) | ($spl1 | $spl170)),
  inference(forward_demodulation,[],[f2141,f1128])).
fof(f2141,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))))) | sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X6,X7) | ~sP3(X6,X7,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ~sP2(X7,X6,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))))) ) | $spl1),
  inference(resolution,[],[f852,f376])).
fof(f852,plain,(
  ( ! [X0,X1] : (sP4('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f380])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f2126,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl218),
  inference(cnf_transformation,[],[f2126_D])).
fof(f2126_D,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f2120,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f267,f520,f519,f851,f375])).
fof(f2108,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl217),
  inference(cnf_transformation,[],[f2108_D])).
fof(f2108_D,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl217),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl217])])).
fof(f2104,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl55),
  inference(unit_resulting_resolution,[],[f733,f372])).
fof(f2096,plain,(
  ~sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl215),
  inference(cnf_transformation,[],[f2096_D])).
fof(f2096_D,plain,(
  ~sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl215),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl215])])).
fof(f2092,plain,(
  ~sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl49),
  inference(unit_resulting_resolution,[],[f727,f371])).
fof(f2072,plain,(
  sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl212),
  inference(cnf_transformation,[],[f2072_D])).
fof(f2072_D,plain,(
  sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f2065,plain,(
  sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl36),
  inference(resolution,[],[f715,f393])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f237])).
fof(f237,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f155])).
fof(f2059,plain,(
  sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f715,f393])).
fof(f1855,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,'T_0')),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl210),
  inference(cnf_transformation,[],[f1855_D])).
fof(f1855_D,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,'T_0')),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f1825,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,'T_0')),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ($spl2 | $spl76)),
  inference(superposition,[],[f527,f755])).
fof(f1853,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,'T_0'))) | $spl208),
  inference(cnf_transformation,[],[f1853_D])).
fof(f1853_D,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,'T_0'))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f1824,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,'T_0'))) | ($spl2 | $spl76)),
  inference(superposition,[],[f524,f755])).
fof(f1846,plain,(
  sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl206),
  inference(cnf_transformation,[],[f1846_D])).
fof(f1846_D,plain,(
  sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f1752,plain,(
  sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f445,f755,f390])).
fof(f1753,plain,(
  sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f262,f755,f390])).
fof(f1754,plain,(
  sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f261,f755,f390])).
fof(f1844,plain,(
  sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl204),
  inference(cnf_transformation,[],[f1844_D])).
fof(f1844_D,plain,(
  sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f1756,plain,(
  sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl76),
  inference(unit_resulting_resolution,[],[f445,f755,f390])).
fof(f1757,plain,(
  sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl76),
  inference(unit_resulting_resolution,[],[f262,f755,f390])).
fof(f1758,plain,(
  sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl76),
  inference(unit_resulting_resolution,[],[f261,f755,f390])).
fof(f1672,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') | $spl203),
  inference(cnf_transformation,[],[f1672_D])).
fof(f1672_D,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') <=> ~$spl203),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl203])])).
fof(f1668,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') | $spl59),
  inference(unit_resulting_resolution,[],[f737,f372])).
fof(f1662,plain,(
  ~sP5('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl201),
  inference(cnf_transformation,[],[f1662_D])).
fof(f1662_D,plain,(
  ~sP5('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl201),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl201])])).
fof(f1658,plain,(
  ~sP5('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl57),
  inference(unit_resulting_resolution,[],[f735,f372])).
fof(f1652,plain,(
  ~sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl199),
  inference(cnf_transformation,[],[f1652_D])).
fof(f1652_D,plain,(
  ~sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f1648,plain,(
  ~sP5('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl53),
  inference(unit_resulting_resolution,[],[f731,f371])).
fof(f1642,plain,(
  ~sP5('F_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl197),
  inference(cnf_transformation,[],[f1642_D])).
fof(f1642_D,plain,(
  ~sP5('F_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f1638,plain,(
  ~sP5('F_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl51),
  inference(unit_resulting_resolution,[],[f729,f371])).
fof(f1610,plain,(
  sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl194),
  inference(cnf_transformation,[],[f1610_D])).
fof(f1610_D,plain,(
  sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f1605,plain,(
  sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl40),
  inference(resolution,[],[f719,f393])).
fof(f1599,plain,(
  sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl40),
  inference(unit_resulting_resolution,[],[f719,f393])).
fof(f1593,plain,(
  sP6('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl192),
  inference(cnf_transformation,[],[f1593_D])).
fof(f1593_D,plain,(
  sP6('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f1588,plain,(
  sP6('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl38),
  inference(resolution,[],[f717,f393])).
fof(f1582,plain,(
  sP6('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl38),
  inference(unit_resulting_resolution,[],[f717,f393])).
fof(f1527,plain,(
  sP4('F_0','T_0','F_0') | $spl186),
  inference(cnf_transformation,[],[f1527_D])).
fof(f1527_D,plain,(
  sP4('F_0','T_0','F_0') <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f1523,plain,(
  sP4('F_0','T_0','F_0') | $spl88),
  inference(resolution,[],[f791,f372])).
fof(f1529,plain,(
  sP3('F_0','T_0','F_0') | $spl188),
  inference(cnf_transformation,[],[f1529_D])).
fof(f1529_D,plain,(
  sP3('F_0','T_0','F_0') <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f1522,plain,(
  sP3('F_0','T_0','F_0') | $spl88),
  inference(resolution,[],[f791,f371])).
fof(f1531,plain,(
  sP2('T_0','F_0','F_0') | $spl190),
  inference(cnf_transformation,[],[f1531_D])).
fof(f1531_D,plain,(
  sP2('T_0','F_0','F_0') <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f1521,plain,(
  sP2('T_0','F_0','F_0') | $spl88),
  inference(resolution,[],[f791,f370])).
fof(f1517,plain,(
  sP2('T_0','F_0','F_0') | $spl88),
  inference(unit_resulting_resolution,[],[f791,f370])).
fof(f1518,plain,(
  sP3('F_0','T_0','F_0') | $spl88),
  inference(unit_resulting_resolution,[],[f791,f371])).
fof(f1519,plain,(
  sP4('F_0','T_0','F_0') | $spl88),
  inference(unit_resulting_resolution,[],[f791,f372])).
fof(f1502,plain,(
  sP10('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl184),
  inference(cnf_transformation,[],[f1502_D])).
fof(f1502_D,plain,(
  sP10('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f1500,plain,(
  sP10('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl10),
  inference(resolution,[],[f689,f429])).
fof(f1494,plain,(
  sP10('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl10),
  inference(unit_resulting_resolution,[],[f689,f429])).
fof(f1297,plain,(
  sP4('F_0','F_0','T_0') | $spl178),
  inference(cnf_transformation,[],[f1297_D])).
fof(f1297_D,plain,(
  sP4('F_0','F_0','T_0') <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f1294,plain,(
  sP4('F_0','F_0','T_0') | $spl84),
  inference(resolution,[],[f787,f372])).
fof(f1299,plain,(
  sP3('T_0','F_0','F_0') | $spl180),
  inference(cnf_transformation,[],[f1299_D])).
fof(f1299_D,plain,(
  sP3('T_0','F_0','F_0') <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f1293,plain,(
  sP3('T_0','F_0','F_0') | $spl84),
  inference(resolution,[],[f787,f371])).
fof(f1301,plain,(
  sP2('F_0','T_0','F_0') | $spl182),
  inference(cnf_transformation,[],[f1301_D])).
fof(f1301_D,plain,(
  sP2('F_0','T_0','F_0') <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f1292,plain,(
  sP2('F_0','T_0','F_0') | $spl84),
  inference(resolution,[],[f787,f370])).
fof(f1287,plain,(
  sP2('F_0','T_0','F_0') | $spl84),
  inference(unit_resulting_resolution,[],[f787,f370])).
fof(f1288,plain,(
  sP3('T_0','F_0','F_0') | $spl84),
  inference(unit_resulting_resolution,[],[f787,f371])).
fof(f1289,plain,(
  sP4('F_0','F_0','T_0') | $spl84),
  inference(unit_resulting_resolution,[],[f787,f372])).
fof(f1206,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl176),
  inference(cnf_transformation,[],[f1206_D])).
fof(f1206_D,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f1201,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f510,f519,f520,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | ~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f226])).
fof(f1202,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f510,f519,f520,f374])).
fof(f1203,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f510,f519,f520,f374])).
fof(f1128,plain,(
  s(bool,'F_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl170),
  inference(cnf_transformation,[],[f1128_D])).
fof(f1128_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f1023,plain,(
  s(bool,'F_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(resolution,[],[f450,f284])).
fof(f1022,plain,(
  s(bool,'F_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(resolution,[],[f450,f282])).
fof(f1133,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))))) | $spl175),
  inference(cnf_transformation,[],[f1133_D])).
fof(f1133_D,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))))) <=> ~$spl175),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl175])])).
fof(f1131,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))))) | $spl1),
  inference(forward_demodulation,[],[f792,f332])).
fof(f332,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6] : ! [X20,X21,X22] : s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X20),s(X6,X21),s(X6,X22))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),X20))),s(X6,X21))),s(X6,X22)))),
  file('thm.relation.reflexive_TC','ARITY_1_3_const.relation.TC')).
fof(f792,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12)))),s(sK12,sK14('const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)),sK12))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f309])).
fof(f309,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK14(X1,X0)))),s(X0,sK14(X1,X0))))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f1130,plain,(
  s(bool,'T_0') != s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl173),
  inference(cnf_transformation,[],[f1130_D])).
fof(f1130_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f793,plain,(
  s(bool,'T_0') != s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f277])).
fof(f794,plain,(
  s(bool,'T_0') != s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f279])).
fof(f795,plain,(
  s(bool,'F_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f282])).
fof(f796,plain,(
  s(bool,'F_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f284])).
fof(f801,plain,(
  s(bool,'F_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f450,f334])).
fof(f803,plain,(
  s(bool,'F_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f450,f334])).
fof(f1126,plain,(
  ~sP0('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl169),
  inference(cnf_transformation,[],[f1126_D])).
fof(f1126_D,plain,(
  ~sP0('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl169),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl169])])).
fof(f805,plain,(
  ~sP0('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f267,f450,f342])).
fof(f1122,plain,(
  ~sP0('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl165),
  inference(cnf_transformation,[],[f1122_D])).
fof(f1122_D,plain,(
  ~sP0('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl165),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl165])])).
fof(f806,plain,(
  ~sP0('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f267,f450,f342])).
fof(f1124,plain,(
  ~sP0('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl167),
  inference(cnf_transformation,[],[f1124_D])).
fof(f1124_D,plain,(
  ~sP0('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl167),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl167])])).
fof(f807,plain,(
  ~sP0('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f267,f450,f342])).
fof(f808,plain,(
  ~sP0('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f267,f450,f342])).
fof(f809,plain,(
  ~sP0('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f267,f450,f346])).
fof(f810,plain,(
  ~sP0('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f267,f450,f346])).
fof(f811,plain,(
  ~sP0('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f267,f450,f346])).
fof(f812,plain,(
  ~sP0('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f267,f450,f346])).
fof(f1114,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl157),
  inference(cnf_transformation,[],[f1114_D])).
fof(f1114_D,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f813,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f450,f357])).
fof(f1112,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl155),
  inference(cnf_transformation,[],[f1112_D])).
fof(f1112_D,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f814,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f450,f357])).
fof(f1116,plain,(
  ~sP1('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl159),
  inference(cnf_transformation,[],[f1116_D])).
fof(f1116_D,plain,(
  ~sP1('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f815,plain,(
  ~sP1('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f450,f357])).
fof(f816,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f450,f357])).
fof(f1118,plain,(
  sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl160),
  inference(cnf_transformation,[],[f1118_D])).
fof(f1118_D,plain,(
  sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f817,plain,(
  sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f450,f359])).
fof(f818,plain,(
  sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f450,f359])).
fof(f819,plain,(
  sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f450,f359])).
fof(f1120,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl162),
  inference(cnf_transformation,[],[f1120_D])).
fof(f1120_D,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f820,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f450,f362])).
fof(f821,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f450,f362])).
fof(f822,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f450,f362])).
fof(f823,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f450,f363])).
fof(f824,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f450,f363])).
fof(f825,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f450,f363])).
fof(f826,plain,(
  sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f450,f364])).
fof(f827,plain,(
  sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f450,f364])).
fof(f828,plain,(
  sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f450,f364])).
fof(f829,plain,(
  ~sP1('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f450,f366])).
fof(f830,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f450,f366])).
fof(f831,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f450,f366])).
fof(f832,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f450,f366])).
fof(f1110,plain,(
  ~sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl153),
  inference(cnf_transformation,[],[f1110_D])).
fof(f1110_D,plain,(
  ~sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f833,plain,(
  ~sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f445,f450,f373])).
fof(f834,plain,(
  ~sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f445,f450,f373])).
fof(f835,plain,(
  ~sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f445,f450,f373])).
fof(f836,plain,(
  ~sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f262,f450,f373])).
fof(f837,plain,(
  ~sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f262,f450,f373])).
fof(f838,plain,(
  ~sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f262,f450,f373])).
fof(f839,plain,(
  ~sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f261,f450,f373])).
fof(f840,plain,(
  ~sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f261,f450,f373])).
fof(f841,plain,(
  ~sP5('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f261,f450,f373])).
fof(f1108,plain,(
  ~sP4('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl151),
  inference(cnf_transformation,[],[f1108_D])).
fof(f1108_D,plain,(
  ~sP4('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f842,plain,(
  ~sP4('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f445,f450,f377])).
fof(f843,plain,(
  ~sP4('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f445,f450,f377])).
fof(f844,plain,(
  ~sP4('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f445,f450,f377])).
fof(f845,plain,(
  ~sP4('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f262,f450,f377])).
fof(f846,plain,(
  ~sP4('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f262,f450,f377])).
fof(f847,plain,(
  ~sP4('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f262,f450,f377])).
fof(f848,plain,(
  ~sP4('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f261,f450,f377])).
fof(f849,plain,(
  ~sP4('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f261,f450,f377])).
fof(f850,plain,(
  ~sP4('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f261,f450,f377])).
fof(f1106,plain,(
  ~sP3('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl149),
  inference(cnf_transformation,[],[f1106_D])).
fof(f1106_D,plain,(
  ~sP3('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f853,plain,(
  ~sP3('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f445,f450,f381])).
fof(f854,plain,(
  ~sP3('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f445,f450,f381])).
fof(f855,plain,(
  ~sP3('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f445,f450,f381])).
fof(f856,plain,(
  ~sP3('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f262,f450,f381])).
fof(f857,plain,(
  ~sP3('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f262,f450,f381])).
fof(f858,plain,(
  ~sP3('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f262,f450,f381])).
fof(f859,plain,(
  ~sP3('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f261,f450,f381])).
fof(f860,plain,(
  ~sP3('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f261,f450,f381])).
fof(f861,plain,(
  ~sP3('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f261,f450,f381])).
fof(f1104,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','F_0') | $spl147),
  inference(cnf_transformation,[],[f1104_D])).
fof(f1104_D,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','F_0') <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f864,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f267,f450,f385])).
fof(f1100,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl143),
  inference(cnf_transformation,[],[f1100_D])).
fof(f1100_D,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f865,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f450,f267,f450,f385])).
fof(f1096,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl139),
  inference(cnf_transformation,[],[f1096_D])).
fof(f1096_D,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f866,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f450,f450,f385])).
fof(f1092,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl135),
  inference(cnf_transformation,[],[f1092_D])).
fof(f1092_D,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f867,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f450,f450,f385])).
fof(f1102,plain,(
  ~sP2('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl145),
  inference(cnf_transformation,[],[f1102_D])).
fof(f1102_D,plain,(
  ~sP2('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f868,plain,(
  ~sP2('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f267,f450,f385])).
fof(f869,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f450,f267,f450,f385])).
fof(f1094,plain,(
  ~sP2('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl137),
  inference(cnf_transformation,[],[f1094_D])).
fof(f1094_D,plain,(
  ~sP2('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f870,plain,(
  ~sP2('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f450,f450,f385])).
fof(f871,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f450,f450,f385])).
fof(f1098,plain,(
  ~sP2('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl141),
  inference(cnf_transformation,[],[f1098_D])).
fof(f1098_D,plain,(
  ~sP2('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f872,plain,(
  ~sP2('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f267,f450,f385])).
fof(f873,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f267,f450,f385])).
fof(f874,plain,(
  ~sP2('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f450,f450,f385])).
fof(f875,plain,(
  ~sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f450,f450,f385])).
fof(f1090,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','F_0') | $spl132),
  inference(cnf_transformation,[],[f1090_D])).
fof(f1090_D,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','F_0') <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f876,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','F_0') | ($spl1 | $spl81)),
  inference(unit_resulting_resolution,[],[f763,f450,f389])).
fof(f1088,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','T_0') | $spl130),
  inference(cnf_transformation,[],[f1088_D])).
fof(f1088_D,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','T_0') <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f877,plain,(
  sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','T_0') | ($spl1 | $spl81)),
  inference(unit_resulting_resolution,[],[f763,f450,f389])).
fof(f1086,plain,(
  ~sP6('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl129),
  inference(cnf_transformation,[],[f1086_D])).
fof(f1086_D,plain,(
  ~sP6('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f890,plain,(
  ~sP6('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f445,f450,f400])).
fof(f891,plain,(
  ~sP6('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f445,f450,f400])).
fof(f892,plain,(
  ~sP6('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f445,f450,f400])).
fof(f893,plain,(
  ~sP6('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f262,f450,f400])).
fof(f894,plain,(
  ~sP6('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f262,f450,f400])).
fof(f895,plain,(
  ~sP6('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f262,f450,f400])).
fof(f896,plain,(
  ~sP6('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f261,f450,f400])).
fof(f897,plain,(
  ~sP6('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f261,f450,f400])).
fof(f898,plain,(
  ~sP6('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f261,f450,f400])).
fof(f1084,plain,(
  ~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') | $spl127),
  inference(cnf_transformation,[],[f1084_D])).
fof(f1084_D,plain,(
  ~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f915,plain,(
  ~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f445,f450,f409])).
fof(f916,plain,(
  ~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f445,f450,f409])).
fof(f917,plain,(
  ~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f445,f450,f409])).
fof(f918,plain,(
  ~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f262,f450,f409])).
fof(f919,plain,(
  ~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f262,f450,f409])).
fof(f920,plain,(
  ~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f262,f450,f409])).
fof(f921,plain,(
  ~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f261,f450,f409])).
fof(f922,plain,(
  ~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f261,f450,f409])).
fof(f923,plain,(
  ~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f261,f450,f409])).
fof(f1082,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl125),
  inference(cnf_transformation,[],[f1082_D])).
fof(f1082_D,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f930,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f267,f450,f417])).
fof(f931,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f267,f450,f417])).
fof(f932,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f267,f450,f417])).
fof(f1078,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl121),
  inference(cnf_transformation,[],[f1078_D])).
fof(f1078_D,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f933,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f450,f450,f417])).
fof(f934,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f450,f450,f417])).
fof(f935,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f450,f450,f417])).
fof(f1080,plain,(
  ~sP8('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl123),
  inference(cnf_transformation,[],[f1080_D])).
fof(f1080_D,plain,(
  ~sP8('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f936,plain,(
  ~sP8('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f267,f450,f417])).
fof(f937,plain,(
  ~sP8('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f267,f450,f417])).
fof(f938,plain,(
  ~sP8('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f267,f450,f417])).
fof(f939,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f445,f450,f450,f417])).
fof(f940,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f450,f450,f417])).
fof(f941,plain,(
  ~sP8('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f261,f450,f450,f417])).
fof(f1076,plain,(
  sP9('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl118),
  inference(cnf_transformation,[],[f1076_D])).
fof(f1076_D,plain,(
  sP9('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f943,plain,(
  sP9('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f267,f450,f421])).
fof(f1068,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl110),
  inference(cnf_transformation,[],[f1068_D])).
fof(f1068_D,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f944,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f450,f267,f450,f421])).
fof(f1072,plain,(
  sP9('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl114),
  inference(cnf_transformation,[],[f1072_D])).
fof(f1072_D,plain,(
  sP9('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f945,plain,(
  sP9('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f450,f450,f421])).
fof(f1064,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl106),
  inference(cnf_transformation,[],[f1064_D])).
fof(f1064_D,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f946,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f450,f450,f421])).
fof(f1074,plain,(
  sP9('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl116),
  inference(cnf_transformation,[],[f1074_D])).
fof(f1074_D,plain,(
  sP9('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f947,plain,(
  sP9('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f267,f450,f421])).
fof(f1066,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl108),
  inference(cnf_transformation,[],[f1066_D])).
fof(f1066_D,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f948,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f267,f450,f421])).
fof(f949,plain,(
  sP9('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f450,f450,f421])).
fof(f950,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f450,f450,f421])).
fof(f1070,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','F_0') | $spl112),
  inference(cnf_transformation,[],[f1070_D])).
fof(f1070_D,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','F_0') <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f951,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f267,f450,f421])).
fof(f952,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f450,f267,f450,f421])).
fof(f953,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f450,f450,f421])).
fof(f954,plain,(
  sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f450,f450,f421])).
fof(f1062,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl105),
  inference(cnf_transformation,[],[f1062_D])).
fof(f1062_D,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f955,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f445,f450,f424])).
fof(f1058,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl101),
  inference(cnf_transformation,[],[f1058_D])).
fof(f1058_D,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f956,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f445,f450,f424])).
fof(f957,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f262,f450,f424])).
fof(f958,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f262,f450,f424])).
fof(f959,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f261,f450,f424])).
fof(f960,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f261,f450,f424])).
fof(f1060,plain,(
  ~sP9('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl103),
  inference(cnf_transformation,[],[f1060_D])).
fof(f1060_D,plain,(
  ~sP9('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f961,plain,(
  ~sP9('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f445,f450,f424])).
fof(f962,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f445,f450,f424])).
fof(f963,plain,(
  ~sP9('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f262,f450,f424])).
fof(f964,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f262,f450,f424])).
fof(f965,plain,(
  ~sP9('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f267,f261,f450,f424])).
fof(f966,plain,(
  ~sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f261,f450,f424])).
fof(f1056,plain,(
  ~sP10('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl99),
  inference(cnf_transformation,[],[f1056_D])).
fof(f1056_D,plain,(
  ~sP10('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f980,plain,(
  ~sP10('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f445,f450,f434])).
fof(f981,plain,(
  ~sP10('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f445,f450,f434])).
fof(f982,plain,(
  ~sP10('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f445,f450,f434])).
fof(f983,plain,(
  ~sP10('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f262,f450,f434])).
fof(f984,plain,(
  ~sP10('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f262,f450,f434])).
fof(f985,plain,(
  ~sP10('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f262,f450,f434])).
fof(f986,plain,(
  ~sP10('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f261,f450,f434])).
fof(f987,plain,(
  ~sP10('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f261,f450,f434])).
fof(f988,plain,(
  ~sP10('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f261,f450,f434])).
fof(f1054,plain,(
  sP11('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl96),
  inference(cnf_transformation,[],[f1054_D])).
fof(f1054_D,plain,(
  sP11('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f991,plain,(
  sP11('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f445,f450,f438])).
fof(f1050,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl92),
  inference(cnf_transformation,[],[f1050_D])).
fof(f1050_D,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f992,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f450,f445,f450,f438])).
fof(f993,plain,(
  sP11('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f262,f450,f438])).
fof(f994,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f450,f262,f450,f438])).
fof(f995,plain,(
  sP11('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f267,f261,f450,f438])).
fof(f996,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f450,f261,f450,f438])).
fof(f1052,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','T_0') | $spl94),
  inference(cnf_transformation,[],[f1052_D])).
fof(f1052_D,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','T_0') <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f997,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f267,f450,f438])).
fof(f998,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f267,f450,f438])).
fof(f999,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f267,f450,f438])).
fof(f1000,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f450,f450,f438])).
fof(f1001,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f450,f450,f438])).
fof(f1002,plain,(
  sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f450,f450,f438])).
fof(f1048,plain,(
  ~sP11('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl91),
  inference(cnf_transformation,[],[f1048_D])).
fof(f1048_D,plain,(
  ~sP11('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f1006,plain,(
  ~sP11('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f445,f450,f441])).
fof(f1007,plain,(
  ~sP11('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f445,f450,f441])).
fof(f1008,plain,(
  ~sP11('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f445,f450,f441])).
fof(f1009,plain,(
  ~sP11('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f262,f450,f441])).
fof(f1010,plain,(
  ~sP11('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f262,f450,f441])).
fof(f1011,plain,(
  ~sP11('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f262,f450,f441])).
fof(f1012,plain,(
  ~sP11('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f445,f261,f450,f441])).
fof(f1013,plain,(
  ~sP11('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f262,f261,f450,f441])).
fof(f1014,plain,(
  ~sP11('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f261,f261,f450,f441])).
fof(f791,plain,(
  sP5('F_0','F_0','T_0') | $spl88),
  inference(cnf_transformation,[],[f791_D])).
fof(f791_D,plain,(
  sP5('F_0','F_0','T_0') <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f767,plain,(
  sP5('F_0','F_0','T_0') | $spl81),
  inference(unit_resulting_resolution,[],[f267,f763,f389])).
fof(f789,plain,(
  ~sP5('T_0','F_0','T_0') | $spl87),
  inference(cnf_transformation,[],[f789_D])).
fof(f789_D,plain,(
  ~sP5('T_0','F_0','T_0') <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f768,plain,(
  ~sP5('T_0','F_0','T_0') | $spl81),
  inference(unit_resulting_resolution,[],[f445,f763,f391])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f769,plain,(
  ~sP5('T_0','F_0','T_0') | $spl81),
  inference(unit_resulting_resolution,[],[f262,f763,f391])).
fof(f770,plain,(
  ~sP5('T_0','F_0','T_0') | $spl81),
  inference(unit_resulting_resolution,[],[f261,f763,f391])).
fof(f787,plain,(
  sP5('F_0','T_0','F_0') | $spl84),
  inference(cnf_transformation,[],[f787_D])).
fof(f787_D,plain,(
  sP5('F_0','T_0','F_0') <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f776,plain,(
  sP5('F_0','T_0','F_0') | $spl81),
  inference(unit_resulting_resolution,[],[f267,f763,f389])).
fof(f785,plain,(
  ~sP5('T_0','T_0','F_0') | $spl83),
  inference(cnf_transformation,[],[f785_D])).
fof(f785_D,plain,(
  ~sP5('T_0','T_0','F_0') <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f777,plain,(
  ~sP5('T_0','T_0','F_0') | $spl81),
  inference(unit_resulting_resolution,[],[f445,f763,f391])).
fof(f778,plain,(
  ~sP5('T_0','T_0','F_0') | $spl81),
  inference(unit_resulting_resolution,[],[f262,f763,f391])).
fof(f779,plain,(
  ~sP5('T_0','T_0','F_0') | $spl81),
  inference(unit_resulting_resolution,[],[f261,f763,f391])).
fof(f763,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl81),
  inference(cnf_transformation,[],[f763_D])).
fof(f763_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f759,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl76 | $spl79)),
  inference(backward_demodulation,[],[f755,f757])).
fof(f755,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl76),
  inference(cnf_transformation,[],[f755_D])).
fof(f755_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f758,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(global_subsumption,[],[f260,f259,f452,f650,f649,f648,f647,f646,f645,f644,f643,f642,f641,f640,f636,f633,f632,f631,f630,f629,f628,f627,f626,f625,f624,f623,f621,f620,f619,f618,f617,f613,f610,f609,f605,f602,f601,f600,f597,f596,f595,f594,f593,f592,f591,f590,f589,f588,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f544,f543,f542,f541,f540,f539,f538,f537,f536,f535,f534,f533,f532,f531,f528,f527,f526,f525,f524,f523,f522,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f507,f506,f505,f504,f503,f502,f501,f500,f499,f498,f497,f496,f495,f494,f493,f492,f491,f490,f489,f488,f487,f486,f485,f484,f483,f482,f481,f480,f479,f478,f477,f476,f475,f474,f473,f472,f468,f467,f466,f464,f463,f462,f458,f457,f456,f455,f453,f651,f652,f655,f660,f661,f662,f665,f672,f675,f676,f678,f679,f680,f681])).
fof(f681,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(superposition,[],[f452,f268])).
fof(f268,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.reflexive_TC',aHL_BOOL_CASES)).
fof(f680,plain,(
  ( ! [X39,X40] : (sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X39,X40) | ~p(s(bool,X39))) ) | $spl2),
  inference(resolution,[],[f452,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f258])).
fof(f679,plain,(
  ( ! [X37,X38] : (sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X38,X37) | p(s(bool,X37))) ) | $spl2),
  inference(resolution,[],[f452,f439])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f258])).
fof(f676,plain,(
  ( ! [X31,X32] : (sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X31,X32) | ~p(s(bool,X31))) ) | $spl2),
  inference(resolution,[],[f452,f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f250])).
fof(f675,plain,(
  ( ! [X30,X29] : (sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X30,X29) | ~p(s(bool,X29))) ) | $spl2),
  inference(resolution,[],[f452,f422])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f250])).
fof(f672,plain,(
  ( ! [X24,X23] : (sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X23,X24) | ~p(s(bool,X24)) | ~p(s(bool,X23))) ) | $spl2),
  inference(resolution,[],[f452,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f665,plain,(
  ( ! [X10] : (~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X10) | ~p(s(bool,X10))) ) | $spl2),
  inference(resolution,[],[f452,f365])).
fof(f365,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f219])).
fof(f660,plain,(
  ( ! [X5] : (~sP1(X5,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ~p(s(bool,X5))) ) | $spl2),
  inference(resolution,[],[f452,f358])).
fof(f358,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f655,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) ) | $spl2),
  inference(resolution,[],[f452,f337])).
fof(f472,plain,(
  ( ! [X0] : (sP0(X0,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f348])).
fof(f348,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'F_0')) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f473,plain,(
  ( ! [X0] : (sP0('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f349])).
fof(f349,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f522,plain,(
  ( ! [X0,X1] : (sP2('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f388])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f234])).
fof(f523,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)),happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f332,f452,f390])).
fof(f525,plain,(
  ( ! [X0] : (sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.combin.I_1'(s(bool,X0)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f452,f390])).
fof(f305,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ! [X20] : s(X6,'const.combin.I_1'(s(X6,X20))) = s(X6,happ(s(fun(X6,X6),'const.combin.I_0'),s(X6,X20)))),
  file('thm.relation.reflexive_TC','ARITY_0_1_const.combin.I')).
fof(f526,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)),'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f332,f452,f390])).
fof(f528,plain,(
  ( ! [X0] : (sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)),'const.combin.I_1'(s(bool,X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f452,f390])).
fof(f539,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f588,plain,(
  ( ! [X0] : (~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f407])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f242])).
fof(f589,plain,(
  ( ! [X0] : (~sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f242])).
fof(f601,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f418])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f248])).
fof(f602,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f248])).
fof(f605,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f422])).
fof(f609,plain,(
  ( ! [X0] : (sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0,'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f422])).
fof(f610,plain,(
  ( ! [X0] : (sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f452,f422])).
fof(f613,plain,(
  ( ! [X0] : (sP9('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f423])).
fof(f617,plain,(
  ( ! [X0] : (sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f423])).
fof(f618,plain,(
  ( ! [X0] : (sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f452,f423])).
fof(f620,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f250])).
fof(f621,plain,(
  ( ! [X0] : (~sP9('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f250])).
fof(f631,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f256])).
fof(f633,plain,(
  ( ! [X0] : (sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f439])).
fof(f636,plain,(
  ( ! [X0] : (sP11('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f440])).
fof(f640,plain,(
  ( ! [X0] : (sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f440])).
fof(f641,plain,(
  ( ! [X0] : (sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f452,f440])).
fof(f650,plain,(
  ( ! [X0] : (~sP11('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f652,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(resolution,[],[f452,f280])).
fof(f280,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f651,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(resolution,[],[f452,f278])).
fof(f278,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f166])).
fof(f455,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f278])).
fof(f456,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f280])).
fof(f757,plain,(
  s(bool,'F_0') != s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl79),
  inference(cnf_transformation,[],[f757_D])).
fof(f757_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f457,plain,(
  s(bool,'F_0') != s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f281])).
fof(f281,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f168])).
fof(f458,plain,(
  s(bool,'F_0') != s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f283])).
fof(f283,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f169])).
fof(f462,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f452,f337])).
fof(f463,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f452,f337])).
fof(f464,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f337])).
fof(f466,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f452,f337])).
fof(f467,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f452,f337])).
fof(f468,plain,(
  s(bool,'T_0') = s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f337])).
fof(f747,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl69),
  inference(cnf_transformation,[],[f747_D])).
fof(f747_D,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f474,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f452,f358])).
fof(f475,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f452,f358])).
fof(f476,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f358])).
fof(f745,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl67),
  inference(cnf_transformation,[],[f745_D])).
fof(f745_D,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f477,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f452,f358])).
fof(f749,plain,(
  ~sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl71),
  inference(cnf_transformation,[],[f749_D])).
fof(f749_D,plain,(
  ~sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f478,plain,(
  ~sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f452,f358])).
fof(f479,plain,(
  ~sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f452,f358])).
fof(f480,plain,(
  ~sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f358])).
fof(f481,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f452,f358])).
fof(f751,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl72),
  inference(cnf_transformation,[],[f751_D])).
fof(f751_D,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f482,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f359])).
fof(f753,plain,(
  sP1('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl74),
  inference(cnf_transformation,[],[f753_D])).
fof(f753_D,plain,(
  sP1('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f483,plain,(
  sP1('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f362])).
fof(f484,plain,(
  sP1('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f363])).
fof(f485,plain,(
  sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f364])).
fof(f486,plain,(
  ~sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f452,f365])).
fof(f487,plain,(
  ~sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f452,f365])).
fof(f488,plain,(
  ~sP1('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f365])).
fof(f489,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f452,f365])).
fof(f490,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f452,f365])).
fof(f491,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f452,f365])).
fof(f492,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f365])).
fof(f493,plain,(
  ~sP1('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f452,f365])).
fof(f743,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl65),
  inference(cnf_transformation,[],[f743_D])).
fof(f743_D,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f494,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f267,f452,f373])).
fof(f495,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f267,f452,f373])).
fof(f496,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f267,f452,f373])).
fof(f739,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl61),
  inference(cnf_transformation,[],[f739_D])).
fof(f739_D,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f497,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f267,f452,f373])).
fof(f741,plain,(
  ~sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl63),
  inference(cnf_transformation,[],[f741_D])).
fof(f741_D,plain,(
  ~sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f498,plain,(
  ~sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f267,f452,f373])).
fof(f499,plain,(
  ~sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f267,f452,f373])).
fof(f500,plain,(
  ~sP5('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f267,f452,f373])).
fof(f501,plain,(
  ~sP5('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f267,f452,f373])).
fof(f737,plain,(
  ~sP4('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl59),
  inference(cnf_transformation,[],[f737_D])).
fof(f737_D,plain,(
  ~sP4('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f502,plain,(
  ~sP4('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f267,f452,f377])).
fof(f503,plain,(
  ~sP4('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f267,f452,f377])).
fof(f504,plain,(
  ~sP4('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f267,f452,f377])).
fof(f733,plain,(
  ~sP4('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl55),
  inference(cnf_transformation,[],[f733_D])).
fof(f733_D,plain,(
  ~sP4('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f505,plain,(
  ~sP4('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f267,f452,f377])).
fof(f735,plain,(
  ~sP4('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl57),
  inference(cnf_transformation,[],[f735_D])).
fof(f735_D,plain,(
  ~sP4('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f506,plain,(
  ~sP4('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f267,f452,f377])).
fof(f507,plain,(
  ~sP4('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f267,f452,f377])).
fof(f508,plain,(
  ~sP4('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f267,f452,f377])).
fof(f509,plain,(
  ~sP4('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f267,f452,f377])).
fof(f731,plain,(
  ~sP3('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl53),
  inference(cnf_transformation,[],[f731_D])).
fof(f731_D,plain,(
  ~sP3('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f511,plain,(
  ~sP3('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f267,f452,f381])).
fof(f512,plain,(
  ~sP3('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f267,f452,f381])).
fof(f513,plain,(
  ~sP3('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f267,f452,f381])).
fof(f727,plain,(
  ~sP3('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl49),
  inference(cnf_transformation,[],[f727_D])).
fof(f727_D,plain,(
  ~sP3('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f514,plain,(
  ~sP3('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f267,f452,f381])).
fof(f729,plain,(
  ~sP3('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl51),
  inference(cnf_transformation,[],[f729_D])).
fof(f729_D,plain,(
  ~sP3('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f515,plain,(
  ~sP3('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f267,f452,f381])).
fof(f516,plain,(
  ~sP3('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f267,f452,f381])).
fof(f517,plain,(
  ~sP3('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f267,f452,f381])).
fof(f518,plain,(
  ~sP3('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f267,f452,f381])).
fof(f725,plain,(
  ~sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl47),
  inference(cnf_transformation,[],[f725_D])).
fof(f725_D,plain,(
  ~sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f531,plain,(
  ~sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f267,f452,f400])).
fof(f532,plain,(
  ~sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f267,f452,f400])).
fof(f533,plain,(
  ~sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f267,f452,f400])).
fof(f721,plain,(
  ~sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl43),
  inference(cnf_transformation,[],[f721_D])).
fof(f721_D,plain,(
  ~sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f534,plain,(
  ~sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f267,f452,f400])).
fof(f723,plain,(
  ~sP6('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl45),
  inference(cnf_transformation,[],[f723_D])).
fof(f723_D,plain,(
  ~sP6('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f535,plain,(
  ~sP6('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f267,f452,f400])).
fof(f536,plain,(
  ~sP6('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f267,f452,f400])).
fof(f537,plain,(
  ~sP6('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f267,f452,f400])).
fof(f538,plain,(
  ~sP6('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f267,f452,f400])).
fof(f719,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl40),
  inference(cnf_transformation,[],[f719_D])).
fof(f719_D,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f540,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f452,f406])).
fof(f541,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f445,f452,f406])).
fof(f542,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f445,f452,f406])).
fof(f711,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl32),
  inference(cnf_transformation,[],[f711_D])).
fof(f711_D,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f543,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f445,f452,f406])).
fof(f544,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f262,f452,f406])).
fof(f545,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f262,f452,f406])).
fof(f546,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f262,f452,f406])).
fof(f547,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f262,f452,f406])).
fof(f548,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f261,f452,f406])).
fof(f549,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f261,f452,f406])).
fof(f550,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f261,f452,f406])).
fof(f551,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f261,f452,f406])).
fof(f715,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl36),
  inference(cnf_transformation,[],[f715_D])).
fof(f715_D,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f552,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f452,f452,f406])).
fof(f553,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f452,f452,f406])).
fof(f554,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f452,f406])).
fof(f707,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl28),
  inference(cnf_transformation,[],[f707_D])).
fof(f707_D,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f555,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f452,f452,f406])).
fof(f717,plain,(
  sP7('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl38),
  inference(cnf_transformation,[],[f717_D])).
fof(f717_D,plain,(
  sP7('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f556,plain,(
  sP7('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f452,f406])).
fof(f557,plain,(
  sP7('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f445,f452,f406])).
fof(f558,plain,(
  sP7('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f445,f452,f406])).
fof(f559,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f445,f452,f406])).
fof(f560,plain,(
  sP7('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f262,f452,f406])).
fof(f561,plain,(
  sP7('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f262,f452,f406])).
fof(f562,plain,(
  sP7('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f262,f452,f406])).
fof(f563,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f262,f452,f406])).
fof(f564,plain,(
  sP7('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f261,f452,f406])).
fof(f565,plain,(
  sP7('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f261,f452,f406])).
fof(f566,plain,(
  sP7('T_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f261,f452,f406])).
fof(f567,plain,(
  sP7('T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f261,f452,f406])).
fof(f709,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl30),
  inference(cnf_transformation,[],[f709_D])).
fof(f709_D,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f568,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f452,f452,f406])).
fof(f569,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f452,f452,f406])).
fof(f570,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f452,f406])).
fof(f571,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f452,f452,f406])).
fof(f713,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl34),
  inference(cnf_transformation,[],[f713_D])).
fof(f713_D,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f572,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f452,f406])).
fof(f573,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f445,f452,f406])).
fof(f574,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f445,f452,f406])).
fof(f575,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f445,f452,f406])).
fof(f576,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f262,f452,f406])).
fof(f577,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f262,f452,f406])).
fof(f578,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f262,f452,f406])).
fof(f579,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f262,f452,f406])).
fof(f580,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f261,f452,f406])).
fof(f581,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f261,f452,f406])).
fof(f582,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f261,f261,f452,f406])).
fof(f583,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f452,f261,f452,f406])).
fof(f584,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f452,f452,f406])).
fof(f585,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f452,f452,f406])).
fof(f586,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f452,f452,f406])).
fof(f587,plain,(
  sP7('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f452,f452,f406])).
fof(f705,plain,(
  ~sP7('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl27),
  inference(cnf_transformation,[],[f705_D])).
fof(f705_D,plain,(
  ~sP7('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f590,plain,(
  ~sP7('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f445,f452,f409])).
fof(f591,plain,(
  ~sP7('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f262,f452,f409])).
fof(f592,plain,(
  ~sP7('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f261,f452,f409])).
fof(f701,plain,(
  ~sP7('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl23),
  inference(cnf_transformation,[],[f701_D])).
fof(f701_D,plain,(
  ~sP7('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f593,plain,(
  ~sP7('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f452,f409])).
fof(f703,plain,(
  ~sP7('F_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl25),
  inference(cnf_transformation,[],[f703_D])).
fof(f703_D,plain,(
  ~sP7('F_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f594,plain,(
  ~sP7('F_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f445,f452,f409])).
fof(f595,plain,(
  ~sP7('F_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f262,f452,f409])).
fof(f596,plain,(
  ~sP7('F_0','T_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f261,f452,f409])).
fof(f597,plain,(
  ~sP7('F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f452,f409])).
fof(f699,plain,(
  ~sP8('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','F_0') | $spl21),
  inference(cnf_transformation,[],[f699_D])).
fof(f699_D,plain,(
  ~sP8('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','F_0') <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f600,plain,(
  ~sP8('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f267,f452,f417])).
fof(f697,plain,(
  ~sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','F_0') | $spl19),
  inference(cnf_transformation,[],[f697_D])).
fof(f697_D,plain,(
  ~sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','F_0') <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f619,plain,(
  ~sP9('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f267,f452,f424])).
fof(f695,plain,(
  ~sP10('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') | $spl17),
  inference(cnf_transformation,[],[f695_D])).
fof(f695_D,plain,(
  ~sP10('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f623,plain,(
  ~sP10('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f445,f452,f434])).
fof(f624,plain,(
  ~sP10('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f262,f452,f434])).
fof(f625,plain,(
  ~sP10('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f261,f452,f434])).
fof(f691,plain,(
  ~sP10('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl13),
  inference(cnf_transformation,[],[f691_D])).
fof(f691_D,plain,(
  ~sP10('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f626,plain,(
  ~sP10('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f452,f434])).
fof(f693,plain,(
  ~sP10('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl15),
  inference(cnf_transformation,[],[f693_D])).
fof(f693_D,plain,(
  ~sP10('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f627,plain,(
  ~sP10('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f267,f452,f434])).
fof(f628,plain,(
  ~sP10('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f267,f452,f434])).
fof(f629,plain,(
  ~sP10('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f261,f267,f452,f434])).
fof(f630,plain,(
  ~sP10('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f452,f267,f452,f434])).
fof(f689,plain,(
  sP11('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl10),
  inference(cnf_transformation,[],[f689_D])).
fof(f689_D,plain,(
  sP11('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f632,plain,(
  sP11('F_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f267,f452,f438])).
fof(f687,plain,(
  ~sP11('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl9),
  inference(cnf_transformation,[],[f687_D])).
fof(f687_D,plain,(
  ~sP11('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f642,plain,(
  ~sP11('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f445,f452,f441])).
fof(f643,plain,(
  ~sP11('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f262,f452,f441])).
fof(f644,plain,(
  ~sP11('T_0','F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f261,f452,f441])).
fof(f683,plain,(
  ~sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl5),
  inference(cnf_transformation,[],[f683_D])).
fof(f683_D,plain,(
  ~sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f645,plain,(
  ~sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f452,f441])).
fof(f685,plain,(
  ~sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') | $spl7),
  inference(cnf_transformation,[],[f685_D])).
fof(f685_D,plain,(
  ~sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f646,plain,(
  ~sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f445,f452,f441])).
fof(f647,plain,(
  ~sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f262,f452,f441])).
fof(f648,plain,(
  ~sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f267,f261,f452,f441])).
fof(f649,plain,(
  ~sP11('const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)),'F_0','const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f267,f452,f452,f441])).
fof(f452,plain,(
  p(s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl2),
  inference(cnf_transformation,[],[f452_D])).
fof(f452_D,plain,(
  p(s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f259,plain,(
  p(s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))))),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  p(s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) & ~p(s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f122])).
fof(f122,plain,(
  ? [X0,X1] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & ~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ~! [X0,X1] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) => p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) => p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X6,fun(X6,bool)),X7)))) => p(s(bool,'const.relation.reflexive_1'(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X6,fun(X6,bool)),X7)))) => p(s(bool,'const.relation.reflexive_1'(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),X7)))))))),
  file('thm.relation.reflexive_TC','thm.relation.reflexive_TC')).
fof(f450,plain,(
  ~p(s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))))) | $spl1),
  inference(cnf_transformation,[],[f450_D])).
fof(f450_D,plain,(
  ~p(s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f260,plain,(
  ~p(s(bool,'const.relation.reflexive_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.TC_1'(s(fun(sK12,fun(sK12,bool)),sK13))))))),
  inference(cnf_transformation,[],[f163])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 312
% Passive clauses: 421
% Generated clauses: 3385
% Final active clauses: 309
% Final passive clauses: 240
% Input formulas: 46
% Initial clauses: 185
% 
% Duplicate literals: 56
% Trivial inequalities: 5
% Fw demodulations: 549
% Bw demodulations: 7
% Global subsumptions: 3
% 
% Simple tautologies: 185
% Equational tautologies: 9
% Forward subsumptions: 1627
% Backward subsumptions: 2
% 
% Binary resolution: 674
% Unit resulting resolution: 1743
% Forward superposition: 61
% Backward superposition: 155
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 1594
% SAT solver unit clauses: 4
% SAT solver binary clauses: 1385
% 
% Sat splits: 1
% 
% Memory used [KB]: 3837
% Time elapsed: 0.064 s
% ------------------------------
% ------------------------------
% Success in time 18.121 s
