Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.pred_set
% (110280)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 638
% Passive clauses: 9524
% Generated clauses: 32372
% Final active clauses: 473
% Final passive clauses: 6751
% Input formulas: 69
% Initial clauses: 184
% 
% Duplicate literals: 349
% Trivial inequalities: 11
% Fw subsumption resolutions: 1446
% Fw demodulations: 4077
% Bw demodulations: 109
% 
% Simple tautologies: 197
% Equational tautologies: 92
% Forward subsumptions: 15657
% Fw demodulations to eq. taut.: 8
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 12666
% Factoring: 821
% Forward superposition: 4590
% Backward superposition: 7642
% Self superposition: 53
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 982
% Split components: 1966
% SAT solver clauses: 1443
% SAT solver unit clauses: 5
% SAT solver binary clauses: 66
% 
% Sat splits: 982
% Sat splitting refutations: 14
% 
% TWLsolver clauses: 1443
% TWLsolver calls for satisfiability: 118
% 
% Memory used [KB]: 16119
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (110299)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 299
% Passive clauses: 2282
% Generated clauses: 14082
% Final active clauses: 261
% Final passive clauses: 1163
% Input formulas: 69
% Initial clauses: 184
% 
% Duplicate literals: 239
% Trivial inequalities: 6
% Fw subsumption resolutions: 421
% Fw demodulations: 1854
% Bw demodulations: 832
% Condensations: 299
% Global subsumptions: 9
% 
% Simple tautologies: 164
% Equational tautologies: 19
% Forward subsumptions: 4757
% Fw demodulations to eq. taut.: 19
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 2317
% Factoring: 393
% Forward superposition: 1993
% Backward superposition: 1670
% Self superposition: 32
% Equality factoring: 2
% Equality resolution: 4
% 
% Split clauses: 39
% Split components: 78
% SAT solver clauses: 2823
% SAT solver unit clauses: 208
% SAT solver binary clauses: 62
% SAT solver learnt clauses: 20
% SAT solver learnt literals: 3
% 
% Sat splits: 39
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 2785
% TWLsolver calls for satisfiability: 1446
% 
% Memory used [KB]: 14711
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f120874,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f120873])).
fof(f120873,plain,(
  s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11))))) != s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11)))))),
  inference(forward_demodulation,[],[f120868,f3613])).
fof(f3613,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))) )),
  inference(superposition,[],[f354,f503])).
fof(f503,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f501,f476])).
fof(f476,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(forward_demodulation,[],[f276,f266])).
fof(f266,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0')),
  inference(cnf_transformation,[],[f20])).
fof(f20,axiom,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0')),
  file('thm.pred_set.chooser_def_compute|split|1','thm.arithmetic.ALT_ZERO')).
fof(f276,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X14] : s('type.num.num',X14) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))),
  file('thm.pred_set.chooser_def_compute|split|1','thm.arithmetic.ADD_|48|')).
fof(f501,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(backward_demodulation,[],[f354,f482])).
fof(f482,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(forward_demodulation,[],[f282,f266])).
fof(f282,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X8] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  file('thm.pred_set.chooser_def_compute|split|1','thm.arithmetic.BIT|49|')).
fof(f354,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8,X14] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X8))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))),
  file('thm.pred_set.chooser_def_compute|split|1','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f120868,plain,(
  s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11))))) != s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),sK10),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK11),s('type.num.num',sK11)))))))),
  inference(backward_demodulation,[],[f120245,f4564])).
fof(f4564,plain,(
  s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11))))) != s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),sK10),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(superposition,[],[f471,f344])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) = s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))),s('type.num.num',X2)))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) = s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))),s('type.num.num',X2)))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) = s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))),s('type.num.num',X2)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X8] : s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) = s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),'const.pred_set.REST_1'(s(fun(X6,bool),X7))),s('type.num.num',X8)))),
  file('thm.pred_set.chooser_def_compute|split|1','thm.pred_set.chooser_def|split|1')).
fof(f471,plain,(
  s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11))))) != s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),'const.pred_set.REST_1'(s(fun(sK9,bool),sK10))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f470,f270])).
fof(f270,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X13] : s('type.num.num',X13) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13)))),
  file('thm.pred_set.chooser_def_compute|split|1','thm.arithmetic.NUMERAL_DEF')).
fof(f470,plain,(
  s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),sK10),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11))))))) != s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),'const.pred_set.REST_1'(s(fun(sK9,bool),sK10))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(backward_demodulation,[],[f270,f263])).
fof(f263,plain,(
  s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),sK10),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11))))))) != s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),'const.pred_set.REST_1'(s(fun(sK9,bool),sK10))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),sK10),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11))))))) != s(sK9,'const.pred_set.chooser_2'(s(fun(sK9,bool),'const.pred_set.REST_1'(s(fun(sK9,bool),sK10))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK11))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9,sK10,sK11])],[f153])).
fof(f153,plain,(
  ? [X0,X1,X2] : s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))))))) != s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2] : s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))))))) = s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2] : s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))))))) = s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8))))))) = s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),'const.pred_set.REST_1'(s(fun(X6,bool),X7))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8))))))) = s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),'const.pred_set.REST_1'(s(fun(X6,bool),X7))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.pred_set.chooser_def_compute|split|1','thm.pred_set.chooser_def_compute|split|1')).
fof(f120245,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f3457,f3613])).
fof(f3457,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f3386,f475])).
fof(f475,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(forward_demodulation,[],[f274,f266])).
fof(f274,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X14] : s('type.num.num',X14) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))),
  file('thm.pred_set.chooser_def_compute|split|1','thm.arithmetic.SUB_|48||split|1')).
fof(f3386,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f352,f1789])).
fof(f1789,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f503,f474])).
fof(f474,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f273,f266])).
fof(f273,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X14] : s('type.num.num',X14) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X14)))),
  file('thm.pred_set.chooser_def_compute|split|1','thm.arithmetic.ADD_CLAUSES|split|0')).
fof(f352,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X8,X14] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X14))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X14)))))),
  file('thm.pred_set.chooser_def_compute|split|1','thm.arithmetic.SUB_MONO_EQ')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1535
% Passive clauses: 9569
% Generated clauses: 122765
% Final active clauses: 1317
% Final passive clauses: 3308
% Input formulas: 69
% Initial clauses: 184
% Discarded non-redundant clauses: 2068
% 
% Duplicate literals: 304
% Trivial inequalities: 81
% Fw subsumption resolutions: 4510
% Fw demodulations: 26861
% Bw demodulations: 2760
% Condensations: 261
% Global subsumptions: 2564
% 
% Simple tautologies: 89
% Equational tautologies: 113
% Forward subsumptions: 71853
% Fw demodulations to eq. taut.: 107
% Bw demodulations to eq. taut.: 141
% 
% Unit resulting resolution: 23128
% Factoring: 1
% Forward superposition: 43779
% Backward superposition: 15384
% Self superposition: 90
% Equality factoring: 26
% Equality resolution: 8
% 
% Split clauses: 116
% Split components: 237
% SAT solver clauses: 17892
% SAT solver unit clauses: 3953
% SAT solver binary clauses: 7492
% 
% Sat splits: 116
% Sat splitting refutations: 5525
% 
% Memory used [KB]: 274665
% Time elapsed: 9.856 s
% ------------------------------
% ------------------------------
% Success in time 11.211 s
