Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (102351)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 197
% Passive clauses: 3171
% Generated clauses: 28969
% Final active clauses: 164
% Final passive clauses: 2605
% Input formulas: 36
% Initial clauses: 79
% 
% Duplicate literals: 995
% Trivial inequalities: 422
% Fw subsumption resolutions: 4675
% Fw demodulations: 5358
% Bw demodulations: 30
% 
% Simple tautologies: 418
% Equational tautologies: 49
% Forward subsumptions: 13006
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 1143
% Factoring: 2
% Forward superposition: 8518
% Backward superposition: 7310
% Self superposition: 167
% Equality factoring: 97
% Equality resolution: 31
% 
% Split clauses: 423
% Split components: 874
% SAT solver clauses: 1162
% SAT solver unit clauses: 21
% SAT solver binary clauses: 156
% 
% Sat splits: 423
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 1157
% TWLsolver calls for satisfiability: 233
% 
% Memory used [KB]: 11513
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (102363)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 268
% Passive clauses: 1848
% Generated clauses: 18430
% Final active clauses: 157
% Final passive clauses: 1006
% Input formulas: 36
% Initial clauses: 79
% 
% Duplicate literals: 273
% Trivial inequalities: 290
% Fw subsumption resolutions: 1917
% Fw demodulations: 2966
% Bw demodulations: 649
% Condensations: 442
% Global subsumptions: 9
% 
% Simple tautologies: 159
% Equational tautologies: 38
% Forward subsumptions: 8466
% Bw demodulations to eq. taut.: 10
% 
% Binary resolution: 1140
% Factoring: 9
% Forward superposition: 4443
% Backward superposition: 4157
% Self superposition: 66
% Equality factoring: 42
% Equality resolution: 18
% 
% Split clauses: 155
% Split components: 312
% SAT solver clauses: 2011
% SAT solver unit clauses: 87
% SAT solver binary clauses: 40
% SAT solver learnt clauses: 19
% SAT solver learnt literals: 3
% 
% Sat splits: 155
% Sat splitting refutations: 46
% 
% TWLsolver clauses: 1974
% TWLsolver calls for satisfiability: 459
% 
% Memory used [KB]: 12281
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f43880,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f133,f233,f231,f229,f132,f232,f131,f230,f228,f626,f631,f1146,f1152,f629,f1147,f1486,f1489,f1487,f1649,f1651,f1650,f2190,f2192,f2191,f2195,f2209,f2211,f2219,f2221,f2223,f2225,f2227,f2229,f2233,f2235,f2237,f2239,f2241,f2243,f2245,f2247,f2252,f2355,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2793,f2809,f2807,f2838,f2839,f2847,f2840,f2841,f2850,f2842,f2852,f2923,f2924,f2925,f2926,f2927,f2928,f2851,f2937,f2938,f2939,f2950,f2849,f2940,f2846,f2845,f3285,f3288,f3290,f3293,f3297,f3300,f3303,f3385,f3387,f3389,f3391,f3436,f3448,f4360,f4343,f4324,f4319,f4389,f4318,f4309,f4369,f4429,f4430,f4450,f4451,f4456,f4457,f4267,f4462,f4266,f4463,f4599,f4601,f4604,f4605,f4607,f4609,f4818,f4820,f4822,f4824,f4869,f6203,f6209,f9148,f9168,f9165,f9288,f13448,f13449,f13415,f13488,f13846,f13847,f13808,f13889,f14317,f14318,f14276,f14361,f14905,f14906,f14981,f14980,f15667,f15668,f15623,f15712,f17079,f17084,f17085,f17083,f17081,f17157,f17158,f17258,f17278,f17287,f17288,f19189,f19190,f19152,f19226,f9164,f19412,f19454,f19604,f19616,f19648,f19650,f19659,f19660,f19663,f19664,f19665,f19666,f19667,f19668,f19669,f19670,f19671,f19672,f19673,f19674,f19675,f19676,f19677,f19678,f19679,f19681,f19683,f19686,f19688,f19689,f19690,f19691,f19692,f19693,f19694,f19695,f19696,f19697,f19698,f19699,f19700,f19701,f19702,f19703,f19704,f19705,f19706,f19707,f19708,f19709,f19711,f19713,f19714,f19716,f19718,f19719,f19720,f19721,f19722,f19725,f19727,f19729,f19731,f19733,f19735,f19736,f19737,f19739,f19741,f19742,f19743,f19745,f19747,f19748,f19749,f19758,f19759,f19760,f19762,f19768,f19770,f19771,f19773,f19774,f19775,f19776,f19778,f19780,f19783,f19786,f19788,f19790,f19791,f19792,f19793,f19795,f19798,f19801,f19803,f19805,f19807,f19808,f19811,f19812,f19813,f19814,f19815,f19816,f19817,f19818,f19820,f19821,f19822,f19823,f19826,f19827,f19829,f19831,f19833,f19834,f19835,f19836,f19838,f19840,f19842,f19844,f19845,f19846,f19847,f19848,f19849,f19850,f19851,f19852,f19854,f19855,f19857,f19858,f19860,f19861,f19863,f19864,f19866,f19867,f19868,f19869,f19871,f19872,f19873,f19878,f19879,f19880,f19882,f19885,f19887,f19889,f19890,f19895,f19897,f19898,f19900,f19904,f19905,f19906,f19907,f19908,f19910,f19911,f19912,f19915,f19917,f19919,f19921,f19922,f19924,f19925,f19927,f19928,f19929,f19930,f19931,f19932,f19933,f19934,f19935,f19936,f19937,f19939,f19940,f19943,f19945,f19947,f19948,f19949,f19951,f19952,f19953,f19954,f19955,f19956,f19957,f19959,f19962,f19966,f19968,f19970,f19971,f19972,f19973,f19975,f19976,f19977,f19978,f19981,f19982,f19984,f19986,f19987,f19988,f19989,f19990,f19991,f19993,f19995,f19996,f19998,f20026,f20027,f20029,f20030,f20031,f20032,f20033,f20034,f20036,f20037,f20038,f20040,f20041,f20043,f20044,f20048,f20050,f20051,f20052,f20053,f20055,f20056,f20057,f20058,f20060,f20061,f20062,f20063,f20064,f20065,f20067,f20068,f20069,f20070,f20071,f20073,f20074,f20075,f20077,f20079,f20082,f20084,f20085,f20086,f20104,f20105,f20107,f20108,f20109,f20110,f20111,f20112,f20113,f20114,f20115,f20116,f20117,f20118,f20119,f20120,f20121,f20122,f20123,f20124,f20125,f20126,f20127,f20128,f20129,f20131,f20132,f20137,f20138,f20140,f20141,f20142,f20143,f20144,f20145,f20146,f20147,f20148,f20149,f20150,f20151,f20152,f20153,f20154,f20155,f20156,f20157,f20158,f20159,f20160,f20161,f20162,f20163,f20164,f20165,f20166,f20167,f20168,f20169,f20170,f20171,f20172,f20173,f20174,f20175,f20176,f20177,f20178,f20179,f20180,f20181,f20182,f20183,f20184,f20185,f20186,f20187,f20188,f20189,f20190,f20191,f20192,f20193,f20194,f20195,f20309,f20311,f20312,f20318,f20320,f20322,f20324,f20370,f20381,f20487,f20489,f20492,f20493,f20509,f20511,f20525,f20536,f20641,f20643,f20646,f20647,f20662,f20663,f20665,f20676,f20677,f20679,f4277,f4418,f4308,f4388,f20781,f20788,f20795,f5080,f20811,f9632,f20845,f20989,f20995,f20991,f20998,f20994,f20997,f21000,f21055,f21067,f21070,f21072,f21711,f21723,f22711,f22721,f23528,f23539,f23889,f23899,f25352,f25353,f25365,f25366,f1488,f26266,f26570,f26572,f26574,f26576,f26578,f26562,f26579,f26580,f26586,f26588,f26595,f26596,f26597,f26599,f26601,f26604,f26541,f26606,f26582,f26571,f26345,f26614,f26615,f628,f26569,f27166,f27169,f27298,f27300,f27303,f27304,f27310,f27312,f27313,f27315,f27317,f27392,f27396,f30114,f30133,f30136,f30088,f30051,f30137,f30045,f30226,f30134,f30042,f30227,f30030,f30128,f29958,f30419,f29957,f30421,f29956,f30427,f30422,f20678,f29951,f30428,f32774,f32775,f33155,f33156,f33654,f33655,f33917,f33933,f34491,f34492,f38199,f38200,f4368,f43654,f43703,f43714,f43804,f20796,f43805,f43810,f17281,f17279,f43863,f43876])).
fof(f43876,plain,(
  $false | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(subsumption_resolution,[],[f43875,f137])).
fof(f137,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.arithmetic.ONE_MOD_IFF',aHL_FALSITY)).
fof(f43875,plain,(
  p(s(bool,'F_0')) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(forward_demodulation,[],[f43783,f261])).
fof(f261,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f184,f150])).
fof(f150,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.bool.EQ_CLAUSES|split|3')).
fof(f184,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X6))))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.prim_rec.LESS_REFL')).
fof(f43783,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(backward_demodulation,[],[f43715,f39894])).
fof(f39894,plain,(
  ( ! [X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))))) )),
  inference(subsumption_resolution,[],[f39855,f134])).
fof(f134,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.ONE_MOD_IFF',aHL_TRUTH)).
fof(f39855,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))))) )),
  inference(superposition,[],[f534,f274])).
fof(f274,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f138,f146])).
fof(f146,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.bool.EQ_CLAUSES|split|0')).
fof(f138,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.prim_rec.LESS_SUC_REFL')).
fof(f534,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num',X5)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X5))))) )),
  inference(subsumption_resolution,[],[f522,f134])).
fof(f522,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num',X5)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X5))))) )),
  inference(superposition,[],[f195,f274])).
fof(f195,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(ennf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7,X6,X8] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X8))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8)))))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.arithmetic.LESS_TRANS')).
fof(f43715,plain,(
  ( ! [X45] : (s(bool,X45) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(subsumption_resolution,[],[f43712,f27505])).
fof(f27505,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') != s(bool,f(s(bool,X41))) | s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X42),s('type.num.num',X42)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f27463,f134])).
fof(f27463,plain,(
  ( ! [X41,X42] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,f(s(bool,X41))) | s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X42),s('type.num.num',X42)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f149,f12026])).
fof(f12026,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl24 | $spl30)),
  inference(condensation,[],[f12025])).
fof(f12025,plain,(
  ( ! [X37,X35,X36] : (s(bool,X35) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X37),s('type.num.num',X37))) | s(bool,'T_0') = s(bool,f(s(bool,X35))) | s(bool,'T_0') = s(bool,X36)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f11957,f141])).
fof(f141,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.arithmetic.ONE_MOD_IFF',aHL_BOOL_CASES)).
fof(f11957,plain,(
  ( ! [X37,X35,X36] : (s(bool,'F_0') != s(bool,X36) | s(bool,X35) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X37),s('type.num.num',X37))) | s(bool,'T_0') = s(bool,f(s(bool,X35))) | s(bool,'T_0') = s(bool,X36)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f10889,f300])).
fof(f300,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f141,f141])).
fof(f10889,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl24 | $spl30)),
  inference(forward_demodulation,[],[f10827,f261])).
fof(f10827,plain,(
  ( ! [X0,X1] : (s(bool,f(s(bool,X0))) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1))) | s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl24 | $spl30)),
  inference(equality_factoring,[],[f4252])).
fof(f4252,plain,(
  ( ! [X74,X75] : (s(bool,X75) = s(bool,f(s(bool,X75))) | s(bool,X75) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X74),s('type.num.num',X74)))) ) | ($spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f4217])).
fof(f4217,plain,(
  ( ! [X74,X75] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X75) = s(bool,f(s(bool,X75))) | s(bool,X75) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X74),s('type.num.num',X74)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f242,f2970])).
fof(f2970,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2966,f300])).
fof(f2966,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl24 | $spl30)),
  inference(backward_demodulation,[],[f2849,f2808])).
fof(f2808,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) | $spl24),
  inference(cnf_transformation,[],[f2808_D])).
fof(f2808_D,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f242,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f184,f145])).
fof(f145,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f100])).
fof(f149,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f102])).
fof(f43712,plain,(
  ( ! [X45] : (s(bool,X45) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X45)))) ) | ($spl80 | $spl156)),
  inference(backward_demodulation,[],[f17081,f38200])).
fof(f43863,plain,(
  $false | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(subsumption_resolution,[],[f43831,f242])).
fof(f43831,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(backward_demodulation,[],[f43822,f1298])).
fof(f1298,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))) )),
  inference(subsumption_resolution,[],[f1281,f134])).
fof(f1281,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))) )),
  inference(superposition,[],[f471,f280])).
fof(f280,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f139,f146])).
fof(f139,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.prim_rec.LESS_|48|')).
fof(f471,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(subsumption_resolution,[],[f462,f137])).
fof(f462,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f189,f141])).
fof(f189,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X7,X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.prim_rec.LESS_MONO')).
fof(f43822,plain,(
  ( ! [X10] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(subsumption_resolution,[],[f43821,f137])).
fof(f43821,plain,(
  ( ! [X10] : (p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(forward_demodulation,[],[f43741,f261])).
fof(f43741,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(backward_demodulation,[],[f43715,f37824])).
fof(f37824,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X10)),s('type.num.num',X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) )),
  inference(trivial_inequality_removal,[],[f37750])).
fof(f37750,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X10)),s('type.num.num',X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) )),
  inference(superposition,[],[f145,f402])).
fof(f402,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X0)),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f274,f142])).
fof(f142,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f43])).
fof(f43,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) | ? [X6] : s('type.num.num',X7) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.arithmetic.num_CASES')).
fof(f17279,plain,(
  ( ! [X118,X116,X117] : (s(fun(X116,bool),X117) = s(fun(X116,bool),X118)) ) | $spl86),
  inference(cnf_transformation,[],[f17279_D])).
fof(f17279_D,plain,(
  ( ! [X118,X116,X117] : (s(fun(X116,bool),X117) = s(fun(X116,bool),X118)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f17281,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl88),
  inference(cnf_transformation,[],[f17281_D])).
fof(f17281_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f43810,plain,(
  ( ! [X177,X175,X176] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(fun(X175,bool),X176) = s(fun(X175,bool),X177)) ) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(forward_demodulation,[],[f43809,f261])).
fof(f43809,plain,(
  ( ! [X177,X175,X176] : (s(bool,'T_0') = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) | s(fun(X175,bool),X176) = s(fun(X175,bool),X177)) ) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(forward_demodulation,[],[f43808,f43715])).
fof(f43808,plain,(
  ( ! [X177,X175,X176] : (s(fun(X175,bool),X176) = s(fun(X175,bool),X177) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X175,sK2(X176,X177,bool,X175))))))) ) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(subsumption_resolution,[],[f43723,f43715])).
fof(f43723,plain,(
  ( ! [X177,X175,X176] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(fun(X175,bool),X176) = s(fun(X175,bool),X177) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X175,sK2(X176,X177,bool,X175))))))) ) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(backward_demodulation,[],[f43715,f30773])).
fof(f30773,plain,(
  ( ! [X177,X175,X176] : (s(bool,'F_0') != s(bool,f(s(X175,sK2(X176,X177,bool,X175)))) | s(fun(X175,bool),X176) = s(fun(X175,bool),X177) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X175,sK2(X176,X177,bool,X175))))))) ) | ($spl24 | $spl30)),
  inference(forward_demodulation,[],[f22563,f261])).
fof(f22563,plain,(
  ( ! [X177,X175,X178,X176] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X178),s('type.num.num',X178))) != s(bool,f(s(X175,sK2(X176,X177,bool,X175)))) | s(fun(X175,bool),X176) = s(fun(X175,bool),X177) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X175,sK2(X176,X177,bool,X175))))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f194,f12026])).
fof(f194,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.arithmetic.ONE_MOD_IFF',aHL_EXT)).
fof(f43805,plain,(
  ( ! [X305] : (s(bool,X305) = s(bool,g(s(bool,X305)))) ) | $spl158),
  inference(cnf_transformation,[],[f43805_D])).
fof(f43805_D,plain,(
  ( ! [X305] : (s(bool,X305) = s(bool,g(s(bool,X305)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f20796,plain,(
  ( ! [X121,X122,X120] : (s(fun(X120,bool),X121) = s(fun(X120,bool),X122) | s(bool,'T_0') = s(bool,g(s(X120,sK2(X121,X122,bool,X120))))) ) | $spl98),
  inference(cnf_transformation,[],[f20796_D])).
fof(f20796_D,plain,(
  ( ! [X121,X122,X120] : (s(fun(X120,bool),X121) = s(fun(X120,bool),X122) | s(bool,'T_0') = s(bool,g(s(X120,sK2(X121,X122,bool,X120))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f43804,plain,(
  ( ! [X302,X304,X303,X305] : (s(fun(X302,bool),X303) = s(fun(X302,bool),X304) | s(bool,X305) = s(bool,g(s(bool,X305))) | s(bool,'T_0') = s(bool,g(s(X302,sK2(X303,X304,bool,X302))))) ) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(subsumption_resolution,[],[f43720,f30244])).
fof(f30244,plain,(
  ( ! [X72,X73] : (s(bool,X73) = s(bool,g(s(bool,X73))) | s(bool,X73) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X72),s('type.num.num',X72)))) ) | $spl30),
  inference(subsumption_resolution,[],[f29398,f240])).
fof(f240,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f137,f145])).
fof(f29398,plain,(
  ( ! [X72,X73] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X73) = s(bool,g(s(bool,X73))) | s(bool,X73) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X72),s('type.num.num',X72)))) ) | $spl30),
  inference(superposition,[],[f3068,f261])).
fof(f3068,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl30),
  inference(superposition,[],[f2849,f300])).
fof(f43720,plain,(
  ( ! [X302,X304,X303,X305] : (s(bool,X305) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(fun(X302,bool),X303) = s(fun(X302,bool),X304) | s(bool,X305) = s(bool,g(s(bool,X305))) | s(bool,'T_0') = s(bool,g(s(X302,sK2(X303,X304,bool,X302))))) ) | ($spl24 | $spl30 | $spl80 | $spl156)),
  inference(backward_demodulation,[],[f43715,f29727])).
fof(f29727,plain,(
  ( ! [X302,X304,X303,X305] : (s(bool,X305) != s(bool,f(s(X302,sK2(X303,X304,bool,X302)))) | s(fun(X302,bool),X303) = s(fun(X302,bool),X304) | s(bool,X305) = s(bool,g(s(bool,X305))) | s(bool,'T_0') = s(bool,g(s(X302,sK2(X303,X304,bool,X302))))) ) | $spl30),
  inference(superposition,[],[f194,f3068])).
fof(f43714,plain,(
  $false | ($spl2 | $spl80)),
  inference(subsumption_resolution,[],[f43713,f137])).
fof(f43713,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl80)),
  inference(forward_demodulation,[],[f43704,f261])).
fof(f43704,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl2 | $spl80)),
  inference(backward_demodulation,[],[f17081,f230])).
fof(f43703,plain,(
  $false | ($spl1 | $spl4 | $spl81)),
  inference(subsumption_resolution,[],[f43702,f183])).
fof(f183,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.arithmetic.SUC_NOT')).
fof(f43702,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl4 | $spl81)),
  inference(forward_demodulation,[],[f6212,f43612])).
fof(f43612,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',sK0)))) ) | ($spl1 | $spl81)),
  inference(backward_demodulation,[],[f43578,f140])).
fof(f140,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.arithmetic.MOD_ONE')).
fof(f43578,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl81)),
  inference(backward_demodulation,[],[f43577,f17790])).
fof(f17790,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(sK0)))) | $spl81),
  inference(unit_resulting_resolution,[],[f17080,f142])).
fof(f17080,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0) | $spl81),
  inference(cnf_transformation,[],[f17080_D])).
fof(f17080_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f43577,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1(sK0)) | ($spl1 | $spl81)),
  inference(subsumption_resolution,[],[f43576,f137])).
fof(f43576,plain,(
  p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1(sK0)) | ($spl1 | $spl81)),
  inference(forward_demodulation,[],[f43529,f21376])).
fof(f21376,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | ($spl1 | $spl81)),
  inference(forward_demodulation,[],[f21366,f261])).
fof(f21366,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) ) | ($spl1 | $spl81)),
  inference(unit_resulting_resolution,[],[f242,f19662,f340])).
fof(f340,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f330,f137])).
fof(f330,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | p(s(bool,X7)) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f190,f141])).
fof(f190,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.bool.IMP_ANTISYM_AX')).
fof(f19662,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) | ($spl1 | $spl81)),
  inference(global_subsumption,[],[f19661,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502])).
fof(f1502,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',sK1(X5))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))))),s('type.num.num',X5)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f459,f142])).
fof(f1414,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(subsumption_resolution,[],[f1403,f134])).
fof(f1403,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(superposition,[],[f454,f274])).
fof(f1979,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X2))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(subsumption_resolution,[],[f1967,f134])).
fof(f1967,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X2))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(superposition,[],[f517,f274])).
fof(f1400,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X5)),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f454,f142])).
fof(f1715,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(subsumption_resolution,[],[f1699,f134])).
fof(f1699,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f458,f274])).
fof(f1269,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X5)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f471,f142])).
fof(f1695,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X5))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f458,f142])).
fof(f1840,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X7)))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) )),
  inference(subsumption_resolution,[],[f1822,f134])).
fof(f1822,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X7)))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) )),
  inference(superposition,[],[f512,f274])).
fof(f1841,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X9)))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) )),
  inference(subsumption_resolution,[],[f1823,f134])).
fof(f1823,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X9)))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) )),
  inference(superposition,[],[f512,f280])).
fof(f1627,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1))))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(subsumption_resolution,[],[f1615,f134])).
fof(f1615,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1))))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(superposition,[],[f453,f274])).
fof(f1274,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X4)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f471,f142])).
fof(f1612,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X5)))),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f453,f142])).
fof(f1816,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X8)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X9)))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X7))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) )),
  inference(superposition,[],[f512,f142])).
fof(f136,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.arithmetic.ONE')).
fof(f235,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f234,f136])).
fof(f234,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0))))) )),
  inference(forward_demodulation,[],[f144,f136])).
fof(f144,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) => s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))) => s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6))))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.arithmetic.ONE_MOD')).
fof(f1144,plain,(
  ( ! [X50,X51] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X50))) | s(bool,X51) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X50))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f901,f134])).
fof(f901,plain,(
  ( ! [X50,X51] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X50))) | s(bool,X51) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X50))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f235,f300])).
fof(f662,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(subsumption_resolution,[],[f655,f134])).
fof(f655,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f235,f274])).
fof(f650,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(superposition,[],[f235,f142])).
fof(f2104,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s(bool,X2) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f2085,f134])).
fof(f2085,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s(bool,X2) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f650,f300])).
fof(f389,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f183,f142])).
fof(f453,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))),s('type.num.num',X5)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f189,f142])).
fof(f454,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X0)),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f189,f142])).
fof(f458,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f189,f142])).
fof(f459,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',sK1(X0))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f189,f142])).
fof(f477,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f189,f389])).
fof(f478,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num',X3)))),s('type.num.num',X4))))) )),
  inference(superposition,[],[f189,f389])).
fof(f480,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num',X6)))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f274,f389])).
fof(f535,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X7)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))) )),
  inference(subsumption_resolution,[],[f523,f134])).
fof(f523,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X7)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))) )),
  inference(superposition,[],[f195,f280])).
fof(f1299,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))) | s(bool,X6) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1282,f134])).
fof(f1282,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))) | s(bool,X6) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f471,f300])).
fof(f1415,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s(bool,X4) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X2)),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1404,f134])).
fof(f1404,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s(bool,X4) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X2)),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f454,f300])).
fof(f1515,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s(bool,X3) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',sK1(X2)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1505,f134])).
fof(f1505,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s(bool,X3) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',sK1(X2)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f459,f300])).
fof(f1628,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s(bool,X4) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X2)))),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1616,f134])).
fof(f1616,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s(bool,X4) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X2)))),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f453,f300])).
fof(f1712,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(forward_demodulation,[],[f1711,f140])).
fof(f1711,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(subsumption_resolution,[],[f1710,f134])).
fof(f1710,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(forward_demodulation,[],[f1693,f280])).
fof(f1693,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X2))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(superposition,[],[f458,f140])).
fof(f1716,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4) | s(bool,X5) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f1702,f134])).
fof(f1702,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4) | s(bool,X5) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f458,f300])).
fof(f512,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X7)))),s('type.num.num',X9)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) )),
  inference(superposition,[],[f195,f142])).
fof(f517,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X7)))),s('type.num.num',X9)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) )),
  inference(superposition,[],[f195,f142])).
fof(f1811,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X7))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X9)))),s('type.num.num',X8)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) )),
  inference(superposition,[],[f512,f142])).
fof(f1843,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X13)))),s('type.num.num',X12)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s(bool,X14) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num',X12))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f1825,f134])).
fof(f1825,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X13)))),s('type.num.num',X12)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s(bool,X14) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num',X12))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f512,f300])).
fof(f1964,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X8)))),s('type.num.num',X7)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X8)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X7))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X8) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) )),
  inference(superposition,[],[f517,f142])).
fof(f1980,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X5)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4) | s(bool,X7) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))),s('type.num.num',X5))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1968,f134])).
fof(f1968,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X5)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4) | s(bool,X7) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))),s('type.num.num',X5))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f517,f300])).
fof(f621,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(X3,sK2(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,g(s(X3,sK2(X4,X5,bool,X3))))) )),
  inference(superposition,[],[f194,f141])).
fof(f828,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,X84) != s(bool,f(s(X81,sK2(X82,X83,bool,X81)))) | s(fun(X81,bool),X82) = s(fun(X81,bool),X83) | s(bool,'T_0') = s(bool,X84) | s(bool,'T_0') = s(bool,g(s(X81,sK2(X82,X83,bool,X81))))) )),
  inference(superposition,[],[f194,f300])).
fof(f908,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'T_0') != s(bool,f(s(X72,sK2(X73,X74,bool,X72)))) | s(fun(X72,bool),X73) = s(fun(X72,bool),X74) | s(bool,X75) = s(bool,g(s(X72,sK2(X73,X74,bool,X72)))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f194,f300])).
fof(f14974,plain,(
  ( ! [X26,X27,X25] : (s(fun(X25,bool),X26) = s(fun(X25,bool),X27) | s(bool,'T_0') = s(bool,g(s(X25,sK2(X26,X27,bool,X25)))) | s(bool,'T_0') = s(bool,f(s(X25,sK2(X26,X27,bool,X25))))) )),
  inference(trivial_inequality_removal,[],[f14949])).
fof(f14949,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X25,bool),X26) = s(fun(X25,bool),X27) | s(bool,'T_0') = s(bool,g(s(X25,sK2(X26,X27,bool,X25)))) | s(bool,'T_0') = s(bool,f(s(X25,sK2(X26,X27,bool,X25))))) )),
  inference(superposition,[],[f621,f141])).
fof(f1145,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X53),s('type.num.num',X54)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X52),s('type.num.num',X54)))) | s(bool,X55) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X52),s('type.num.num',X53))) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f902,f134])).
fof(f902,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X53),s('type.num.num',X54)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X52),s('type.num.num',X54)))) | s(bool,X55) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X52),s('type.num.num',X53))) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f195,f300])).
fof(f7508,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,sK2(X5,X6,X7,bool))) | s(fun(bool,X7),X5) = s(fun(bool,X7),X6) | s(X7,f(s(bool,'F_0'))) != s(X7,g(s(bool,'F_0')))) )),
  inference(trivial_inequality_removal,[],[f7466])).
fof(f7466,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,sK2(X5,X6,X7,bool))) | s(fun(bool,X7),X5) = s(fun(bool,X7),X6) | s(X7,f(s(bool,'F_0'))) != s(X7,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f145,f619])).
fof(f7515,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,X28)) | s(bool,X28) = s(bool,sK2(X25,X26,X27,bool)) | s(fun(bool,X27),X25) = s(fun(bool,X27),X26) | s(X27,f(s(bool,'F_0'))) != s(X27,g(s(bool,'F_0')))) )),
  inference(subsumption_resolution,[],[f7472,f134])).
fof(f7472,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X28)) | s(bool,X28) = s(bool,sK2(X25,X26,X27,bool)) | s(fun(bool,X27),X25) = s(fun(bool,X27),X26) | s(X27,f(s(bool,'F_0'))) != s(X27,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f193,f619])).
fof(f7516,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X32) != s(bool,sK2(X29,X30,X31,bool)) | p(s(bool,X32)) | s(fun(bool,X31),X29) = s(fun(bool,X31),X30) | s(X31,f(s(bool,'F_0'))) != s(X31,g(s(bool,'F_0')))) )),
  inference(subsumption_resolution,[],[f7473,f134])).
fof(f7473,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'T_0')) | s(bool,X32) != s(bool,sK2(X29,X30,X31,bool)) | p(s(bool,X32)) | s(fun(bool,X31),X29) = s(fun(bool,X31),X30) | s(X31,f(s(bool,'F_0'))) != s(X31,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f221,f619])).
fof(f7500,plain,(
  ( ! [X127,X128,X126] : (s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(fun(bool,X128),X126) = s(fun(bool,X128),X127) | s(X128,f(s(bool,'F_0'))) != s(X128,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f7499])).
fof(f7499,plain,(
  ( ! [X127,X128,X126] : (s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(fun(bool,X128),X126) = s(fun(bool,X128),X127) | s(fun(bool,X128),X126) = s(fun(bool,X128),X127) | s(X128,f(s(bool,'F_0'))) != s(X128,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f194,f619])).
fof(f7514,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,sK2(X11,X12,X13,bool)) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(X13,f(s(bool,'F_0'))) != s(X13,g(s(bool,'F_0')))) )),
  inference(subsumption_resolution,[],[f7468,f134])).
fof(f7468,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,sK2(X11,X12,X13,bool)) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(X13,f(s(bool,'F_0'))) != s(X13,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f149,f619])).
fof(f619,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,sK2(X3,X4,X5,bool)) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(X5,f(s(bool,'F_0'))) != s(X5,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f194,f141])).
fof(f863,plain,(
  ( ! [X80,X78,X79,X77] : (s(X79,f(s(bool,X80))) != s(X79,g(s(bool,X80))) | s(fun(bool,X79),X77) = s(fun(bool,X79),X78) | s(bool,'T_0') = s(bool,sK2(X77,X78,X79,bool)) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f194,f300])).
fof(f260,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f150])).
fof(f622,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,X2,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f618,f260])).
fof(f618,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))))) != s(X2,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,X2,bool))) ) | $spl1),
  inference(superposition,[],[f194,f506])).
fof(f623,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0))))) ) | $spl1),
  inference(forward_demodulation,[],[f620,f260])).
fof(f620,plain,(
  ( ! [X2,X0,X1] : (s(bool,f(s(X0,sK2(X1,X2,bool,X0)))) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0))))) ) | $spl1),
  inference(superposition,[],[f194,f506])).
fof(f907,plain,(
  ( ! [X70,X68,X71,X69] : (s(X70,f(s(bool,'T_0'))) != s(X70,g(s(bool,'T_0'))) | s(fun(bool,X70),X68) = s(fun(bool,X70),X69) | s(bool,X71) = s(bool,sK2(X68,X69,X70,bool)) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f194,f300])).
fof(f1142,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f885,f134])).
fof(f885,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f193,f300])).
fof(f1141,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f881,f134])).
fof(f881,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f149,f300])).
fof(f867,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f145,f300])).
fof(f766,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f300,f300])).
fof(f624,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(bool,X0))) != s(X1,g(s(bool,X0))) | s(fun(bool,X1),X2) = s(fun(bool,X1),X3) | s(bool,'T_0') = s(bool,sK2(X2,X3,X1,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f622,f141])).
fof(f507,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f501,f137])).
fof(f501,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | s(bool,X3) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f340,f261])).
fof(f506,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(subsumption_resolution,[],[f500,f137])).
fof(f500,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,X1) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(superposition,[],[f340,f260])).
fof(f429,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f280,f142])).
fof(f241,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f145])).
fof(f237,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) | $spl1),
  inference(forward_demodulation,[],[f229,f136])).
fof(f221,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f220])).
fof(f220,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f209])).
fof(f209,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X14,X17,X15,X18] : ((s(bool,X14) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X15) = s(bool,X18))) => ((p(s(bool,X14)) => p(s(bool,X15))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.bool.IMP_CONG')).
fof(f193,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f88])).
fof(f19661,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) | $spl81),
  inference(forward_demodulation,[],[f18448,f17790])).
fof(f18448,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(sK0)))))) | $spl81),
  inference(superposition,[],[f235,f17790])).
fof(f43529,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1(sK0)) | $spl81),
  inference(superposition,[],[f1712,f17790])).
fof(f6212,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl4),
  inference(forward_demodulation,[],[f232,f136])).
fof(f43654,plain,(
  $false | ($spl1 | $spl81 | $spl82)),
  inference(subsumption_resolution,[],[f43653,f17080])).
fof(f43653,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | ($spl1 | $spl81 | $spl82)),
  inference(forward_demodulation,[],[f43622,f43612])).
fof(f43622,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK0),s('type.num.num',sK0))) | ($spl1 | $spl81 | $spl82)),
  inference(backward_demodulation,[],[f43578,f17803])).
fof(f17803,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | ($spl81 | $spl82)),
  inference(backward_demodulation,[],[f17790,f17083])).
fof(f4368,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,X26)) ) | $spl38),
  inference(cnf_transformation,[],[f4368_D])).
fof(f4368_D,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,X26)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f38200,plain,(
  ( ! [X45] : (s(bool,X45) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,f(s(bool,X45)))) ) | $spl156),
  inference(cnf_transformation,[],[f38200_D])).
fof(f38200_D,plain,(
  ( ! [X45] : (s(bool,X45) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,f(s(bool,X45)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f38199,plain,(
  ( ! [X45,X46] : (s(bool,X45) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,f(s(bool,X45))) | s(bool,'T_0') = s(bool,X46)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78 | $spl81)),
  inference(subsumption_resolution,[],[f37896,f31970])).
fof(f31970,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(superposition,[],[f31103,f141])).
fof(f31103,plain,(
  ( ! [X62] : (s(bool,'F_0') != s(bool,X62) | s(bool,'F_0') = s(bool,f(s(bool,X62)))) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(forward_demodulation,[],[f31045,f261])).
fof(f31045,plain,(
  ( ! [X62,X63] : (s(bool,'F_0') != s(bool,X62) | s(bool,f(s(bool,X62))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num',X63)))) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(superposition,[],[f11978,f30501])).
fof(f30501,plain,(
  ( ! [X53] : (s(bool,X53) = s(bool,f(s(bool,X53)))) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(global_subsumption,[],[f30500,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f4360,f4373,f4343,f4324,f4318,f4010,f4011,f4309,f4450,f4456,f4466,f4160,f4166,f4167,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f12066,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f17054,f19490,f19492,f621,f4412,f15668,f15712,f12026,f6127,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f471,f828,f14974,f10889,f6133,f5080,f5051,f5114,f5115,f19994,f20047,f19489,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f21711,f21727,f21894,f21895,f26883,f26884,f26975,f1715,f1712,f1627,f477,f1716,f458,f1695,f1414,f1841,f459,f1515,f26243,f1628,f1612,f453,f1502,f1400,f454,f1415,f478,f26976,f20042,f1980,f1979,f1964,f1843,f1840,f1816,f1811,f1299,f1274,f1269,f1145,f534,f517,f512,f195,f189,f26245,f20139,f4498,f2104,f1144,f650,f235,f26248,f26247,f26244,f26231,f535,f4226,f4161,f908,f26246,f11978,f7516,f7515,f7514,f7508,f7500,f4490,f619,f261,f150,f149,f137,f30114,f30088,f30051,f30045,f30226,f30238,f29360,f29362,f29363,f30030,f30244,f29956,f30427,f30434,f29726,f29727,f29732,f29735,f29736,f29741,f30473,f30479,f29887,f29888,f29902,f30486,f29891,f30488,f29738,f29737,f30499])).
fof(f30499,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') = s(bool,f(s(bool,X54))) | s(bool,X53) = s(bool,f(s(bool,X53))) | s(bool,X54) = s(bool,g(s(bool,X54)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f29617,f4466])).
fof(f29617,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') = s(bool,f(s(bool,X54))) | s(bool,X53) = s(bool,f(s(bool,X53))) | s(bool,X54) = s(bool,g(s(bool,X54))) | s(bool,'T_0') = s(bool,X53)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f3068])).
fof(f29737,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) != s(bool,X31) | s(bool,'T_0') = s(bool,f(s(bool,X31))) | s(bool,X32) = s(bool,g(s(bool,X32))) | s(bool,X30) = s(bool,X32)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f6117,f3068])).
fof(f29738,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) != s(bool,f(s(bool,X33))) | s(bool,'T_0') = s(bool,f(s(bool,X34))) | s(bool,X35) = s(bool,g(s(bool,X35))) | s(bool,X33) = s(bool,X35)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f6127,f3068])).
fof(f30488,plain,(
  ( ! [X70,X72] : (s(bool,'T_0') = s(bool,f(s(bool,X70))) | s(bool,X72) = s(bool,g(s(bool,X72))) | s(bool,X70) = s(bool,X72)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f29812,f242])).
fof(f29812,plain,(
  ( ! [X70,X72,X71] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X71),s('type.num.num',X71))) | s(bool,'T_0') = s(bool,f(s(bool,X70))) | s(bool,X72) = s(bool,g(s(bool,X72))) | s(bool,X70) = s(bool,X72)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f12026,f3068])).
fof(f29891,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl30),
  inference(equality_factoring,[],[f3068])).
fof(f30486,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1)))) ) | $spl30),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f2849,f3068,f863,f621,f471,f828,f14974,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f26883,f26884,f26975,f1715,f1712,f1627,f477,f1716,f458,f1695,f1414,f1841,f459,f1515,f26243,f1628,f1612,f453,f1502,f1400,f454,f1415,f478,f26976,f1980,f1979,f1964,f1843,f1840,f1816,f1811,f1299,f1274,f1269,f1145,f534,f517,f512,f195,f189,f26245,f2104,f1144,f650,f235,f26248,f26247,f26244,f26231,f535,f908,f26246,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f30114,f30088,f30051,f30045,f30238,f29360,f29363,f30030,f30244,f29726,f29727,f29732,f30473,f30479,f29887,f29888,f29902])).
fof(f29902,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(equality_factoring,[],[f3068])).
fof(f29888,plain,(
  ( ! [X265,X267,X264,X266] : (s(bool,'T_0') != s(bool,f(s(X264,sK2(X265,X266,bool,X264)))) | s(fun(X264,bool),X265) = s(fun(X264,bool),X266) | s(bool,X267) = s(bool,g(s(bool,X267))) | s(bool,X267) = s(bool,g(s(X264,sK2(X265,X266,bool,X264))))) ) | $spl30),
  inference(superposition,[],[f194,f3068])).
fof(f29887,plain,(
  ( ! [X261,X263,X260,X262] : (s(X262,f(s(bool,'T_0'))) != s(X262,g(s(bool,'T_0'))) | s(fun(bool,X262),X260) = s(fun(bool,X262),X261) | s(bool,X263) = s(bool,g(s(bool,X263))) | s(bool,X263) = s(bool,sK2(X260,X261,X262,bool))) ) | $spl30),
  inference(superposition,[],[f194,f3068])).
fof(f30479,plain,(
  ( ! [X231,X230] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X230))) | s(bool,X231) = s(bool,g(s(bool,X231))) | s(bool,X231) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X230)))) ) | $spl30),
  inference(subsumption_resolution,[],[f29877,f134])).
fof(f29877,plain,(
  ( ! [X231,X230] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X230))) | s(bool,X231) = s(bool,g(s(bool,X231))) | s(bool,X231) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X230)))) ) | $spl30),
  inference(superposition,[],[f235,f3068])).
fof(f30473,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,g(s(bool,X7))) | s(bool,X6) = s(bool,X7)) ) | $spl30),
  inference(subsumption_resolution,[],[f29788,f134])).
fof(f29788,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,g(s(bool,X7))) | s(bool,X6) = s(bool,X7)) ) | $spl30),
  inference(superposition,[],[f149,f3068])).
fof(f29741,plain,(
  ( ! [X43,X44,X42] : (s(bool,X43) != s(bool,f(s(bool,X42))) | s(bool,X43) = s(bool,f(s(bool,X43))) | s(bool,X44) = s(bool,g(s(bool,X44))) | s(bool,X42) = s(bool,X44)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f19994,f3068])).
fof(f29736,plain,(
  ( ! [X28,X29,X27] : (s(bool,X27) != s(bool,X28) | s(bool,X28) = s(bool,f(s(bool,X28))) | s(bool,X29) = s(bool,g(s(bool,X29))) | s(bool,X27) = s(bool,X29)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4466,f3068])).
fof(f29735,plain,(
  ( ! [X26,X24,X25] : (s(bool,X25) != s(bool,f(s(bool,X24))) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,g(s(bool,X26))) | s(bool,X24) = s(bool,X26)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4243,f3068])).
fof(f29732,plain,(
  ( ! [X17,X16] : (s(bool,X16) = s(bool,g(s(bool,X16))) | s(bool,X17) = s(bool,g(s(bool,X17))) | s(bool,X16) = s(bool,X17)) ) | $spl30),
  inference(superposition,[],[f2849,f3068])).
fof(f29726,plain,(
  ( ! [X300,X298,X301,X299] : (s(X300,f(s(bool,X301))) != s(X300,g(s(bool,X301))) | s(fun(bool,X300),X298) = s(fun(bool,X300),X299) | s(bool,X301) = s(bool,g(s(bool,X301))) | s(bool,'T_0') = s(bool,sK2(X298,X299,X300,bool))) ) | $spl30),
  inference(superposition,[],[f194,f3068])).
fof(f30434,plain,(
  ( ! [X182,X183,X181] : (s(bool,X181) != s(bool,X182) | s(bool,X182) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X183),s('type.num.num',X183))) | s(bool,X181) = s(bool,g(s(bool,X181)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f29672,f240])).
fof(f29672,plain,(
  ( ! [X182,X183,X181] : (s(bool,X181) != s(bool,X182) | s(bool,X182) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X183),s('type.num.num',X183))) | s(bool,X181) = s(bool,g(s(bool,X181))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f11978,f3068])).
fof(f29363,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,X13) | s(bool,X14) = s(bool,g(s(bool,X14))) | s(bool,X13) = s(bool,X14) | s(bool,X12) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl30),
  inference(superposition,[],[f3068,f300])).
fof(f29362,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X10) = s(bool,g(s(bool,X10))) | s(bool,X9) = s(bool,X10) | s(bool,X11) = s(bool,f(s(bool,X11))) | s(bool,X8) = s(bool,X11)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f3068,f2970])).
fof(f29360,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,g(s(bool,X2))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,g(s(bool,X3))) | s(bool,X0) = s(bool,X3)) ) | $spl30),
  inference(superposition,[],[f3068,f3068])).
fof(f30238,plain,(
  ( ! [X101] : (s(bool,X101) = s(bool,g(s(bool,X101))) | s(bool,'F_0') = s(bool,X101)) ) | $spl30),
  inference(subsumption_resolution,[],[f29341,f242])).
fof(f29341,plain,(
  ( ! [X101,X100] : (s(bool,'F_0') = s(bool,X101) | s(bool,X101) = s(bool,g(s(bool,X101))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X100),s('type.num.num',X100)))) ) | $spl30),
  inference(superposition,[],[f3068,f261])).
fof(f4490,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,X6) = s(bool,X7)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4188,f134])).
fof(f4188,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,X6) = s(bool,X7)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f149,f2970])).
fof(f26246,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) )),
  inference(subsumption_resolution,[],[f26179,f134])).
fof(f26179,plain,(
  ( ! [X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) )),
  inference(superposition,[],[f149,f429])).
fof(f4161,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,X123) != s(bool,f(s(X120,sK2(X121,X122,bool,X120)))) | s(fun(X120,bool),X121) = s(fun(X120,bool),X122) | s(bool,X123) = s(bool,f(s(bool,X123))) | s(bool,'T_0') = s(bool,g(s(X120,sK2(X121,X122,bool,X120))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f194,f2970])).
fof(f4226,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') != s(bool,f(s(X102,sK2(X103,X104,bool,X102)))) | s(fun(X102,bool),X103) = s(fun(X102,bool),X104) | s(bool,X105) = s(bool,f(s(bool,X105))) | s(bool,X105) = s(bool,g(s(X102,sK2(X103,X104,bool,X102))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f194,f2970])).
fof(f26231,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X11)) )),
  inference(trivial_inequality_removal,[],[f26177])).
fof(f26177,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X11)) )),
  inference(superposition,[],[f145,f429])).
fof(f26244,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X5)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(subsumption_resolution,[],[f26172,f134])).
fof(f26172,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X5)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f195,f429])).
fof(f26247,plain,(
  ( ! [X19,X20] : (~p(s(bool,X20)) | s(bool,X20) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19)) )),
  inference(subsumption_resolution,[],[f26183,f134])).
fof(f26183,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | s(bool,X20) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19)) )),
  inference(superposition,[],[f193,f429])).
fof(f26248,plain,(
  ( ! [X21,X22] : (s(bool,X22) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X21))) | p(s(bool,X22)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X21)) )),
  inference(subsumption_resolution,[],[f26184,f134])).
fof(f26184,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X21))) | p(s(bool,X22)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X21)) )),
  inference(superposition,[],[f221,f429])).
fof(f4498,plain,(
  ( ! [X78,X79] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X78))) | s(bool,X79) = s(bool,f(s(bool,X79))) | s(bool,X79) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X78)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4219,f134])).
fof(f4219,plain,(
  ( ! [X78,X79] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X78))) | s(bool,X79) = s(bool,f(s(bool,X79))) | s(bool,X79) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X78)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f235,f2970])).
fof(f20139,plain,(
  ( ! [X17,X18] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X17)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,X18) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X17)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f17053,f134])).
fof(f17053,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X17)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,X18) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X17)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f650,f2970])).
fof(f26245,plain,(
  ( ! [X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X9)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1(X9))) )),
  inference(subsumption_resolution,[],[f26175,f134])).
fof(f26175,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X9)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1(X9))) )),
  inference(superposition,[],[f459,f429])).
fof(f20042,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))))) | s(bool,X29) = s(bool,f(s(bool,X29))) | s(bool,X29) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num',X28)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f12065,f134])).
fof(f12065,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))))) | s(bool,X29) = s(bool,f(s(bool,X29))) | s(bool,X29) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num',X28)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f471,f2970])).
fof(f26976,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num',X5)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num',X4)))),s('type.num.num',X5))))) )),
  inference(subsumption_resolution,[],[f26898,f134])).
fof(f26898,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num',X5)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num',X4)))),s('type.num.num',X5))))) )),
  inference(superposition,[],[f195,f480])).
fof(f26243,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X3)))),s('type.num.num',X2)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(subsumption_resolution,[],[f26171,f134])).
fof(f26171,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X3)))),s('type.num.num',X2)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(superposition,[],[f512,f429])).
fof(f26975,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num',X2))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X3)))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3)) )),
  inference(subsumption_resolution,[],[f26897,f134])).
fof(f26897,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num',X2))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X3)))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3)) )),
  inference(superposition,[],[f512,f480])).
fof(f26884,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num',X1))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f480,f389])).
fof(f26883,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num',sK1(X0))))),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f480,f142])).
fof(f21895,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,f(s(bool,X3))) | s(bool,X4) = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f19994,f300])).
fof(f21894,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,f(s(bool,X0))) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,X2) = s(bool,f(s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f19994,f2970])).
fof(f21727,plain,(
  ( ! [X54,X53] : (s(bool,X53) != s(bool,f(s(fun(bool,bool),X54))) | s(bool,X53) = s(bool,f(s(bool,X53)))) ) | ($spl24 | $spl30 | $spl78)),
  inference(subsumption_resolution,[],[f21692,f242])).
fof(f21692,plain,(
  ( ! [X54,X52,X53] : (s(bool,X53) != s(bool,f(s(fun(bool,bool),X54))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X52),s('type.num.num',X52))) | s(bool,X53) = s(bool,f(s(bool,X53)))) ) | ($spl24 | $spl30 | $spl78)),
  inference(superposition,[],[f15712,f4252])).
fof(f19489,plain,(
  ( ! [X15,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X15))) | s(bool,X16) = s(bool,f(s(bool,X16))) | s(bool,X16) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X15)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f18461,f134])).
fof(f18461,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X15))) | s(bool,X16) = s(bool,f(s(bool,X16))) | s(bool,X16) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X15)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f235,f2970])).
fof(f20047,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) != s(bool,X29) | s(bool,X29) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X30))) | s(bool,X28) = s(bool,f(s(bool,X28)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f12112,f240])).
fof(f12112,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) != s(bool,X29) | s(bool,X29) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X30))) | s(bool,X28) = s(bool,f(s(bool,X28))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f11978,f2970])).
fof(f19994,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,f(s(bool,'T_0'))) | s(bool,X66) = s(bool,f(s(bool,X66)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10797,f242])).
fof(f10797,plain,(
  ( ! [X66,X65] : (s(bool,X66) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X65),s('type.num.num',X65))) | s(bool,X66) = s(bool,f(s(bool,X66)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4243,f4252])).
fof(f5115,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X4) = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4466,f300])).
fof(f5114,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,X2) = s(bool,f(s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4466,f2970])).
fof(f5051,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4243,f2970])).
fof(f6133,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'T_0') = s(bool,f(s(bool,X7)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f6057])).
fof(f6057,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'T_0') = s(bool,f(s(bool,X7))) | s(bool,'T_0') = s(bool,f(s(bool,X7)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f141,f4245])).
fof(f2917,plain,(
  ( ! [X0,X1] : (s(bool,f(s(bool,X0))) = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(superposition,[],[f2808,f300])).
fof(f4163,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X4) = s(bool,f(s(bool,X4))) | s(bool,X2) = s(bool,X4)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f145,f2970])).
fof(f4350,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,X0)))))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f4349])).
fof(f4349,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl24 | $spl30)),
  inference(condensation,[],[f3957])).
fof(f3957,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) = s(bool,X21) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,f(s(bool,X20))) = s(bool,f(s(bool,f(s(bool,X20))))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f2970])).
fof(f4413,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X26)) | s(bool,X26) = s(bool,f(s(bool,X25))) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4042,f340])).
fof(f4042,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) = s(bool,f(s(bool,X25))) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f193,f2970])).
fof(f4414,plain,(
  ( ! [X30,X28,X29] : (s(bool,X29) != s(bool,f(s(bool,X28))) | p(s(bool,X29)) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4043,f340])).
fof(f4043,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X28)) | s(bool,X29) != s(bool,f(s(bool,X28))) | p(s(bool,X29)) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f221,f2970])).
fof(f4420,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1)))) ) | ($spl24 | $spl30)),
  inference(condensation,[],[f4419])).
fof(f4419,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X10) = s(bool,X12) | s(bool,X10) = s(bool,f(s(bool,X10)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4061,f145])).
fof(f4061,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X10) = s(bool,X12) | s(bool,X10) = s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f190,f2970])).
fof(f4491,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,X16) = s(bool,X18)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4192,f134])).
fof(f4192,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,X16) = s(bool,X18)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f193,f2970])).
fof(f4499,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X81),s('type.num.num',X82)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X80),s('type.num.num',X82)))) | s(bool,X83) = s(bool,f(s(bool,X83))) | s(bool,X83) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X80),s('type.num.num',X81)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4220,f134])).
fof(f4220,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X81),s('type.num.num',X82)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X80),s('type.num.num',X82)))) | s(bool,X83) = s(bool,f(s(bool,X83))) | s(bool,X83) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X80),s('type.num.num',X81)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f195,f2970])).
fof(f4502,plain,(
  ( ! [X12,X11] : (p(s(bool,f(s(bool,X11)))) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4036,f300])).
fof(f4036,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,X11) | p(s(bool,f(s(bool,X11)))) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f145,f2970])).
fof(f4108,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,X123) != s(bool,f(s(X120,sK2(X121,X122,bool,X120)))) | s(fun(X120,bool),X121) = s(fun(X120,bool),X122) | s(bool,g(s(X120,sK2(X121,X122,bool,X120)))) = s(bool,f(s(bool,g(s(X120,sK2(X121,X122,bool,X120)))))) | s(bool,'T_0') = s(bool,X123)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f194,f2970])).
fof(f4107,plain,(
  ( ! [X118,X116,X119,X117] : (s(X118,f(s(bool,X119))) != s(X118,g(s(bool,X119))) | s(fun(bool,X118),X116) = s(fun(bool,X118),X117) | s(bool,sK2(X116,X117,X118,bool)) = s(bool,f(s(bool,sK2(X116,X117,X118,bool)))) | s(bool,'T_0') = s(bool,X119)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f194,f2970])).
fof(f20680,plain,(
  ( ! [X94] : (s(bool,X94) = s(bool,f(s(bool,X94))) | ~p(s(bool,X94))) ) | ($spl24 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f4360,f4373,f4343,f4324,f4318,f4010,f4011,f4309,f4450,f4456,f4466,f4160,f4166,f4167,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f4277,f4418,f4308,f4498,f4161,f4226,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f19710,f19712,f4253,f19715,f19717,f5080,f5051,f5114,f5115,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f19994,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f20042,f12066,f20047,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f20139,f17054,f19490,f19489,f19492,f621,f828,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f4412,f10889,f137,f149,f150,f261,f7508,f7514,f7515,f7516,f11978,f195,f1145,f1299,f12026,f6127,f6133,f619,f7500,f14974,f908,f4153,f20677,f4150])).
fof(f4150,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X94)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X93),s('type.num.num',X95)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X92),s('type.num.num',X95)))) | s(bool,X94) = s(bool,f(s(bool,X94))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X92),s('type.num.num',X93)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f195,f2970])).
fof(f4153,plain,(
  ( ! [X101,X102,X100] : (p(s(bool,X102)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X100),s('type.num.num',X101)))) | s(bool,X102) = s(bool,f(s(bool,X102))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X100))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f189,f2970])).
fof(f6127,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X39)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f6076])).
fof(f6076,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X39))) | s(bool,'T_0') = s(bool,f(s(bool,X39)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4243,f4245])).
fof(f4412,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') != s(bool,f(s(bool,X15))) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4038,f340])).
fof(f4038,plain,(
  ( ! [X15,X16] : (~p(s(bool,X15)) | s(bool,'F_0') != s(bool,f(s(bool,X15))) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f149,f2970])).
fof(f19492,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f19491,f137])).
fof(f19491,plain,(
  ( ! [X28,X29] : (p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))))) ) | ($spl24 | $spl30)),
  inference(forward_demodulation,[],[f18189,f261])).
fof(f18189,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X30)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f189,f4252])).
fof(f19490,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | s(bool,X14) = s(bool,f(s(bool,X14)))) ) | ($spl24 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2917,f4360,f4373,f4350,f4343,f4324,f4318,f4010,f4011,f4309,f4413,f4414,f4420,f4450,f4107,f4108,f4456,f4276,f4466,f4160,f4161,f4163,f4166,f4167,f4490,f4491,f4252,f4498,f4499,f4225,f4226,f4229,f4243,f4245,f4502,f4253,f4153,f863,f4412,f619,f261,f150,f149,f137,f5052,f6117,f6405,f7514,f7500,f7516,f7515,f7508,f1145,f195,f10889,f11978,f12105,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f19489,f18460])).
fof(f17054,plain,(
  ( ! [X19,X20] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num',X20)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X19)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X19))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X19)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f650,f4252])).
fof(f17052,plain,(
  ( ! [X15,X16] : (~p(s(bool,X16)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X15)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X15) | s(bool,X16) = s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X15)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f650,f2970])).
fof(f20136,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) = s(bool,f(s(bool,X2)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f20135,f183])).
fof(f20135,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X2)) | s(bool,X2) = s(bool,f(s(bool,X2)))) ) | ($spl24 | $spl30)),
  inference(forward_demodulation,[],[f20134,f140])).
fof(f20134,plain,(
  ( ! [X2] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X2)) | s(bool,X2) = s(bool,f(s(bool,X2)))) ) | ($spl24 | $spl30)),
  inference(forward_demodulation,[],[f20133,f389])).
fof(f20133,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | s(bool,X2) = s(bool,f(s(bool,X2)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f17042,f183])).
fof(f17042,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X2) = s(bool,f(s(bool,X2)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f650,f4252])).
fof(f15966,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X46),s('type.num.num',X46)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X47),s('type.num.num',X44)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X47),s('type.num.num',X45)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X44) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X44)))),s('type.num.num',X45))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X44)))),s('type.num.num',X45)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f517,f4252])).
fof(f20130,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X42),s('type.num.num',X40)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X42),s('type.num.num',X41)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40) | s(bool,X43) = s(bool,f(s(bool,X43))) | s(bool,X43) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X40)))),s('type.num.num',X41)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f15965,f134])).
fof(f15965,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X42),s('type.num.num',X40)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X42),s('type.num.num',X41)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40) | s(bool,X43) = s(bool,f(s(bool,X43))) | s(bool,X43) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X40)))),s('type.num.num',X41)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f517,f2970])).
fof(f15964,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X38)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num',X36)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num',X37)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X36) | s(bool,X38) = s(bool,f(s(bool,X38))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X36)))),s('type.num.num',X37)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f517,f2970])).
fof(f15954,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X4)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X3))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3) | s(bool,X4) = s(bool,f(s(bool,X4)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f517,f4252])).
fof(f15298,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X61),s('type.num.num',X61)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X62),s('type.num.num',X59)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X62)))),s('type.num.num',X60)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X62) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X60))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X60)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f512,f4252])).
fof(f20106,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X57),s('type.num.num',X55)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X57)))),s('type.num.num',X56)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X57) | s(bool,X58) = s(bool,f(s(bool,X58))) | s(bool,X58) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num',X56)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f15297,f134])).
fof(f15297,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X57),s('type.num.num',X55)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X57)))),s('type.num.num',X56)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X57) | s(bool,X58) = s(bool,f(s(bool,X58))) | s(bool,X58) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num',X56)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f512,f2970])).
fof(f15296,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X53)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X54),s('type.num.num',X51)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X54)))),s('type.num.num',X52)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X54) | s(bool,X53) = s(bool,f(s(bool,X53))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X51),s('type.num.num',X52)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f512,f2970])).
fof(f15277,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X4)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X5)))),s('type.num.num',X3)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5) | s(bool,X4) = s(bool,f(s(bool,X4)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f512,f4252])).
fof(f14957,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X59))) | s(fun(X56,bool),X57) = s(fun(X56,bool),X58) | s(bool,'T_0') = s(bool,g(s(X56,sK2(X57,X58,bool,X56)))) | s(bool,f(s(X56,sK2(X57,X58,bool,X56)))) = s(bool,f(s(bool,f(s(X56,sK2(X57,X58,bool,X56))))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f621,f4252])).
fof(f14987,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,f(s(X0,sK2(X1,X2,bool,X0)))) = s(bool,f(s(bool,f(s(X0,sK2(X1,X2,bool,X0))))))) ) | ($spl24 | $spl30)),
  inference(condensation,[],[f14986])).
fof(f14986,plain,(
  ( ! [X47,X45,X46,X44] : (s(fun(X44,bool),X45) = s(fun(X44,bool),X46) | s(bool,'T_0') = s(bool,g(s(X44,sK2(X45,X46,bool,X44)))) | s(bool,f(s(X44,sK2(X45,X46,bool,X44)))) = s(bool,f(s(bool,f(s(X44,sK2(X45,X46,bool,X44)))))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f14954,f141])).
fof(f14954,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X47) | s(fun(X44,bool),X45) = s(fun(X44,bool),X46) | s(bool,'T_0') = s(bool,g(s(X44,sK2(X45,X46,bool,X44)))) | s(bool,f(s(X44,sK2(X45,X46,bool,X44)))) = s(bool,f(s(bool,f(s(X44,sK2(X45,X46,bool,X44)))))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f621,f2970])).
fof(f14587,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num',X36)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X35)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X35) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X35)))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X35))))))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f458,f4252])).
fof(f20072,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X32)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X32) | s(bool,X33) = s(bool,f(s(bool,X33))) | s(bool,X33) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X31),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X32))))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f14586,f134])).
fof(f14586,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X32)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X32) | s(bool,X33) = s(bool,f(s(bool,X33))) | s(bool,X33) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X31),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X32))))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f458,f2970])).
fof(f14585,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X30)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X29) | s(bool,X30) = s(bool,f(s(bool,X30))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X29))))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f458,f2970])).
fof(f14575,plain,(
  ( ! [X6,X5] : (~p(s(bool,X6)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X5)))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5) | s(bool,X6) = s(bool,f(s(bool,X6)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f458,f4252])).
fof(f13986,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X33)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X32)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X31) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X31)))),s('type.num.num',X32))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X31)))),s('type.num.num',X32)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f453,f4252])).
fof(f20066,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28) | s(bool,X30) = s(bool,f(s(bool,X30))) | s(bool,X30) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X28)))),s('type.num.num',X29)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f13985,f134])).
fof(f13985,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28) | s(bool,X30) = s(bool,f(s(bool,X30))) | s(bool,X30) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X28)))),s('type.num.num',X29)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f453,f2970])).
fof(f13984,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X27)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s(bool,X27) = s(bool,f(s(bool,X27))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X25)))),s('type.num.num',X26)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f453,f2970])).
fof(f13974,plain,(
  ( ! [X2,X3] : (~p(s(bool,X3)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X2))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s(bool,X3) = s(bool,f(s(bool,X3)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f453,f4252])).
fof(f13156,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X30)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num',X29)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X29) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',sK1(X29)))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',sK1(X29))))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f459,f4252])).
fof(f20059,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num',X26)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s(bool,X27) = s(bool,f(s(bool,X27))) | s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num',sK1(X26))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f13155,f134])).
fof(f13155,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num',X26)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s(bool,X27) = s(bool,f(s(bool,X27))) | s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num',sK1(X26))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f459,f2970])).
fof(f13154,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X24)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.num.num',X23)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X23) | s(bool,X24) = s(bool,f(s(bool,X24))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num',sK1(X23))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f459,f2970])).
fof(f13146,plain,(
  ( ! [X2,X3] : (~p(s(bool,X3)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X2)))),s('type.num.num',X2)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s(bool,X3) = s(bool,f(s(bool,X3)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f459,f4252])).
fof(f12852,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X31),s('type.num.num',X31)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X29) | s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X29)),s('type.num.num',X30))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X29)),s('type.num.num',X30)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f454,f4252])).
fof(f20054,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s(bool,X28) = s(bool,f(s(bool,X28))) | s(bool,X28) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X26)),s('type.num.num',X27)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f12851,f134])).
fof(f12851,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s(bool,X28) = s(bool,f(s(bool,X28))) | s(bool,X28) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X26)),s('type.num.num',X27)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f454,f2970])).
fof(f12850,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X25)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X23) | s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X23)),s('type.num.num',X24)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f454,f2970])).
fof(f12841,plain,(
  ( ! [X2,X3] : (~p(s(bool,X3)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X2))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s(bool,X3) = s(bool,f(s(bool,X3)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f454,f4252])).
fof(f12156,plain,(
  ( ! [X94,X95,X93] : (s(bool,'F_0') != s(bool,X94) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X93),s('type.num.num',X93))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X95),s('type.num.num',X95))) | s(bool,X94) = s(bool,f(s(bool,X94)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f11978,f4252])).
fof(f20049,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X32))) | s(bool,X33) = s(bool,f(s(bool,X33))) | s(bool,'F_0') = s(bool,X33)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f12113,f507])).
fof(f12113,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,X31) | s(bool,X31) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X32))) | s(bool,X33) = s(bool,f(s(bool,X33))) | s(bool,'F_0') = s(bool,X33)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f11978,f2970])).
fof(f12066,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X32)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X31))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X31)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f471,f4252])).
fof(f12064,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X26)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))))) | s(bool,X26) = s(bool,f(s(bool,X26))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num',X25)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f471,f2970])).
fof(f20039,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,f(s(bool,X1)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f12048,f242])).
fof(f12048,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s(bool,X1) = s(bool,f(s(bool,X1)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f471,f4252])).
fof(f11962,plain,(
  ( ! [X52,X50,X51] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X51),s('type.num.num',X51))) | s(bool,X50) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X52),s('type.num.num',X52))) | s(bool,f(s(bool,X50))) = s(bool,f(s(bool,f(s(bool,X50)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f10889,f4252])).
fof(f11960,plain,(
  ( ! [X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | s(bool,X44) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X46),s('type.num.num',X46))) | s(bool,X45) = s(bool,f(s(bool,X45))) | s(bool,'T_0') = s(bool,f(s(bool,X44)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f10889,f2970])).
fof(f20035,plain,(
  ( ! [X43,X41,X42] : (s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X43),s('type.num.num',X43))) | s(bool,f(s(bool,X41))) = s(bool,f(s(bool,f(s(bool,X41))))) | s(bool,'T_0') = s(bool,X42)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f11959,f141])).
fof(f11959,plain,(
  ( ! [X43,X41,X42] : (s(bool,'F_0') != s(bool,X42) | s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X43),s('type.num.num',X43))) | s(bool,f(s(bool,X41))) = s(bool,f(s(bool,f(s(bool,X41))))) | s(bool,'T_0') = s(bool,X42)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f10889,f2970])).
fof(f12024,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl24 | $spl30)),
  inference(condensation,[],[f12023])).
fof(f12023,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num',X34))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,f(s(bool,X32)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f11956,f141])).
fof(f11956,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') != s(bool,X33) | s(bool,X32) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num',X34))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,f(s(bool,X32)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f10889,f300])).
fof(f11975,plain,(
  ( ! [X28,X27] : (s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X28))) | s(bool,'T_0') = s(bool,f(s(bool,X27)))) ) | ($spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f11954])).
fof(f11954,plain,(
  ( ! [X28,X27] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X28))) | s(bool,'T_0') = s(bool,f(s(bool,X27)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f10889,f141])).
fof(f11943,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,X5) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,f(s(bool,X5)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f10889,f4245])).
fof(f11930,plain,(
  ( ! [X103,X102] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X102),s('type.num.num',X102))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X103),s('type.num.num',X103)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f10889,f261])).
fof(f11929,plain,(
  ( ! [X101,X99,X100] : (s(bool,'F_0') != s(bool,f(s(bool,X100))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X99),s('type.num.num',X99))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X101),s('type.num.num',X101))) | s(bool,X100) = s(bool,f(s(bool,X100)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f10889,f4252])).
fof(f20025,plain,(
  ( ! [X52,X50] : (s(bool,'F_0') != s(bool,f(s(bool,X50))) | s(bool,X52) = s(bool,f(s(bool,X52))) | s(bool,X50) = s(bool,X52)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f11907,f242])).
fof(f11907,plain,(
  ( ! [X52,X50,X51] : (s(bool,'F_0') != s(bool,f(s(bool,X50))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X51),s('type.num.num',X51))) | s(bool,X52) = s(bool,f(s(bool,X52))) | s(bool,X50) = s(bool,X52)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f10889,f2970])).
fof(f10829,plain,(
  ( ! [X68,X67] : (s(bool,X68) = s(bool,f(s(bool,X68))) | s(bool,'T_0') = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X67),s('type.num.num',X67)))))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f10798])).
fof(f10798,plain,(
  ( ! [X68,X67] : (s(bool,X68) = s(bool,f(s(bool,X68))) | s(bool,'T_0') = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X67),s('type.num.num',X67))))) | s(bool,X68) = s(bool,f(s(bool,X68)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f4252])).
fof(f10786,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X34)) | s(bool,X35) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X33))) | p(s(bool,X35)) | s(bool,X34) = s(bool,f(s(bool,X34)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f221,f4252])).
fof(f10785,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X30))) | s(bool,X31) = s(bool,f(s(bool,X31)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f193,f4252])).
fof(f10784,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num',X27))) | s(bool,X28) = s(bool,f(s(bool,X28)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f190,f4252])).
fof(f10783,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num',X24))) | s(bool,X25) = s(bool,f(s(bool,X25)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f190,f4252])).
fof(f10781,plain,(
  ( ! [X21,X20] : (~p(s(bool,X21)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num',X20))) | s(bool,X21) = s(bool,f(s(bool,X21)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f149,f4252])).
fof(f19985,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | s(bool,X19) = s(bool,f(s(bool,X19)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10780,f242])).
fof(f10780,plain,(
  ( ! [X19,X18] : (~p(s(bool,X19)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num',X18))) | s(bool,X19) = s(bool,f(s(bool,X19)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f146,f4252])).
fof(f19983,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,X15) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10778,f242])).
fof(f10778,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num',X14))) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f141,f4252])).
fof(f10777,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num',X12)))) | s(bool,X13) = s(bool,f(s(bool,X13)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f189,f4252])).
fof(f19980,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) = s(bool,f(s(bool,X7)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f19979,f183])).
fof(f19979,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X7)) | s(bool,X7) = s(bool,f(s(bool,X7)))) ) | ($spl24 | $spl30)),
  inference(forward_demodulation,[],[f10774,f140])).
fof(f10774,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X7) = s(bool,f(s(bool,X7)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f235,f4252])).
fof(f10772,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,X4) | s(bool,X4) = s(bool,f(s(bool,X4)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f261,f4252])).
fof(f10770,plain,(
  ( ! [X222,X220,X219,X221] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X222),s('type.num.num',X222))) != s(bool,f(s(X219,sK2(X220,X221,bool,X219)))) | s(fun(X219,bool),X220) = s(fun(X219,bool),X221) | s(bool,g(s(X219,sK2(X220,X221,bool,X219)))) = s(bool,f(s(bool,g(s(X219,sK2(X220,X221,bool,X219))))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f194,f4252])).
fof(f10769,plain,(
  ( ! [X218,X216,X215,X217] : (s(X217,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X218),s('type.num.num',X218))))) != s(X217,g(s(bool,'const.prim_rec.<_2'(s('type.num.num',X218),s('type.num.num',X218))))) | s(fun(bool,X217),X215) = s(fun(bool,X217),X216) | s(bool,sK2(X215,X216,X217,bool)) = s(bool,f(s(bool,sK2(X215,X216,X217,bool))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f194,f4252])).
fof(f19974,plain,(
  ( ! [X208,X209,X207] : (s(fun(bool,X209),X207) = s(fun(bool,X209),X208) | s(X209,f(s(bool,'F_0'))) != s(X209,g(s(bool,'F_0'))) | s(bool,sK2(X207,X208,X209,bool)) = s(bool,f(s(bool,sK2(X207,X208,X209,bool))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10768,f242])).
fof(f10768,plain,(
  ( ! [X210,X208,X209,X207] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X210),s('type.num.num',X210))) | s(fun(bool,X209),X207) = s(fun(bool,X209),X208) | s(X209,f(s(bool,'F_0'))) != s(X209,g(s(bool,'F_0'))) | s(bool,sK2(X207,X208,X209,bool)) = s(bool,f(s(bool,sK2(X207,X208,X209,bool))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f619,f4252])).
fof(f10760,plain,(
  ( ! [X194,X195,X193] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X195),s('type.num.num',X195)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X193),s('type.num.num',X194)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X193))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X194))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X193))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X194)))))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f189,f4252])).
fof(f10759,plain,(
  ( ! [X191,X189,X192,X190] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X191),s('type.num.num',X191)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X190),s('type.num.num',X192)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X189),s('type.num.num',X192)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X189),s('type.num.num',X190))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X189),s('type.num.num',X190)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f195,f4252])).
fof(f10757,plain,(
  ( ! [X185,X186] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X186),s('type.num.num',X186)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X185))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X185))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X185)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f235,f4252])).
fof(f19969,plain,(
  ( ! [X156] : (s(bool,X156) = s(bool,f(s(bool,X156))) | s(bool,f(s(bool,X156))) = s(bool,f(s(bool,f(s(bool,X156)))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10748,f242])).
fof(f10748,plain,(
  ( ! [X156,X157] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X157),s('type.num.num',X157))) | s(bool,X156) = s(bool,f(s(bool,X156))) | s(bool,f(s(bool,X156))) = s(bool,f(s(bool,f(s(bool,X156)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f4252])).
fof(f19967,plain,(
  ( ! [X154,X155] : (s(bool,X154) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X155),s('type.num.num',X155))) | s(bool,f(s(bool,X154))) = s(bool,f(s(bool,f(s(bool,X154)))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10747,f4466])).
fof(f10747,plain,(
  ( ! [X154,X155] : (s(bool,X154) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X155),s('type.num.num',X155))) | s(bool,'T_0') = s(bool,f(s(bool,X154))) | s(bool,f(s(bool,X154))) = s(bool,f(s(bool,f(s(bool,X154)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f4252])).
fof(f19965,plain,(
  ( ! [X152,X151,X153] : (s(bool,X151) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X152),s('type.num.num',X152))) | s(bool,'T_0') = s(bool,X153) | s(bool,f(s(bool,X151))) = s(bool,f(s(bool,f(s(bool,X151)))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10746,f12105])).
fof(f10746,plain,(
  ( ! [X152,X151,X153] : (s(bool,X151) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X152),s('type.num.num',X152))) | s(bool,X151) = s(bool,X153) | s(bool,'T_0') = s(bool,X153) | s(bool,f(s(bool,X151))) = s(bool,f(s(bool,f(s(bool,X151)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f4252])).
fof(f10833,plain,(
  ( ! [X56,X55] : (s(bool,'T_0') = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X56),s('type.num.num',X56))))) | s(bool,X55) = s(bool,f(s(bool,X55)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f10717])).
fof(f10717,plain,(
  ( ! [X56,X55] : (s(bool,'T_0') = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X56),s('type.num.num',X56))))) | s(bool,X55) = s(bool,f(s(bool,X55))) | s(bool,X55) = s(bool,f(s(bool,X55)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f4252])).
fof(f10704,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num',X20)))) | s(bool,X19) != s(bool,X21) | p(s(bool,X21)) | s(bool,X19) = s(bool,f(s(bool,X19)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f221,f4252])).
fof(f10703,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X17),s('type.num.num',X17)))) | ~p(s(bool,X18)) | s(bool,X16) = s(bool,X18) | s(bool,X16) = s(bool,f(s(bool,X16)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f193,f4252])).
fof(f10702,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num',X14)))) | p(s(bool,X15)) | s(bool,X13) = s(bool,X15) | s(bool,X13) = s(bool,f(s(bool,X13)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f190,f4252])).
fof(f10701,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num',X11)))) | p(s(bool,X12)) | s(bool,X10) = s(bool,X12) | s(bool,X10) = s(bool,f(s(bool,X10)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f190,f4252])).
fof(f10700,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X9)))) | s(bool,'F_0') = s(bool,X8) | s(bool,X8) = s(bool,f(s(bool,X8)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f150,f4252])).
fof(f10699,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X7)))) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,f(s(bool,X6)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f149,f4252])).
fof(f19950,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X5)))) | s(bool,X4) = s(bool,f(s(bool,X4)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10698,f4466])).
fof(f10698,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X5)))) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,f(s(bool,X4)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f146,f4252])).
fof(f10871,plain,(
  ( ! [X127,X126] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X126))))) | s(bool,X126) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X127),s('type.num.num',X127)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10688,f507])).
fof(f10688,plain,(
  ( ! [X127,X126] : (s(bool,'T_0') != s(bool,X126) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X126))))) | s(bool,X126) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X127),s('type.num.num',X127)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f6117,f4252])).
fof(f19946,plain,(
  ( ! [X80,X81] : (s(bool,f(s(bool,X80))) = s(bool,f(s(bool,f(s(bool,X80))))) | s(bool,X80) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X81),s('type.num.num',X81)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10669,f507])).
fof(f10669,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,X80) | s(bool,f(s(bool,X80))) = s(bool,f(s(bool,f(s(bool,X80))))) | s(bool,X80) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X81),s('type.num.num',X81)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4466,f4252])).
fof(f19944,plain,(
  ( ! [X78,X79] : (s(bool,f(s(bool,X78))) = s(bool,f(s(bool,f(s(bool,X78))))) | s(bool,X78) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num',X79)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10668,f4466])).
fof(f10668,plain,(
  ( ! [X78,X79] : (s(bool,'T_0') = s(bool,f(s(bool,X78))) | s(bool,f(s(bool,X78))) = s(bool,f(s(bool,f(s(bool,X78))))) | s(bool,X78) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num',X79)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f4252])).
fof(f19942,plain,(
  ( ! [X74,X75] : (s(bool,'T_0') = s(bool,f(s(bool,X74))) | s(bool,X74) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X75),s('type.num.num',X75)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f19941,f507])).
fof(f19941,plain,(
  ( ! [X74,X75] : (s(bool,'T_0') != s(bool,X74) | s(bool,'T_0') = s(bool,f(s(bool,X74))) | s(bool,X74) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X75),s('type.num.num',X75)))) ) | ($spl24 | $spl30)),
  inference(forward_demodulation,[],[f10666,f2966])).
fof(f10666,plain,(
  ( ! [X74,X75] : (s(bool,X74) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X74))) | s(bool,X74) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X75),s('type.num.num',X75)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4243,f4252])).
fof(f10664,plain,(
  ( ! [X70,X68,X69] : (s(bool,X68) = s(bool,X69) | s(bool,X69) = s(bool,f(s(bool,X69))) | s(bool,'T_0') = s(bool,f(s(bool,X68))) | s(bool,X68) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X70)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f4252])).
fof(f19938,plain,(
  ( ! [X66,X67,X65] : (s(bool,f(s(bool,X65))) = s(bool,f(s(bool,f(s(bool,X65))))) | s(bool,'T_0') = s(bool,X66) | s(bool,X65) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X67),s('type.num.num',X67)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10663,f12105])).
fof(f10663,plain,(
  ( ! [X66,X67,X65] : (s(bool,X65) = s(bool,X66) | s(bool,f(s(bool,X65))) = s(bool,f(s(bool,f(s(bool,X65))))) | s(bool,'T_0') = s(bool,X66) | s(bool,X65) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X67),s('type.num.num',X67)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f4252])).
fof(f10660,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,X57)) | s(bool,X58) = s(bool,f(s(bool,X57))) | s(bool,'T_0') = s(bool,X58) | s(bool,X57) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X59)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f340,f4252])).
fof(f10655,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,X42)) | s(bool,X43) != s(bool,f(s(bool,X42))) | p(s(bool,X43)) | s(bool,X42) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X44),s('type.num.num',X44)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f221,f4252])).
fof(f10654,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) = s(bool,f(s(bool,X39))) | s(bool,X39) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X41),s('type.num.num',X41)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f193,f4252])).
fof(f10653,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X37) = s(bool,f(s(bool,X36))) | s(bool,X36) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X38),s('type.num.num',X38)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f190,f4252])).
fof(f10652,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) = s(bool,f(s(bool,X33))) | s(bool,X33) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num',X35)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f190,f4252])).
fof(f10651,plain,(
  ( ! [X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,f(s(bool,X31))) | s(bool,X31) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X32)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f150,f4252])).
fof(f10650,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,'F_0') != s(bool,f(s(bool,X29))) | s(bool,X29) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X30)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f149,f4252])).
fof(f10649,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | s(bool,'T_0') = s(bool,f(s(bool,X27))) | s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X28)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f146,f4252])).
fof(f10851,plain,(
  ( ! [X26,X25] : (p(s(bool,f(s(bool,X25)))) | s(bool,X25) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num',X26)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10648,f507])).
fof(f10648,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,X25) | p(s(bool,f(s(bool,X25)))) | s(bool,X25) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num',X26)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f145,f4252])).
fof(f10647,plain,(
  ( ! [X24,X23] : (s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,f(s(bool,X23))) | s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num',X24)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f141,f4252])).
fof(f19926,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,X6) | s(bool,X6) = s(bool,f(s(bool,X6)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10619,f242])).
fof(f10619,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') = s(bool,X6) | s(bool,X6) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X5)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f141])).
fof(f10618,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,X4) | s(bool,X4) = s(bool,f(s(bool,X4)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f261])).
fof(f19923,plain,(
  ( ! [X103,X105,X104] : (s(bool,sK2(X103,X104,X105,bool)) = s(bool,f(s(bool,sK2(X103,X104,X105,bool)))) | s(fun(bool,X105),X103) = s(fun(bool,X105),X104) | s(X105,f(s(bool,'F_0'))) != s(X105,g(s(bool,'F_0')))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10612,f242])).
fof(f10612,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X106),s('type.num.num',X106))) | s(bool,sK2(X103,X104,X105,bool)) = s(bool,f(s(bool,sK2(X103,X104,X105,bool)))) | s(fun(bool,X105),X103) = s(fun(bool,X105),X104) | s(X105,f(s(bool,'F_0'))) != s(X105,g(s(bool,'F_0')))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f619])).
fof(f10835,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1))) | s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,X0)))))) ) | ($spl24 | $spl30)),
  inference(condensation,[],[f10600])).
fof(f10600,plain,(
  ( ! [X70,X71,X69] : (s(bool,X69) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X70))) | s(bool,f(s(bool,X69))) = s(bool,f(s(bool,f(s(bool,X69))))) | s(bool,X69) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X71),s('type.num.num',X71)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f4252])).
fof(f19920,plain,(
  ( ! [X67] : (s(bool,f(s(bool,X67))) = s(bool,f(s(bool,f(s(bool,X67))))) | s(bool,X67) = s(bool,f(s(bool,X67)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10599,f242])).
fof(f10599,plain,(
  ( ! [X68,X67] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X68),s('type.num.num',X68))) | s(bool,f(s(bool,X67))) = s(bool,f(s(bool,f(s(bool,X67))))) | s(bool,X67) = s(bool,f(s(bool,X67)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f4245])).
fof(f19918,plain,(
  ( ! [X66,X65] : (s(bool,X65) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X66),s('type.num.num',X66))) | s(bool,f(s(bool,X65))) = s(bool,f(s(bool,f(s(bool,X65)))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10598,f4466])).
fof(f10598,plain,(
  ( ! [X66,X65] : (s(bool,X65) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X66),s('type.num.num',X66))) | s(bool,f(s(bool,X65))) = s(bool,f(s(bool,f(s(bool,X65))))) | s(bool,'T_0') = s(bool,f(s(bool,X65)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f4245])).
fof(f19916,plain,(
  ( ! [X64,X62,X63] : (s(bool,X62) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num',X63))) | s(bool,f(s(bool,X62))) = s(bool,f(s(bool,f(s(bool,X62))))) | s(bool,'T_0') = s(bool,X64)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10597,f12105])).
fof(f10597,plain,(
  ( ! [X64,X62,X63] : (s(bool,X62) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num',X63))) | s(bool,f(s(bool,X62))) = s(bool,f(s(bool,f(s(bool,X62))))) | s(bool,X62) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f2970])).
fof(f10836,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num',X24))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num',X25))) | s(bool,X23) = s(bool,f(s(bool,X23)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f10586])).
fof(f10586,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num',X24))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num',X25))) | s(bool,X23) = s(bool,f(s(bool,X23))) | s(bool,X23) = s(bool,f(s(bool,X23)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f4252])).
fof(f10838,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1))) | s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,X0)))))) ) | ($spl24 | $spl30)),
  inference(condensation,[],[f10568])).
fof(f10568,plain,(
  ( ! [X43,X41,X42] : (s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X42),s('type.num.num',X42))) | s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X43),s('type.num.num',X43))) | s(bool,f(s(bool,X41))) = s(bool,f(s(bool,f(s(bool,X41)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f4252])).
fof(f10566,plain,(
  ( ! [X37,X35,X36] : (s(bool,X35) = s(bool,X36) | s(bool,X35) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X37),s('type.num.num',X37))) | s(bool,X36) = s(bool,f(s(bool,X36))) | s(bool,'T_0') = s(bool,f(s(bool,X35)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f2970])).
fof(f19909,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num',X34))) | s(bool,f(s(bool,X32))) = s(bool,f(s(bool,f(s(bool,X32))))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10565,f12105])).
fof(f10565,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) = s(bool,X33) | s(bool,X32) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num',X34))) | s(bool,f(s(bool,X32))) = s(bool,f(s(bool,f(s(bool,X32))))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f2970])).
fof(f10560,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') = s(bool,X18) | s(bool,X18) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X19))) | s(bool,'T_0') = s(bool,f(s(bool,X18)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f141])).
fof(f10546,plain,(
  ( ! [X99,X98] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X98),s('type.num.num',X98))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X99),s('type.num.num',X99)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f261])).
fof(f10839,plain,(
  ( ! [X97,X95,X96] : (s(bool,X96) = s(bool,f(s(bool,X96))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X95),s('type.num.num',X95))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X97),s('type.num.num',X97)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f10545])).
fof(f10545,plain,(
  ( ! [X97,X95,X96] : (s(bool,X96) = s(bool,f(s(bool,X96))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X95),s('type.num.num',X95))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X97),s('type.num.num',X97))) | s(bool,X96) = s(bool,f(s(bool,X96)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f4252])).
fof(f19903,plain,(
  ( ! [X64,X63] : (s(bool,X63) = s(bool,f(s(bool,X63))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X64),s('type.num.num',X64)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10529,f4466])).
fof(f10529,plain,(
  ( ! [X64,X63] : (s(bool,X63) = s(bool,f(s(bool,X63))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X64),s('type.num.num',X64))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4252,f141])).
fof(f19899,plain,(
  ( ! [X14,X13] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X13))) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X14) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X13)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10452,f134])).
fof(f10452,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X13))) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X14) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X13)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f235,f2970])).
fof(f10451,plain,(
  ( ! [X12,X11] : (~p(s(bool,X12)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X11))) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X11)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f235,f2970])).
fof(f19896,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num',X35)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X35)))) | s(bool,X36) = s(bool,f(s(bool,X36))) | s(bool,X36) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X34)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10206,f134])).
fof(f10206,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num',X35)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X35)))) | s(bool,X36) = s(bool,f(s(bool,X36))) | s(bool,X36) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X34)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f195,f2970])).
fof(f10205,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X31)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X32)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X29),s('type.num.num',X32)))) | s(bool,X31) = s(bool,f(s(bool,X31))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X29),s('type.num.num',X30)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f195,f2970])).
fof(f9954,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num',X19)))) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f189,f2970])).
fof(f9975,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num',X16)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16)))))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f9953,f145])).
fof(f9953,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num',X16)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f189,f2970])).
fof(f19870,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s(bool,X24) = s(bool,f(s(bool,X24))) | s(bool,X24) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X22)))),s('type.num.num',X23)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f8085,f134])).
fof(f8085,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s(bool,X24) = s(bool,f(s(bool,X24))) | s(bool,X24) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X22)))),s('type.num.num',X23)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f453,f2970])).
fof(f8084,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X21)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s(bool,X21) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X19)))),s('type.num.num',X20)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f453,f2970])).
fof(f7501,plain,(
  ( ! [X114,X115,X113] : (s(bool,'T_0') = s(bool,f(s(bool,sK2(X113,X114,X115,bool)))) | s(fun(bool,X115),X113) = s(fun(bool,X115),X114) | s(X115,f(s(bool,'F_0'))) != s(X115,g(s(bool,'F_0')))) ) | ($spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f7495])).
fof(f7495,plain,(
  ( ! [X114,X115,X113] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(bool,sK2(X113,X114,X115,bool)))) | s(fun(bool,X115),X113) = s(fun(bool,X115),X114) | s(X115,f(s(bool,'F_0'))) != s(X115,g(s(bool,'F_0')))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f6117,f619])).
fof(f7502,plain,(
  ( ! [X80,X78,X79] : (s(bool,sK2(X78,X79,X80,bool)) = s(bool,f(s(bool,sK2(X78,X79,X80,bool)))) | s(fun(bool,X80),X78) = s(fun(bool,X80),X79) | s(X80,f(s(bool,'F_0'))) != s(X80,g(s(bool,'F_0')))) ) | ($spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f7484])).
fof(f7484,plain,(
  ( ! [X80,X78,X79] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,sK2(X78,X79,X80,bool)) = s(bool,f(s(bool,sK2(X78,X79,X80,bool)))) | s(fun(bool,X80),X78) = s(fun(bool,X80),X79) | s(X80,f(s(bool,'F_0'))) != s(X80,g(s(bool,'F_0')))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4466,f619])).
fof(f19865,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X65) = s(bool,f(s(bool,X65))) | s(bool,'T_0') = s(bool,sK2(X62,X63,X64,bool)) | s(fun(bool,X64),X62) = s(fun(bool,X64),X63) | s(X64,f(s(bool,'F_0'))) != s(X64,g(s(bool,'F_0')))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f7480,f4466])).
fof(f7480,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,X65) | s(bool,X65) = s(bool,f(s(bool,X65))) | s(bool,'T_0') = s(bool,sK2(X62,X63,X64,bool)) | s(fun(bool,X64),X62) = s(fun(bool,X64),X63) | s(X64,f(s(bool,'F_0'))) != s(X64,g(s(bool,'F_0')))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f619])).
fof(f7503,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'T_0') = s(bool,X61) | s(bool,sK2(X58,X59,X60,bool)) = s(bool,f(s(bool,sK2(X58,X59,X60,bool)))) | s(fun(bool,X60),X58) = s(fun(bool,X60),X59) | s(X60,f(s(bool,'F_0'))) != s(X60,g(s(bool,'F_0')))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f7479])).
fof(f7479,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'T_0') = s(bool,X61) | s(bool,sK2(X58,X59,X60,bool)) = s(bool,f(s(bool,sK2(X58,X59,X60,bool)))) | s(bool,'T_0') = s(bool,X61) | s(fun(bool,X60),X58) = s(fun(bool,X60),X59) | s(X60,f(s(bool,'F_0'))) != s(X60,g(s(bool,'F_0')))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f619])).
fof(f19862,plain,(
  ( ! [X28,X26,X29,X27] : (s(fun(bool,X28),X26) = s(fun(bool,X28),X27) | s(X28,f(s(bool,'F_0'))) != s(X28,g(s(bool,'F_0'))) | s(bool,X29) = s(bool,f(s(bool,X29))) | s(bool,'T_0') = s(bool,sK2(X26,X27,X28,bool))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f7460,f4466])).
fof(f7460,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,X29) | s(fun(bool,X28),X26) = s(fun(bool,X28),X27) | s(X28,f(s(bool,'F_0'))) != s(X28,g(s(bool,'F_0'))) | s(bool,X29) = s(bool,f(s(bool,X29))) | s(bool,'T_0') = s(bool,sK2(X26,X27,X28,bool))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f619,f2970])).
fof(f7509,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X25) | s(fun(bool,X24),X22) = s(fun(bool,X24),X23) | s(X24,f(s(bool,'F_0'))) != s(X24,g(s(bool,'F_0'))) | s(bool,sK2(X22,X23,X24,bool)) = s(bool,f(s(bool,sK2(X22,X23,X24,bool))))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f7459])).
fof(f7459,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X25) | s(fun(bool,X24),X22) = s(fun(bool,X24),X23) | s(X24,f(s(bool,'F_0'))) != s(X24,g(s(bool,'F_0'))) | s(bool,sK2(X22,X23,X24,bool)) = s(bool,f(s(bool,sK2(X22,X23,X24,bool)))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f619,f2970])).
fof(f19859,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.num.num',X22)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s(bool,X23) = s(bool,f(s(bool,X23))) | s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num',sK1(X22))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6940,f134])).
fof(f6940,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.num.num',X22)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s(bool,X23) = s(bool,f(s(bool,X23))) | s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num',sK1(X22))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f459,f2970])).
fof(f6939,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X20)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.num.num',X19)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num',sK1(X19))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f459,f2970])).
fof(f19856,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s(bool,X24) = s(bool,f(s(bool,X24))) | s(bool,X24) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X22)),s('type.num.num',X23)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6745,f134])).
fof(f6745,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s(bool,X24) = s(bool,f(s(bool,X24))) | s(bool,X24) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X22)),s('type.num.num',X23)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f454,f2970])).
fof(f6744,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X21)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s(bool,X21) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X19)),s('type.num.num',X20)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f454,f2970])).
fof(f19853,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X23))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))))) | s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,X25) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X23),s('type.num.num',X24)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6515,f134])).
fof(f6515,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X23))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))))) | s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,X25) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X23),s('type.num.num',X24)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f471,f2970])).
fof(f6514,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X22)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))))) | s(bool,X22) = s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num',X21)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f471,f2970])).
fof(f6483,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X29))))) | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6433,f300])).
fof(f6433,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X29))))) | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f6117,f2970])).
fof(f6479,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X25)))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6429,f6117])).
fof(f6429,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X25))))) | s(bool,'T_0') = s(bool,f(s(bool,X25)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f6117,f4245])).
fof(f6452,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X24))))) | s(bool,X24) = s(bool,f(s(bool,X24)))) ) | ($spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f6428])).
fof(f6428,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X24))))) | s(bool,X24) = s(bool,f(s(bool,X24)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f6117,f4245])).
fof(f6462,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') = s(bool,f(s(bool,X13))) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X13) = s(bool,X14)) ) | ($spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f6421])).
fof(f6421,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(bool,X13))) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X13) = s(bool,X14)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f6117,f2970])).
fof(f6404,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,X1))) | s(bool,X2) = s(bool,f(s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f6117,f2970])).
fof(f6332,plain,(
  ( ! [X38] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X38),s('type.num.num',X38))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X38),s('type.num.num',X38)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f261])).
fof(f6331,plain,(
  ( ! [X37] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X37),s('type.num.num',X37)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f261])).
fof(f19843,plain,(
  ( ! [X33] : (s(bool,'F_0') = s(bool,X33) | s(bool,X33) = s(bool,f(s(bool,X33)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6328,f242])).
fof(f6328,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') = s(bool,X33) | s(bool,X33) = s(bool,f(s(bool,X33))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X32)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f261])).
fof(f19839,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,X13) = s(bool,f(s(bool,X13)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6304,f242])).
fof(f6304,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,X13) = s(bool,f(s(bool,X13))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num',X12)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f261,f2970])).
fof(f19837,plain,(
  ( ! [X14,X13] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X13))) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X14) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X13)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6199,f134])).
fof(f6199,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X13))) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X14) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X13)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f235,f2970])).
fof(f6198,plain,(
  ( ! [X12,X11] : (~p(s(bool,X12)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X11))) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X11)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f235,f2970])).
fof(f6119,plain,(
  ( ! [X38] : (s(bool,f(s(bool,X38))) = s(bool,f(s(bool,f(s(bool,X38))))) | s(bool,X38) = s(bool,f(s(bool,X38)))) ) | ($spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f6108])).
fof(f6108,plain,(
  ( ! [X38] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,f(s(bool,X38))) = s(bool,f(s(bool,f(s(bool,X38))))) | s(bool,X38) = s(bool,f(s(bool,X38)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4466,f4245])).
fof(f6120,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,X33) | s(bool,f(s(bool,X32))) = s(bool,f(s(bool,f(s(bool,X32))))) | s(bool,X32) = s(bool,f(s(bool,X32)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f6105])).
fof(f6105,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,X33) | s(bool,f(s(bool,X32))) = s(bool,f(s(bool,f(s(bool,X32))))) | s(bool,'T_0') = s(bool,X33) | s(bool,X32) = s(bool,f(s(bool,X32)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f4245])).
fof(f19832,plain,(
  ( ! [X19,X20] : (s(bool,X20) != s(bool,f(s(bool,X19))) | p(s(bool,X20)) | s(bool,X19) = s(bool,f(s(bool,X19)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6099,f134])).
fof(f6099,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | s(bool,X20) != s(bool,f(s(bool,X19))) | p(s(bool,X20)) | s(bool,X19) = s(bool,f(s(bool,X19)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f221,f4245])).
fof(f19830,plain,(
  ( ! [X17,X18] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,f(s(bool,X17))) | s(bool,X17) = s(bool,f(s(bool,X17)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6098,f134])).
fof(f6098,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | s(bool,X18) = s(bool,f(s(bool,X17))) | s(bool,X17) = s(bool,f(s(bool,X17)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f193,f4245])).
fof(f19828,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,f(s(bool,X11))) | s(bool,X11) = s(bool,f(s(bool,X11)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6094,f134])).
fof(f6094,plain,(
  ( ! [X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,f(s(bool,X11))) | s(bool,X11) = s(bool,f(s(bool,X11)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f149,f4245])).
fof(f6125,plain,(
  ( ! [X9] : (p(s(bool,f(s(bool,X9)))) | s(bool,X9) = s(bool,f(s(bool,X9)))) ) | ($spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f6092])).
fof(f6092,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,f(s(bool,X9)))) | s(bool,X9) = s(bool,f(s(bool,X9)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f145,f4245])).
fof(f6180,plain,(
  ( ! [X40] : (s(bool,'T_0') != s(bool,X40) | s(bool,f(s(bool,X40))) = s(bool,f(s(bool,f(s(bool,X40)))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6077,f4466])).
fof(f6077,plain,(
  ( ! [X40] : (s(bool,'T_0') != s(bool,X40) | s(bool,f(s(bool,X40))) = s(bool,f(s(bool,f(s(bool,X40))))) | s(bool,'T_0') = s(bool,f(s(bool,X40)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4466,f4245])).
fof(f6128,plain,(
  ( ! [X35,X36] : (s(bool,X35) = s(bool,X36) | s(bool,X36) = s(bool,f(s(bool,X36))) | s(bool,'T_0') = s(bool,f(s(bool,X35)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f6074])).
fof(f6074,plain,(
  ( ! [X35,X36] : (s(bool,X35) = s(bool,X36) | s(bool,X36) = s(bool,f(s(bool,X36))) | s(bool,'T_0') = s(bool,f(s(bool,X35))) | s(bool,'T_0') = s(bool,f(s(bool,X35)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f4245])).
fof(f19819,plain,(
  ( ! [X33,X34] : (s(bool,X33) = s(bool,X34) | s(bool,f(s(bool,X33))) = s(bool,f(s(bool,f(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6073,f4466])).
fof(f6073,plain,(
  ( ! [X33,X34] : (s(bool,X33) = s(bool,X34) | s(bool,f(s(bool,X33))) = s(bool,f(s(bool,f(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,f(s(bool,X33)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f4245])).
fof(f6178,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,'T_0') = s(bool,f(s(bool,X18)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6065,f340])).
fof(f6065,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,X19) != s(bool,f(s(bool,X18))) | p(s(bool,X19)) | s(bool,'T_0') = s(bool,f(s(bool,X18)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f221,f4245])).
fof(f6176,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | ~p(s(bool,X0))) ) | ($spl24 | $spl30)),
  inference(condensation,[],[f6175])).
fof(f6175,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s(bool,X17) = s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,f(s(bool,X16)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6064,f340])).
fof(f6064,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) = s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,f(s(bool,X16)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f193,f4245])).
fof(f6173,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'T_0') = s(bool,f(s(bool,X10)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6060,f141])).
fof(f6060,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'F_0') != s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,f(s(bool,X10)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f149,f4245])).
fof(f6132,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'T_0') = s(bool,f(s(bool,X9)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f6059])).
fof(f6059,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'T_0') = s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,f(s(bool,X9)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f146,f4245])).
fof(f6170,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,f(s(bool,X8))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6058,f145])).
fof(f6058,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,f(s(bool,X8)))) | s(bool,'T_0') = s(bool,f(s(bool,X8)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f145,f4245])).
fof(f6134,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X15) = s(bool,f(s(bool,X15))) | s(bool,f(s(bool,X15))) = s(bool,f(s(bool,f(s(bool,X15)))))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f6048])).
fof(f6048,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X15) = s(bool,f(s(bool,X15))) | s(bool,f(s(bool,X15))) = s(bool,f(s(bool,f(s(bool,X15))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f2970])).
fof(f19806,plain,(
  ( ! [X25] : (s(bool,f(s(bool,X25))) = s(bool,f(s(bool,f(s(bool,X25)))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6139,f4466])).
fof(f6139,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,f(s(bool,X25))) | s(bool,f(s(bool,X25))) = s(bool,f(s(bool,f(s(bool,X25)))))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f6025])).
fof(f6025,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,f(s(bool,X25))) | s(bool,f(s(bool,X25))) = s(bool,f(s(bool,f(s(bool,X25))))) | s(bool,'T_0') = s(bool,f(s(bool,X25)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f4245])).
fof(f19804,plain,(
  ( ! [X23] : (s(bool,'T_0') = s(bool,f(s(bool,X23))) | s(bool,'F_0') = s(bool,f(s(bool,'F_0')))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6023,f6117])).
fof(f6023,plain,(
  ( ! [X23] : (s(bool,'T_0') = s(bool,f(s(bool,X23))) | s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f141])).
fof(f19802,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') = s(bool,f(s(bool,X10))) | s(bool,X9) = s(bool,f(s(bool,X9)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6146,f6117])).
fof(f6146,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') = s(bool,f(s(bool,X10))) | s(bool,X9) = s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f6015])).
fof(f6015,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') = s(bool,f(s(bool,X10))) | s(bool,X9) = s(bool,f(s(bool,X9))) | s(bool,X9) = s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f2970])).
fof(f19800,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,f(s(bool,X6))) | s(bool,X5) = s(bool,f(s(bool,X5)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f19799,f4466])).
fof(f19799,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,f(s(bool,X6))) | s(bool,X5) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6013,f6117])).
fof(f6013,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,f(s(bool,X6))) | s(bool,X5) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f300])).
fof(f19797,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,f(s(bool,X3)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f19796,f4466])).
fof(f19796,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6012,f6117])).
fof(f6012,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f300])).
fof(f19794,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X0) = s(bool,f(s(bool,X0)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6010,f4466])).
fof(f6010,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f141])).
fof(f6153,plain,(
  ( ! [X19,X18] : (s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,f(s(bool,X18))) | s(bool,X19) = s(bool,f(s(bool,X19)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f6006])).
fof(f6006,plain,(
  ( ! [X19,X18] : (s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,f(s(bool,X18))) | s(bool,X19) = s(bool,f(s(bool,X19))) | s(bool,'T_0') = s(bool,f(s(bool,X18)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f2970])).
fof(f6156,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'T_0') = s(bool,f(s(bool,X7)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f6000])).
fof(f6000,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'T_0') = s(bool,f(s(bool,X7))) | s(bool,'T_0') = s(bool,f(s(bool,X7)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f141])).
fof(f19789,plain,(
  ( ! [X23] : (s(bool,X23) = s(bool,f(s(bool,X23))) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5979,f4466])).
fof(f5979,plain,(
  ( ! [X23] : (s(bool,X23) = s(bool,f(s(bool,X23))) | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f141])).
fof(f19787,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,f(s(bool,X11)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6163,f6117])).
fof(f6163,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,f(s(bool,X11))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f5972])).
fof(f5972,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,f(s(bool,X11))) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f2970])).
fof(f19785,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,f(s(bool,X5)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f19784,f6117])).
fof(f19784,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5969,f4466])).
fof(f5969,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f300])).
fof(f19782,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,f(s(bool,X3)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f19781,f6117])).
fof(f19781,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5968,f4466])).
fof(f5968,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f300])).
fof(f19779,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5966,f6117])).
fof(f5966,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4245,f141])).
fof(f19777,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num',X36)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num',X36)))) | s(bool,X37) = s(bool,f(s(bool,X37))) | s(bool,X37) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num',X35)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5927,f134])).
fof(f5927,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num',X36)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num',X36)))) | s(bool,X37) = s(bool,f(s(bool,X37))) | s(bool,X37) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num',X35)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f195,f2970])).
fof(f5926,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X32)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X31),s('type.num.num',X33)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X33)))) | s(bool,X32) = s(bool,f(s(bool,X32))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X31)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f195,f2970])).
fof(f5649,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num',X19)))) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f189,f2970])).
fof(f5662,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num',X16)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16)))))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5648,f145])).
fof(f5648,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num',X16)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f189,f2970])).
fof(f19772,plain,(
  ( ! [X15,X16] : (p(s(bool,X15)) | s(bool,X16) = s(bool,f(s(bool,X16))) | s(bool,X15) = s(bool,X16)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5370,f240])).
fof(f5370,plain,(
  ( ! [X15,X16] : (p(s(bool,X15)) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X16) = s(bool,f(s(bool,X16))) | s(bool,X15) = s(bool,X16)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f150,f2970])).
fof(f5388,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s(bool,'F_0') = s(bool,X9) | s(bool,X9) = s(bool,f(s(bool,X9)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5367,f145])).
fof(f5367,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s(bool,'F_0') = s(bool,X9) | s(bool,X9) = s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f150,f2970])).
fof(f19769,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X13) = s(bool,X14)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5185,f134])).
fof(f5185,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X13) = s(bool,X14)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f149,f2970])).
fof(f19767,plain,(
  ( ! [X12,X11] : (~p(s(bool,X12)) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5184,f141])).
fof(f5184,plain,(
  ( ! [X12,X11] : (~p(s(bool,X12)) | s(bool,'F_0') != s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f149,f2970])).
fof(f19761,plain,(
  ( ! [X24,X25] : (s(bool,f(s(bool,X24))) = s(bool,f(s(bool,f(s(bool,X24))))) | s(bool,X24) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5138,f300])).
fof(f5138,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') != s(bool,X24) | s(bool,f(s(bool,X24))) = s(bool,f(s(bool,f(s(bool,X24))))) | s(bool,X24) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4466,f2970])).
fof(f5116,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4466,f300])).
fof(f19717,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4895,f134])).
fof(f4895,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f137,f2970])).
fof(f19715,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(bool,X6)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4894,f240])).
fof(f4894,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f137,f2970])).
fof(f4253,plain,(
  ( ! [X66] : (s(bool,X66) = s(bool,f(s(bool,X66))) | s(bool,'F_0') = s(bool,X66)) ) | ($spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f4213])).
fof(f4213,plain,(
  ( ! [X66] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X66) = s(bool,f(s(bool,X66))) | s(bool,'F_0') = s(bool,X66)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f240,f2970])).
fof(f19712,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4684,f134])).
fof(f4684,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f137,f2970])).
fof(f19710,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(bool,X6)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4683,f240])).
fof(f4683,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f137,f2970])).
fof(f4276,plain,(
  ( ! [X4,X5] : (~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,f(s(bool,X5)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f4111])).
fof(f4111,plain,(
  ( ! [X4,X5] : (~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f146,f2970])).
fof(f19687,plain,(
  ( ! [X6,X7] : (~p(s(bool,X7)) | s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4112,f141])).
fof(f4112,plain,(
  ( ! [X6,X7] : (~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f149,f2970])).
fof(f19685,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | s(bool,X1) = s(bool,f(s(bool,X1)))) ) | ($spl24 | $spl30)),
  inference(condensation,[],[f19684])).
fof(f19684,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X18) | s(bool,X17) = s(bool,f(s(bool,X17))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4116,f340])).
fof(f4116,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X16) = s(bool,X18) | s(bool,X17) = s(bool,f(s(bool,X17))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f193,f2970])).
fof(f19682,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4117,f340])).
fof(f4117,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | p(s(bool,X21)) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f221,f2970])).
fof(f19680,plain,(
  ( ! [X71] : (~p(s(bool,X71)) | s(bool,X71) = s(bool,f(s(bool,X71)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4137,f240])).
fof(f4137,plain,(
  ( ! [X71] : (~p(s(bool,X71)) | s(bool,X71) = s(bool,f(s(bool,X71))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f137,f2970])).
fof(f4149,plain,(
  ( ! [X90,X91] : (~p(s(bool,X91)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X90))) | s(bool,X91) = s(bool,f(s(bool,X91))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X90)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f235,f2970])).
fof(f18460,plain,(
  ( ! [X14,X13] : (~p(s(bool,X14)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X13))) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X13)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f235,f2970])).
fof(f12105,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,X9) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X10))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f11978,f141])).
fof(f6405,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f6117,f300])).
fof(f6117,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl24 | $spl30)),
  inference(equality_factoring,[],[f4245])).
fof(f5052,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4243,f300])).
fof(f4245,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f4232])).
fof(f4232,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,X0) | s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl24 | $spl30)),
  inference(equality_factoring,[],[f2970])).
fof(f4243,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f4235])).
fof(f4235,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl30)),
  inference(equality_factoring,[],[f2970])).
fof(f4229,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl24 | $spl30)),
  inference(equality_factoring,[],[f2970])).
fof(f4225,plain,(
  ( ! [X101,X99,X100,X98] : (s(X100,f(s(bool,'T_0'))) != s(X100,g(s(bool,'T_0'))) | s(fun(bool,X100),X98) = s(fun(bool,X100),X99) | s(bool,X101) = s(bool,f(s(bool,X101))) | s(bool,X101) = s(bool,sK2(X98,X99,X100,bool))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f194,f2970])).
fof(f4167,plain,(
  ( ! [X19,X20] : (s(bool,X19) = s(bool,f(s(bool,X19))) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,X19) = s(bool,X20)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2966,f2970])).
fof(f4166,plain,(
  ( ! [X17,X18] : (s(bool,X17) = s(bool,g(s(bool,X17))) | s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,X17) = s(bool,X18)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2849,f2970])).
fof(f4160,plain,(
  ( ! [X118,X116,X119,X117] : (s(X118,f(s(bool,X119))) != s(X118,g(s(bool,X119))) | s(fun(bool,X118),X116) = s(fun(bool,X118),X117) | s(bool,X119) = s(bool,f(s(bool,X119))) | s(bool,'T_0') = s(bool,sK2(X116,X117,X118,bool))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f194,f2970])).
fof(f4466,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,X74) | s(bool,X74) = s(bool,f(s(bool,X74)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4139,f240])).
fof(f4139,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,X74) | s(bool,X74) = s(bool,f(s(bool,X74))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f240,f2970])).
fof(f4011,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,f(s(bool,X6))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f300])).
fof(f4010,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X2))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,f(s(bool,X3))) | s(bool,X0) = s(bool,X3)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f2970])).
fof(f4373,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,f(s(bool,X4)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f3951,f1141])).
fof(f3951,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,f(s(bool,X4)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f141])).
fof(f30500,plain,(
  ( ! [X54,X53] : (s(bool,X54) = s(bool,g(s(bool,X54))) | s(bool,X53) = s(bool,f(s(bool,X53)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f4360,f4373,f4343,f4324,f4318,f4010,f4011,f4309,f4450,f4456,f4466,f4160,f4166,f4167,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5019,f4959,f5018,f5013,f19724,f19726,f19728,f19730,f19732,f19734,f5020,f5022,f19738,f19740,f5007,f5008,f19744,f19746,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f5983,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6027,f19810,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f19825,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19841,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10530,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f19958,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f19992,f10829,f20025,f11913,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f12066,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f17054,f19490,f19492,f621,f1650,f4412,f12026,f6127,f20222,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f471,f828,f14974,f10889,f6133,f5080,f5051,f5114,f5115,f19994,f20047,f19489,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f21894,f21895,f26883,f26884,f26975,f1715,f1712,f1627,f477,f1716,f458,f1695,f1414,f1841,f459,f1515,f26243,f1628,f1612,f453,f1502,f1400,f454,f1415,f478,f26976,f20042,f1980,f1979,f1964,f1843,f1840,f1816,f1811,f1299,f1274,f1269,f1145,f534,f517,f512,f195,f189,f26245,f20139,f4498,f2104,f1144,f650,f235,f26248,f26247,f26244,f26231,f535,f4226,f4161,f908,f26246,f11978,f7516,f7515,f7514,f7508,f7500,f4490,f619,f261,f150,f149,f137,f30114,f30088,f30051,f30045,f30226,f30238,f29360,f29362,f29363,f30030,f30244,f30434,f29726,f29727,f29732,f29735,f29736,f29741,f30473,f30479,f29887,f29888,f29902,f30486,f29891,f30488,f29738,f29737,f30499])).
fof(f20222,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | $spl18),
  inference(subsumption_resolution,[],[f5007,f137])).
fof(f11913,plain,(
  ( ! [X68,X67] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(bool,f(s(bool,X67))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X68),s('type.num.num',X68))) | s(bool,'T_0') = s(bool,X67)) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f10889,f1650])).
fof(f19992,plain,(
  ( ! [X52] : (s(bool,'F_0') = s(bool,f(s(bool,X52))) | s(bool,X52) = s(bool,f(s(bool,X52)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10792,f242])).
fof(f10792,plain,(
  ( ! [X52,X51] : (s(bool,'F_0') = s(bool,f(s(bool,X52))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X51),s('type.num.num',X51))) | s(bool,X52) = s(bool,f(s(bool,X52)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f1650,f4252])).
fof(f19958,plain,(
  ( ! [X37,X38] : (s(bool,'F_0') = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X38),s('type.num.num',X38))))) | s(bool,X37) = s(bool,f(s(bool,X37)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10710,f4466])).
fof(f10710,plain,(
  ( ! [X37,X38] : (s(bool,'F_0') = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X38),s('type.num.num',X38))))) | s(bool,'T_0') = s(bool,X37) | s(bool,X37) = s(bool,f(s(bool,X37)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f1650,f4252])).
fof(f10530,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,f(s(bool,X65))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X66),s('type.num.num',X66))) | s(bool,'T_0') = s(bool,X65)) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f4252,f1650])).
fof(f19841,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl18),
  inference(subsumption_resolution,[],[f6325,f242])).
fof(f6325,plain,(
  ( ! [X27] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num',X27)))) ) | $spl18),
  inference(superposition,[],[f1650,f261])).
fof(f19825,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f19824,f240])).
fof(f19824,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,f(s(bool,X0)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6085,f4466])).
fof(f6085,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,f(s(bool,X0)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f1650,f4245])).
fof(f19810,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,f(s(bool,X3)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f19809,f240])).
fof(f19809,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,f(s(bool,X3)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f6041,f4466])).
fof(f6041,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f4245,f1650])).
fof(f6027,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,f(s(bool,X27))) = s(bool,f(s(bool,f(s(bool,X27))))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f4245,f1650])).
fof(f5983,plain,(
  ( ! [X27] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X27))))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f4245,f1650])).
fof(f19746,plain,(
  ( ! [X30,X31] : (s(bool,X31) = s(bool,f(s(bool,X31))) | s(bool,X31) = s(bool,f(s(bool,X30))) | s(bool,'T_0') = s(bool,X30)) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5011,f240])).
fof(f5011,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X31) = s(bool,f(s(bool,X31))) | s(bool,X31) = s(bool,f(s(bool,X30))) | s(bool,'T_0') = s(bool,X30)) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f2970,f1650])).
fof(f19744,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') = s(bool,X29) | s(bool,X29) = s(bool,f(s(bool,X29))) | s(bool,'T_0') = s(bool,f(s(bool,X28)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5010,f6117])).
fof(f5010,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') = s(bool,X29) | s(bool,X29) = s(bool,f(s(bool,X29))) | s(bool,'T_0') = s(bool,f(s(bool,X28))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f2970,f1650])).
fof(f5008,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X25) = s(bool,f(s(bool,X24))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X24)) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f2970,f1650])).
fof(f5007,plain,(
  ( ! [X23,X22] : (p(s(bool,'F_0')) | s(bool,X23) = s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | $spl18),
  inference(superposition,[],[f340,f1650])).
fof(f19740,plain,(
  ( ! [X21,X20] : (s(bool,X21) = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | $spl18),
  inference(subsumption_resolution,[],[f5006,f240])).
fof(f5006,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X21) = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | $spl18),
  inference(superposition,[],[f300,f1650])).
fof(f19738,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,f(s(bool,X14))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) ) | $spl18),
  inference(subsumption_resolution,[],[f5003,f240])).
fof(f5003,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X15) = s(bool,f(s(bool,X14))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) ) | $spl18),
  inference(superposition,[],[f300,f1650])).
fof(f5022,plain,(
  ( ! [X8,X9] : (p(s(bool,X9)) | s(bool,X9) = s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl18),
  inference(subsumption_resolution,[],[f5000,f137])).
fof(f5000,plain,(
  ( ! [X8,X9] : (p(s(bool,'F_0')) | p(s(bool,X9)) | s(bool,X9) = s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl18),
  inference(superposition,[],[f190,f1650])).
fof(f5020,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | s(bool,X7) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,X6)) ) | $spl18),
  inference(subsumption_resolution,[],[f4999,f137])).
fof(f4999,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | p(s(bool,X7)) | s(bool,X7) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,X6)) ) | $spl18),
  inference(superposition,[],[f190,f1650])).
fof(f19734,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,X18) = s(bool,f(s(bool,X17)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4993,f240])).
fof(f4993,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,X18) = s(bool,f(s(bool,X17)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f1650,f2970])).
fof(f19732,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,X16) | s(bool,X16) = s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,f(s(bool,X15)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4992,f6117])).
fof(f4992,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,f(s(bool,X15)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f1650,f2970])).
fof(f19730,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X11))) | s(bool,'T_0') = s(bool,X12)) ) | $spl18),
  inference(subsumption_resolution,[],[f4990,f240])).
fof(f4990,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X11))) | s(bool,'T_0') = s(bool,X12)) ) | $spl18),
  inference(superposition,[],[f1650,f300])).
fof(f19728,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X6)) ) | $spl18),
  inference(subsumption_resolution,[],[f4987,f240])).
fof(f4987,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X6)) ) | $spl18),
  inference(superposition,[],[f1650,f300])).
fof(f19726,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X19)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4973,f6117])).
fof(f4973,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X19))) | s(bool,'T_0') = s(bool,X19)) ) | $spl18),
  inference(superposition,[],[f1650,f1650])).
fof(f19724,plain,(
  ( ! [X18] : (s(bool,X18) = s(bool,f(s(bool,X18)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f19723,f240])).
fof(f19723,plain,(
  ( ! [X18] : (s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4972,f4466])).
fof(f4972,plain,(
  ( ! [X18] : (s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X18)) ) | $spl18),
  inference(superposition,[],[f1650,f141])).
fof(f5013,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | ($spl18 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f4967])).
fof(f4967,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl18 | $spl24 | $spl30)),
  inference(superposition,[],[f1650,f2970])).
fof(f5018,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(duplicate_literal_removal,[],[f4961])).
fof(f4961,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(superposition,[],[f1650,f141])).
fof(f4959,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl18),
  inference(unit_resulting_resolution,[],[f240,f1650])).
fof(f5019,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl18),
  inference(forward_demodulation,[],[f4960,f261])).
fof(f4960,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0)))))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f242,f1650])).
fof(f11978,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X0) | s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl24 | $spl30)),
  inference(condensation,[],[f11945])).
fof(f11945,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,X9) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X10))) | s(bool,X9) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num',X11)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f10889,f4252])).
fof(f37896,plain,(
  ( ! [X45,X46] : (s(bool,X45) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,f(s(bool,X45))) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78 | $spl81)),
  inference(superposition,[],[f32431,f2970])).
fof(f32431,plain,(
  ( ! [X96] : (s(bool,X96) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,X96)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78 | $spl81)),
  inference(subsumption_resolution,[],[f32372,f134])).
fof(f32372,plain,(
  ( ! [X96] : (~p(s(bool,'T_0')) | s(bool,X96) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,X96)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78 | $spl81)),
  inference(superposition,[],[f30825,f19313])).
fof(f19313,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | $spl81),
  inference(superposition,[],[f280,f17790])).
fof(f30825,plain,(
  ( ! [X35,X34] : (~p(s(bool,X35)) | s(bool,X34) = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(forward_demodulation,[],[f30824,f261])).
fof(f30824,plain,(
  ( ! [X35,X36,X34] : (s(bool,X34) = s(bool,X35) | ~p(s(bool,X35)) | s(bool,X34) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num',X36)))) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(forward_demodulation,[],[f22644,f30501])).
fof(f22644,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X35)) | s(bool,X35) = s(bool,f(s(bool,X34))) | s(bool,X34) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num',X36)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f22480,f134])).
fof(f22480,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | s(bool,X35) = s(bool,f(s(bool,X34))) | s(bool,X34) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num',X36)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f193,f12026])).
fof(f34492,plain,(
  ( ! [X62,X63] : (s(bool,X62) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num',X63))) | s(bool,'F_0') = s(bool,f(s(bool,X62)))) ) | $spl154),
  inference(cnf_transformation,[],[f34492_D])).
fof(f34492_D,plain,(
  ( ! [X62,X63] : (s(bool,X62) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num',X63))) | s(bool,'F_0') = s(bool,f(s(bool,X62)))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f34491,plain,(
  ( ! [X64,X62,X63] : (s(bool,X62) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num',X63))) | s(bool,'F_0') = s(bool,f(s(bool,X62))) | s(bool,'T_0') = s(bool,X64)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(subsumption_resolution,[],[f34416,f31970])).
fof(f34416,plain,(
  ( ! [X64,X62,X63] : (s(bool,X62) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num',X63))) | s(bool,'F_0') = s(bool,f(s(bool,X62))) | s(bool,X62) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(superposition,[],[f32324,f2970])).
fof(f32324,plain,(
  ( ! [X87,X86] : (s(bool,X87) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X86),s('type.num.num',X86))) | s(bool,'F_0') = s(bool,X87)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(subsumption_resolution,[],[f32274,f137])).
fof(f32274,plain,(
  ( ! [X87,X86] : (p(s(bool,'F_0')) | s(bool,X87) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X86),s('type.num.num',X86))) | s(bool,'F_0') = s(bool,X87)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(superposition,[],[f30823,f261])).
fof(f30823,plain,(
  ( ! [X37,X38] : (p(s(bool,X38)) | s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X37)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(forward_demodulation,[],[f30822,f261])).
fof(f30822,plain,(
  ( ! [X39,X37,X38] : (s(bool,X37) != s(bool,X38) | p(s(bool,X38)) | s(bool,X37) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num',X39)))) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(forward_demodulation,[],[f22646,f30501])).
fof(f22646,plain,(
  ( ! [X39,X37,X38] : (s(bool,X38) != s(bool,f(s(bool,X37))) | p(s(bool,X38)) | s(bool,X37) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num',X39)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f22481,f134])).
fof(f22481,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | s(bool,X38) != s(bool,f(s(bool,X37))) | p(s(bool,X38)) | s(bool,X37) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num',X39)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f221,f12026])).
fof(f33933,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,g(s(bool,X28))) | s(bool,X27) = s(bool,g(s(bool,X27)))) ) | $spl152),
  inference(cnf_transformation,[],[f33933_D])).
fof(f33933_D,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,g(s(bool,X28))) | s(bool,X27) = s(bool,g(s(bool,X27)))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f33917,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) != s(bool,g(s(bool,X28))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,g(s(bool,X27)))) ) | ($spl18 | $spl24 | $spl30 | $spl74 | $spl76 | $spl78 | $spl126)),
  inference(duplicate_literal_removal,[],[f33776])).
fof(f33776,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) != s(bool,g(s(bool,X28))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,g(s(bool,X27))) | s(bool,'T_0') = s(bool,X26)) ) | ($spl18 | $spl24 | $spl30 | $spl74 | $spl76 | $spl78 | $spl126)),
  inference(superposition,[],[f33716,f3068])).
fof(f33716,plain,(
  ( ! [X198,X197] : (s(bool,X198) != s(bool,g(s(bool,X197))) | s(bool,'T_0') = s(bool,X198)) ) | ($spl18 | $spl24 | $spl30 | $spl74 | $spl76 | $spl78 | $spl126)),
  inference(subsumption_resolution,[],[f33715,f240])).
fof(f33715,plain,(
  ( ! [X198,X197] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X198) != s(bool,g(s(bool,X197))) | s(bool,'T_0') = s(bool,X198)) ) | ($spl18 | $spl24 | $spl30 | $spl74 | $spl76 | $spl78 | $spl126)),
  inference(forward_demodulation,[],[f33508,f15042])).
fof(f15042,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(fun(bool,bool),X0)))) ) | $spl74),
  inference(unit_resulting_resolution,[],[f137,f14981,f340])).
fof(f33508,plain,(
  ( ! [X198,X196,X197] : (s(bool,X198) != s(bool,g(s(bool,X197))) | s(bool,'T_0') = s(bool,X198) | s(bool,'F_0') = s(bool,f(s(fun(bool,bool),X196)))) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78 | $spl126)),
  inference(superposition,[],[f15712,f32425])).
fof(f32425,plain,(
  ( ! [X85,X84] : (s(bool,X85) = s(bool,g(s(bool,X84))) | s(bool,'F_0') = s(bool,X85)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78 | $spl126)),
  inference(subsumption_resolution,[],[f32367,f134])).
fof(f32367,plain,(
  ( ! [X85,X84] : (~p(s(bool,'T_0')) | s(bool,X85) = s(bool,g(s(bool,X84))) | s(bool,'F_0') = s(bool,X85)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78 | $spl126)),
  inference(superposition,[],[f30825,f26571])).
fof(f33655,plain,(
  ( ! [X59,X60] : (s(bool,X59) = s(bool,g(s(bool,X60))) | s(bool,'F_0') = s(bool,f(s(bool,X59)))) ) | $spl150),
  inference(cnf_transformation,[],[f33655_D])).
fof(f33655_D,plain,(
  ( ! [X59,X60] : (s(bool,X59) = s(bool,g(s(bool,X60))) | s(bool,'F_0') = s(bool,f(s(bool,X59)))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f33654,plain,(
  ( ! [X61,X59,X60] : (s(bool,X59) = s(bool,g(s(bool,X60))) | s(bool,'F_0') = s(bool,f(s(bool,X59))) | s(bool,'T_0') = s(bool,X61)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78 | $spl126)),
  inference(subsumption_resolution,[],[f33326,f31970])).
fof(f33326,plain,(
  ( ! [X61,X59,X60] : (s(bool,X59) = s(bool,g(s(bool,X60))) | s(bool,'F_0') = s(bool,f(s(bool,X59))) | s(bool,X59) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78 | $spl126)),
  inference(superposition,[],[f32425,f2970])).
fof(f33156,plain,(
  ( ! [X38] : (s(bool,X38) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X38)))) ) | $spl148),
  inference(cnf_transformation,[],[f33156_D])).
fof(f33156_D,plain,(
  ( ! [X38] : (s(bool,X38) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X38)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f33155,plain,(
  ( ! [X39,X38] : (s(bool,X38) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X38))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(subsumption_resolution,[],[f32899,f31970])).
fof(f32899,plain,(
  ( ! [X39,X38] : (s(bool,X38) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X38))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(superposition,[],[f32421,f2970])).
fof(f32421,plain,(
  ( ! [X78] : (s(bool,X78) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X78)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(subsumption_resolution,[],[f32364,f134])).
fof(f32364,plain,(
  ( ! [X78] : (~p(s(bool,'T_0')) | s(bool,X78) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X78)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(superposition,[],[f30825,f2849])).
fof(f32775,plain,(
  ( ! [X37] : (s(bool,X37) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X37)))) ) | $spl146),
  inference(cnf_transformation,[],[f32775_D])).
fof(f32775_D,plain,(
  ( ! [X37] : (s(bool,X37) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X37)))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f32774,plain,(
  ( ! [X37,X38] : (s(bool,X37) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X37))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(subsumption_resolution,[],[f32510,f31970])).
fof(f32510,plain,(
  ( ! [X37,X38] : (s(bool,X37) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X37))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(superposition,[],[f32419,f2970])).
fof(f32419,plain,(
  ( ! [X75] : (s(bool,X75) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X75)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(subsumption_resolution,[],[f32362,f134])).
fof(f32362,plain,(
  ( ! [X75] : (~p(s(bool,'T_0')) | s(bool,X75) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X75)) ) | ($spl18 | $spl24 | $spl30 | $spl76 | $spl78)),
  inference(superposition,[],[f30825,f2966])).
fof(f30428,plain,(
  ( ! [X112] : (s(bool,X112) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,X112) = s(bool,g(s(bool,X112)))) ) | $spl144),
  inference(cnf_transformation,[],[f30428_D])).
fof(f30428_D,plain,(
  ( ! [X112] : (s(bool,X112) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,X112) = s(bool,g(s(bool,X112)))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f29951,plain,(
  ( ! [X111,X112] : (s(bool,X112) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,g(s(bool,X112)))) ) | ($spl30 | $spl66 | $spl81)),
  inference(duplicate_literal_removal,[],[f29640])).
fof(f29640,plain,(
  ( ! [X111,X112] : (s(bool,X112) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,g(s(bool,X112))) | s(bool,'T_0') = s(bool,X111)) ) | ($spl30 | $spl66 | $spl81)),
  inference(superposition,[],[f22697,f3068])).
fof(f22697,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl66 | $spl81)),
  inference(superposition,[],[f13889,f17790])).
fof(f20678,plain,(
  ( ! [X102] : (s(bool,X102) = s(bool,f(s(bool,X102)))) ) | $spl94),
  inference(cnf_transformation,[],[f20678_D])).
fof(f20678_D,plain,(
  ( ! [X102] : (s(bool,X102) = s(bool,f(s(bool,X102)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f30422,plain,(
  ( ! [X87,X86] : (s(bool,X86) != s(bool,f(s(fun(bool,bool),X87))) | s(bool,X86) = s(bool,g(s(bool,X86)))) ) | $spl142),
  inference(cnf_transformation,[],[f30422_D])).
fof(f30422_D,plain,(
  ( ! [X87,X86] : (s(bool,X86) != s(bool,f(s(fun(bool,bool),X87))) | s(bool,X86) = s(bool,g(s(bool,X86)))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f30427,plain,(
  ( ! [X109,X110,X108] : (s(bool,X109) != s(bool,f(s(fun(bool,bool),X110))) | s(bool,X108) = s(bool,f(s(bool,X108))) | s(bool,X109) = s(bool,g(s(bool,X109)))) ) | ($spl24 | $spl30 | $spl78)),
  inference(subsumption_resolution,[],[f29639,f4466])).
fof(f29639,plain,(
  ( ! [X109,X110,X108] : (s(bool,X109) != s(bool,f(s(fun(bool,bool),X110))) | s(bool,X108) = s(bool,f(s(bool,X108))) | s(bool,X109) = s(bool,g(s(bool,X109))) | s(bool,'T_0') = s(bool,X108)) ) | ($spl24 | $spl30 | $spl78)),
  inference(superposition,[],[f21727,f3068])).
fof(f29956,plain,(
  ( ! [X87,X85,X86] : (s(bool,X86) != s(bool,f(s(fun(bool,bool),X87))) | s(bool,'T_0') = s(bool,X85) | s(bool,X86) = s(bool,g(s(bool,X86)))) ) | ($spl30 | $spl78)),
  inference(duplicate_literal_removal,[],[f29630])).
fof(f29630,plain,(
  ( ! [X87,X85,X86] : (s(bool,X86) != s(bool,f(s(fun(bool,bool),X87))) | s(bool,'T_0') = s(bool,X85) | s(bool,X86) = s(bool,g(s(bool,X86))) | s(bool,'T_0') = s(bool,X85)) ) | ($spl30 | $spl78)),
  inference(superposition,[],[f15712,f3068])).
fof(f30421,plain,(
  ( ! [X83,X84] : (s(bool,X83) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X84),s('type.num.num','const.num.SUC_1'(s('type.num.num',X84))))) | s(bool,X83) = s(bool,g(s(bool,X83)))) ) | $spl140),
  inference(cnf_transformation,[],[f30421_D])).
fof(f30421_D,plain,(
  ( ! [X83,X84] : (s(bool,X83) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X84),s('type.num.num','const.num.SUC_1'(s('type.num.num',X84))))) | s(bool,X83) = s(bool,g(s(bool,X83)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f29957,plain,(
  ( ! [X83,X84,X82] : (s(bool,X83) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X84),s('type.num.num','const.num.SUC_1'(s('type.num.num',X84))))) | s(bool,'T_0') = s(bool,X82) | s(bool,X83) = s(bool,g(s(bool,X83)))) ) | ($spl30 | $spl66)),
  inference(duplicate_literal_removal,[],[f29629])).
fof(f29629,plain,(
  ( ! [X83,X84,X82] : (s(bool,X83) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X84),s('type.num.num','const.num.SUC_1'(s('type.num.num',X84))))) | s(bool,'T_0') = s(bool,X82) | s(bool,X83) = s(bool,g(s(bool,X83))) | s(bool,'T_0') = s(bool,X82)) ) | ($spl30 | $spl66)),
  inference(superposition,[],[f13889,f3068])).
fof(f30419,plain,(
  ( ! [X70] : (s(bool,X70) != s(bool,g(s(bool,'T_0'))) | s(bool,X70) = s(bool,g(s(bool,X70)))) ) | $spl138),
  inference(cnf_transformation,[],[f30419_D])).
fof(f30419_D,plain,(
  ( ! [X70] : (s(bool,X70) != s(bool,g(s(bool,'T_0'))) | s(bool,X70) = s(bool,g(s(bool,X70)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f29958,plain,(
  ( ! [X70,X69] : (s(bool,X70) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X69) | s(bool,X70) = s(bool,g(s(bool,X70)))) ) | ($spl30 | $spl58)),
  inference(duplicate_literal_removal,[],[f29624])).
fof(f29624,plain,(
  ( ! [X70,X69] : (s(bool,X70) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X69) | s(bool,X70) = s(bool,g(s(bool,X70))) | s(bool,'T_0') = s(bool,X69)) ) | ($spl30 | $spl58)),
  inference(superposition,[],[f9165,f3068])).
fof(f30128,plain,(
  ( ! [X54,X52] : (s(bool,X52) = s(bool,g(s(bool,X52))) | s(bool,X52) = s(bool,X54) | s(bool,X54) = s(bool,g(s(bool,X54)))) ) | $spl130),
  inference(cnf_transformation,[],[f30128_D])).
fof(f30128_D,plain,(
  ( ! [X54,X52] : (s(bool,X52) = s(bool,g(s(bool,X52))) | s(bool,X52) = s(bool,X54) | s(bool,X54) = s(bool,g(s(bool,X54)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f30030,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,g(s(bool,X19))) | s(bool,X17) = s(bool,X19) | s(bool,X17) = s(bool,g(s(bool,X17)))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f29383])).
fof(f29383,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,g(s(bool,X19))) | s(bool,X17) = s(bool,X19) | s(bool,X17) = s(bool,g(s(bool,X17))) | s(bool,'T_0') = s(bool,X18)) ) | $spl30),
  inference(superposition,[],[f3068,f3068])).
fof(f30227,plain,(
  ( ! [X48] : (s(bool,X48) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,X48) = s(bool,g(s(bool,X48)))) ) | $spl136),
  inference(cnf_transformation,[],[f30227_D])).
fof(f30227_D,plain,(
  ( ! [X48] : (s(bool,X48) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,X48) = s(bool,g(s(bool,X48)))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f30042,plain,(
  ( ! [X47,X48] : (s(bool,X48) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,X48) = s(bool,g(s(bool,X48))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl30 | $spl81)),
  inference(duplicate_literal_removal,[],[f29317])).
fof(f29317,plain,(
  ( ! [X47,X48] : (s(bool,X48) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,X48) = s(bool,g(s(bool,X48))) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl30 | $spl81)),
  inference(superposition,[],[f3068,f21378])).
fof(f21378,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl81)),
  inference(subsumption_resolution,[],[f21368,f134])).
fof(f21368,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl81)),
  inference(superposition,[],[f19662,f300])).
fof(f30134,plain,(
  ( ! [X78,X76] : (s(bool,X76) = s(bool,g(s(bool,X76))) | s(bool,X76) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X78),s('type.num.num',X78)))) ) | $spl132),
  inference(cnf_transformation,[],[f30134_D])).
fof(f30134_D,plain,(
  ( ! [X78,X76] : (s(bool,X76) = s(bool,g(s(bool,X76))) | s(bool,X76) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X78),s('type.num.num',X78)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f30226,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X33))) | s(bool,X34) = s(bool,g(s(bool,X34))) | s(bool,X32) = s(bool,f(s(bool,X32)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f29312,f4466])).
fof(f29312,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X33))) | s(bool,X34) = s(bool,g(s(bool,X34))) | s(bool,'T_0') = s(bool,X32) | s(bool,X32) = s(bool,f(s(bool,X32)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f3068,f4252])).
fof(f30045,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) = s(bool,X28) | s(bool,X28) = s(bool,g(s(bool,X28))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,g(s(bool,X27)))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f29310])).
fof(f29310,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) = s(bool,X28) | s(bool,X28) = s(bool,g(s(bool,X28))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,g(s(bool,X27))) | s(bool,'T_0') = s(bool,X26)) ) | $spl30),
  inference(superposition,[],[f3068,f3068])).
fof(f30137,plain,(
  ( ! [X79] : (s(bool,X79) = s(bool,g(s(bool,X79))) | s(bool,'F_0') = s(bool,X79)) ) | $spl134),
  inference(cnf_transformation,[],[f30137_D])).
fof(f30137_D,plain,(
  ( ! [X79] : (s(bool,X79) = s(bool,g(s(bool,X79))) | s(bool,'F_0') = s(bool,X79)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f30051,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(duplicate_literal_removal,[],[f29301])).
fof(f29301,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(superposition,[],[f3068,f141])).
fof(f30088,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X30) | s(bool,X29) = s(bool,g(s(bool,X29))) | s(bool,X31) = s(bool,g(s(bool,X31))) | s(bool,X29) = s(bool,X31)) ) | $spl30),
  inference(duplicate_literal_removal,[],[f29252])).
fof(f29252,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X30) | s(bool,X29) = s(bool,g(s(bool,X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,g(s(bool,X31))) | s(bool,X29) = s(bool,X31)) ) | $spl30),
  inference(superposition,[],[f3068,f3068])).
fof(f30136,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') = s(bool,X80) | s(bool,X79) = s(bool,g(s(bool,X79))) | s(bool,'F_0') = s(bool,X79)) ) | ($spl11 | $spl24 | $spl30 | $spl42)),
  inference(duplicate_literal_removal,[],[f30135])).
fof(f30135,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') = s(bool,X80) | s(bool,X79) = s(bool,g(s(bool,X79))) | s(bool,'T_0') = s(bool,X80) | s(bool,'F_0') = s(bool,X79)) ) | ($spl11 | $spl24 | $spl30 | $spl42)),
  inference(forward_demodulation,[],[f29189,f20681])).
fof(f20681,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,f(s(bool,X7)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2066,f1908,f1907,f1889,f1906,f1669,f4360,f4373,f4343,f4334,f4324,f4318,f4010,f4011,f4309,f4410,f4417,f4429,f4450,f4456,f4269,f4267,f4266,f4466,f4160,f4166,f4167,f4492,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5080,f5051,f5114,f5115,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f19994,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f20042,f12066,f20047,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f20139,f17054,f19490,f19489,f19492,f621,f4412,f137,f149,f150,f261,f12026,f6127,f631,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4498,f4161,f4226,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f1627,f477,f1715,f1712,f1695,f458,f1716,f1414,f1841,f478,f1400,f454,f1415,f1979,f1840,f1816,f1299,f1274,f1269,f1145,f471,f195,f189,f1502,f459,f534,f1515,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f235,f650,f1144,f2104,f535,f828,f14974,f7500,f619,f908,f11978,f10889,f7516,f7515,f7514,f7508,f6133])).
fof(f4492,plain,(
  ( ! [X45,X43,X44] : (s(bool,X43) = s(bool,f(s(bool,X44))) | s(bool,'T_0') = s(bool,X44) | s(bool,X45) = s(bool,f(s(bool,X45))) | s(bool,X43) = s(bool,X45)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4202,f134])).
fof(f4202,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'T_0')) | s(bool,X43) = s(bool,f(s(bool,X44))) | s(bool,'T_0') = s(bool,X44) | s(bool,X45) = s(bool,f(s(bool,X45))) | s(bool,X43) = s(bool,X45)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f1907,f2970])).
fof(f4269,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,f(s(bool,X0)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f4268])).
fof(f4268,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,f(s(bool,X42))) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,f(s(bool,X42)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f4125])).
fof(f4125,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,f(s(bool,X42))) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,f(s(bool,X42))) | s(bool,'T_0') = s(bool,X41)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f1720,f2970])).
fof(f4417,plain,(
  ( ! [X57,X58,X56] : (s(bool,f(s(bool,X56))) = s(bool,f(s(bool,X57))) | s(bool,'T_0') = s(bool,X57) | s(bool,X56) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4054,f340])).
fof(f4054,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,X56)) | s(bool,f(s(bool,X56))) = s(bool,f(s(bool,X57))) | s(bool,'T_0') = s(bool,X57) | s(bool,X56) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f1907,f2970])).
fof(f4410,plain,(
  ( ! [X12,X11] : (p(s(bool,f(s(bool,X11)))) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4036,f1889])).
fof(f4334,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,X0)))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f4333])).
fof(f4333,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,X40) | s(bool,f(s(bool,X39))) = s(bool,f(s(bool,f(s(bool,X39))))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f3974])).
fof(f3974,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,X40) | s(bool,f(s(bool,X39))) = s(bool,f(s(bool,f(s(bool,X39))))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X39)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f2970,f1720])).
fof(f1669,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl11),
  inference(unit_resulting_resolution,[],[f137,f631,f340])).
fof(f1906,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl11),
  inference(subsumption_resolution,[],[f1874,f134])).
fof(f1874,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl11),
  inference(superposition,[],[f149,f1720])).
fof(f1889,plain,(
  ( ! [X2] : (p(s(bool,f(s(bool,X2)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl11),
  inference(trivial_inequality_removal,[],[f1872])).
fof(f1872,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,f(s(bool,X2)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl11),
  inference(superposition,[],[f145,f1720])).
fof(f1907,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) | $spl11),
  inference(subsumption_resolution,[],[f1878,f134])).
fof(f1878,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) | $spl11),
  inference(superposition,[],[f193,f1720])).
fof(f1908,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X12)) ) | $spl11),
  inference(subsumption_resolution,[],[f1879,f134])).
fof(f1879,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,f(s(bool,X12))) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,X12)) ) | $spl11),
  inference(superposition,[],[f221,f1720])).
fof(f2066,plain,(
  ( ! [X26,X25] : (s(bool,f(s(bool,X25))) = s(bool,f(s(bool,X26))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) ) | $spl11),
  inference(subsumption_resolution,[],[f2044,f134])).
fof(f2044,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,f(s(bool,X25))) = s(bool,f(s(bool,X26))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) ) | $spl11),
  inference(superposition,[],[f1907,f1720])).
fof(f2054,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl11),
  inference(subsumption_resolution,[],[f2034,f134])).
fof(f2034,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl11),
  inference(superposition,[],[f1907,f300])).
fof(f1720,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl11),
  inference(superposition,[],[f1669,f141])).
fof(f1671,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl11),
  inference(superposition,[],[f631,f141])).
fof(f1670,plain,(
  p(s(bool,f(s(bool,'F_0')))) | $spl11),
  inference(unit_resulting_resolution,[],[f137,f631,f190])).
fof(f1450,plain,(
  ( ! [X26,X25] : (s(bool,f(s(bool,X25))) = s(bool,f(s(bool,X26))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) ) | $spl11),
  inference(subsumption_resolution,[],[f1430,f134])).
fof(f1430,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,f(s(bool,X25))) = s(bool,f(s(bool,X26))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) ) | $spl11),
  inference(superposition,[],[f1214,f666])).
fof(f1440,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl11),
  inference(subsumption_resolution,[],[f1420,f134])).
fof(f1420,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl11),
  inference(superposition,[],[f1214,f300])).
fof(f1215,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X12)) ) | $spl11),
  inference(subsumption_resolution,[],[f1186,f134])).
fof(f1186,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,f(s(bool,X12))) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,X12)) ) | $spl11),
  inference(superposition,[],[f221,f666])).
fof(f1214,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) | $spl11),
  inference(subsumption_resolution,[],[f1185,f134])).
fof(f1185,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) | $spl11),
  inference(superposition,[],[f193,f666])).
fof(f1213,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl11),
  inference(subsumption_resolution,[],[f1181,f134])).
fof(f1181,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl11),
  inference(superposition,[],[f149,f666])).
fof(f1196,plain,(
  ( ! [X2] : (p(s(bool,f(s(bool,X2)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl11),
  inference(trivial_inequality_removal,[],[f1179])).
fof(f1179,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,f(s(bool,X2)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl11),
  inference(superposition,[],[f145,f666])).
fof(f666,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl11),
  inference(superposition,[],[f636,f141])).
fof(f638,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl11),
  inference(superposition,[],[f631,f141])).
fof(f636,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl11),
  inference(unit_resulting_resolution,[],[f137,f631,f340])).
fof(f637,plain,(
  p(s(bool,f(s(bool,'F_0')))) | $spl11),
  inference(unit_resulting_resolution,[],[f137,f631,f190])).
fof(f29189,plain,(
  ( ! [X80,X79] : (s(bool,X79) = s(bool,g(s(bool,X79))) | s(bool,X80) = s(bool,f(s(bool,X79))) | s(bool,'T_0') = s(bool,X80) | s(bool,'F_0') = s(bool,X79)) ) | ($spl30 | $spl42)),
  inference(superposition,[],[f3068,f4388])).
fof(f30133,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') = s(bool,X77) | s(bool,X76) = s(bool,g(s(bool,X76))) | s(bool,X76) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X78),s('type.num.num',X78)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f30132])).
fof(f30132,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') = s(bool,X77) | s(bool,X76) = s(bool,g(s(bool,X76))) | s(bool,'T_0') = s(bool,X77) | s(bool,X76) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X78),s('type.num.num',X78)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f29188,f20681])).
fof(f29188,plain,(
  ( ! [X78,X76,X77] : (s(bool,X76) = s(bool,g(s(bool,X76))) | s(bool,X77) = s(bool,f(s(bool,X76))) | s(bool,'T_0') = s(bool,X77) | s(bool,X76) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X78),s('type.num.num',X78)))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f3068,f4252])).
fof(f30114,plain,(
  ( ! [X54,X52,X53] : (s(bool,X52) = s(bool,g(s(bool,X52))) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,g(s(bool,X54))) | s(bool,X52) = s(bool,X54)) ) | $spl30),
  inference(duplicate_literal_removal,[],[f29179])).
fof(f29179,plain,(
  ( ! [X54,X52,X53] : (s(bool,X52) = s(bool,g(s(bool,X52))) | s(bool,'T_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,g(s(bool,X54))) | s(bool,X52) = s(bool,X54)) ) | $spl30),
  inference(superposition,[],[f3068,f3068])).
fof(f27396,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2066,f1908,f1907,f1889,f1906,f1669,f4360,f4373,f4343,f4334,f4324,f4318,f4010,f4011,f4309,f4410,f4417,f4429,f4450,f4456,f4269,f4267,f4266,f4466,f4160,f4166,f4167,f4492,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f12066,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f17054,f19490,f19492,f621,f4412,f12026,f6127,f631,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f1145,f471,f195,f828,f14974,f7500,f619,f11978,f10889,f7516,f7515,f7514,f7508,f6133,f20681,f20682,f261,f150,f149,f137,f5080,f5051,f5114,f5115,f19994,f20047,f20941,f19489,f20948,f20952,f20959,f20963,f20967,f20971,f478,f1502,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f1414,f1979,f1400,f1716,f458,f1695,f1715,f1712,f189,f477,f1269,f1840,f1841,f20956,f20976,f20697,f4161,f4226,f908,f402,f142,f140,f20042,f1299,f1298,f480,f389,f280,f274,f662,f136,f1627,f1274,f1816,f2104,f650,f1144,f235,f4498,f454,f459,f534,f535,f1415,f1515,f20139,f21894,f21895,f25212,f25224,f25199,f25193,f25190,f25146,f24556,f24557,f24558,f25111,f25268,f25269,f25270,f25290,f25352,f24738,f24739,f25365,f25048,f25374,f25380,f24847,f24848,f24850,f24856,f24857,f24862,f25404,f25405,f25410,f25411,f25412,f25413,f25414,f25415,f25416,f24979,f24980,f25002,f24989,f25556,f25869,f25884,f25885,f25886,f26233,f26243,f26244,f26245,f26231,f26246,f26247,f26248,f26266,f26570,f26572,f26574,f26576,f26578,f26562,f26580,f26555,f26586,f26549,f26595,f26596,f26597,f26599,f26601,f26604,f26541,f26345,f26614,f26528,f26617,f26618,f26619,f26636,f26654,f26503,f26571,f26685,f26883,f26884,f26975,f26976,f27074,f27299,f27302,f27297,f27309,f27311,f27295,f27314,f27316,f27391,f27395])).
fof(f27395,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(subsumption_resolution,[],[f27394,f145])).
fof(f27394,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(condensation,[],[f27393])).
fof(f27393,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X17) | p(s(bool,X16)) | p(s(bool,X17))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(forward_demodulation,[],[f27105,f26571])).
fof(f27105,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) = s(bool,g(s(bool,X16)))) ) | ($spl11 | $spl24 | $spl30 | $spl124)),
  inference(superposition,[],[f190,f26685])).
fof(f27391,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(subsumption_resolution,[],[f27390,f145])).
fof(f27390,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(condensation,[],[f27389])).
fof(f27389,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X14)) | p(s(bool,X15))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(forward_demodulation,[],[f27104,f26571])).
fof(f27104,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) = s(bool,g(s(bool,X14)))) ) | ($spl11 | $spl24 | $spl30 | $spl124)),
  inference(superposition,[],[f190,f26685])).
fof(f27316,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(subsumption_resolution,[],[f27212,f134])).
fof(f27212,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(backward_demodulation,[],[f27074,f650])).
fof(f27314,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(subsumption_resolution,[],[f27211,f134])).
fof(f27211,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(backward_demodulation,[],[f27074,f235])).
fof(f27295,plain,(
  ( ! [X103,X105,X102,X104] : (s(fun(X102,bool),X103) = s(fun(X102,bool),X104) | s(bool,X105) = s(bool,f(s(bool,X105))) | s(bool,X105) = s(bool,g(s(X102,sK2(X103,X104,bool,X102))))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(trivial_inequality_removal,[],[f27198])).
fof(f27198,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X102,bool),X103) = s(fun(X102,bool),X104) | s(bool,X105) = s(bool,f(s(bool,X105))) | s(bool,X105) = s(bool,g(s(X102,sK2(X103,X104,bool,X102))))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(backward_demodulation,[],[f27074,f4226])).
fof(f27311,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(subsumption_resolution,[],[f27195,f429])).
fof(f27195,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(backward_demodulation,[],[f27074,f26246])).
fof(f27309,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(subsumption_resolution,[],[f27185,f27074])).
fof(f27185,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(backward_demodulation,[],[f27074,f11978])).
fof(f27297,plain,(
  ( ! [X127,X128,X126] : (s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(fun(bool,X128),X126) = s(fun(bool,X128),X127)) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(duplicate_literal_removal,[],[f27180])).
fof(f27180,plain,(
  ( ! [X127,X128,X126] : (s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(fun(bool,X128),X126) = s(fun(bool,X128),X127)) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(backward_demodulation,[],[f27074,f7500])).
fof(f27302,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0)))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(condensation,[],[f27301])).
fof(f27301,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,X6) = s(bool,X7)) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(subsumption_resolution,[],[f27179,f27074])).
fof(f27179,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,X6) = s(bool,X7)) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(backward_demodulation,[],[f27074,f4490])).
fof(f27299,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(subsumption_resolution,[],[f27174,f27074])).
fof(f27174,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(backward_demodulation,[],[f27074,f149])).
fof(f27074,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(superposition,[],[f26685,f26571])).
fof(f26685,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,X0)))) ) | ($spl11 | $spl24 | $spl30 | $spl124)),
  inference(subsumption_resolution,[],[f26569,f26503])).
fof(f26503,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) = s(bool,g(s(bool,X0)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(equality_factoring,[],[f25002])).
fof(f26654,plain,(
  ( ! [X79] : (s(bool,'T_0') != s(bool,X79) | s(bool,'T_0') = s(bool,g(s(bool,g(s(bool,X79)))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26498,f25374])).
fof(f26498,plain,(
  ( ! [X79] : (s(bool,'T_0') != s(bool,X79) | s(bool,'T_0') = s(bool,g(s(bool,g(s(bool,X79))))) | s(bool,'T_0') = s(bool,g(s(bool,X79)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25374,f25002])).
fof(f26636,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9) | p(s(bool,g(s(bool,X9))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26453,f145])).
fof(f26453,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9) | p(s(bool,g(s(bool,X9)))) | s(bool,'T_0') = s(bool,g(s(bool,X9)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f145,f25002])).
fof(f26619,plain,(
  ( ! [X19,X18] : (s(bool,X19) != s(bool,g(s(bool,X18))) | p(s(bool,X19)) | s(bool,X18) = s(bool,g(s(bool,X18)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26407,f134])).
fof(f26407,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,g(s(bool,X18))) | p(s(bool,X19)) | s(bool,X18) = s(bool,g(s(bool,X18)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f221,f25002])).
fof(f26618,plain,(
  ( ! [X17,X16] : (~p(s(bool,X17)) | s(bool,X17) = s(bool,g(s(bool,X16))) | s(bool,X16) = s(bool,g(s(bool,X16)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26406,f134])).
fof(f26406,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X17) = s(bool,g(s(bool,X16))) | s(bool,X16) = s(bool,g(s(bool,X16)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f193,f25002])).
fof(f26617,plain,(
  ( ! [X10] : (s(bool,'F_0') != s(bool,g(s(bool,X10))) | s(bool,X10) = s(bool,g(s(bool,X10)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26402,f134])).
fof(f26402,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,g(s(bool,X10))) | s(bool,X10) = s(bool,g(s(bool,X10)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f149,f25002])).
fof(f26528,plain,(
  ( ! [X8] : (p(s(bool,g(s(bool,X8)))) | s(bool,X8) = s(bool,g(s(bool,X8)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f26400])).
fof(f26400,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,g(s(bool,X8)))) | s(bool,X8) = s(bool,g(s(bool,X8)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f145,f25002])).
fof(f26549,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,g(s(bool,X0))))) | s(bool,X0) = s(bool,g(s(bool,X0)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f26548])).
fof(f26548,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,g(s(bool,X20))) | s(bool,'T_0') = s(bool,g(s(bool,g(s(bool,X20)))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f26324])).
fof(f26324,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,g(s(bool,X20))) | s(bool,'T_0') = s(bool,g(s(bool,g(s(bool,X20))))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f20682])).
fof(f26555,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,g(s(bool,X46))) | s(bool,g(s(bool,X46))) = s(bool,g(s(bool,g(s(bool,X46)))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f26295])).
fof(f26295,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,g(s(bool,X46))) | s(bool,g(s(bool,X46))) = s(bool,g(s(bool,g(s(bool,X46))))) | s(bool,'T_0') = s(bool,g(s(bool,X46)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f25002])).
fof(f26233,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f26232])).
fof(f26232,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20) | s(bool,'T_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X20)))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f26164])).
fof(f26164,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20) | s(bool,'T_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X20))))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f429,f20682])).
fof(f25886,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') = s(bool,g(s(bool,X7))) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25374,f300])).
fof(f25885,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,'T_0') = s(bool,g(s(bool,X4))) | s(bool,X5) = s(bool,f(s(bool,X5))) | s(bool,X3) = s(bool,X5)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25374,f2970])).
fof(f25884,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,g(s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25374,f20682])).
fof(f25869,plain,(
  ( ! [X26,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X26)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s(bool,'T_0') = s(bool,g(s(bool,X27))) | s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X26)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f25848,f134])).
fof(f25848,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X26)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s(bool,'T_0') = s(bool,g(s(bool,X27))) | s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X26)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f650,f20682])).
fof(f25556,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X57)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X58)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X57) | s(bool,'T_0') = s(bool,g(s(bool,X60))) | s(bool,X60) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X57)))),s('type.num.num',X58)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f25543,f134])).
fof(f25543,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X57)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X58)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X57) | s(bool,'T_0') = s(bool,g(s(bool,X60))) | s(bool,X60) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X57)))),s('type.num.num',X58)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f517,f20682])).
fof(f24989,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,g(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl11 | $spl24 | $spl30)),
  inference(equality_factoring,[],[f20682])).
fof(f25002,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,g(s(bool,X0)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f24982])).
fof(f24982,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,g(s(bool,X0)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(equality_factoring,[],[f20682])).
fof(f24980,plain,(
  ( ! [X229,X227,X230,X228] : (s(bool,'T_0') != s(bool,f(s(X227,sK2(X228,X229,bool,X227)))) | s(fun(X227,bool),X228) = s(fun(X227,bool),X229) | s(bool,'T_0') = s(bool,g(s(bool,X230))) | s(bool,X230) = s(bool,g(s(X227,sK2(X228,X229,bool,X227))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f194,f20682])).
fof(f24979,plain,(
  ( ! [X225,X223,X226,X224] : (s(X225,f(s(bool,'T_0'))) != s(X225,g(s(bool,'T_0'))) | s(fun(bool,X225),X223) = s(fun(bool,X225),X224) | s(bool,'T_0') = s(bool,g(s(bool,X226))) | s(bool,X226) = s(bool,sK2(X223,X224,X225,bool))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f194,f20682])).
fof(f25416,plain,(
  ( ! [X212,X213,X214] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X212))),s('type.num.num',X213)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X213) | s(bool,'T_0') = s(bool,g(s(bool,X214))) | s(bool,X214) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X212),s('type.num.num',sK1(X213))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24978,f134])).
fof(f24978,plain,(
  ( ! [X212,X213,X214] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X212))),s('type.num.num',X213)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X213) | s(bool,'T_0') = s(bool,g(s(bool,X214))) | s(bool,X214) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X212),s('type.num.num',sK1(X213))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f459,f20682])).
fof(f25415,plain,(
  ( ! [X200,X199,X201] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X199))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X200)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X200) | s(bool,'T_0') = s(bool,g(s(bool,X201))) | s(bool,X201) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X199),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X200))))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24972,f134])).
fof(f24972,plain,(
  ( ! [X200,X199,X201] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X199))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X200)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X200) | s(bool,'T_0') = s(bool,g(s(bool,X201))) | s(bool,X201) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X199),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X200))))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f458,f20682])).
fof(f25414,plain,(
  ( ! [X194,X192,X193] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X192))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X193)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X192) | s(bool,'T_0') = s(bool,g(s(bool,X194))) | s(bool,X194) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X192)))),s('type.num.num',X193)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24970,f134])).
fof(f24970,plain,(
  ( ! [X194,X192,X193] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X192))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X193)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X192) | s(bool,'T_0') = s(bool,g(s(bool,X194))) | s(bool,X194) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X192)))),s('type.num.num',X193)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f453,f20682])).
fof(f25413,plain,(
  ( ! [X191,X189,X190] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X189),s('type.num.num','const.num.SUC_1'(s('type.num.num',X190)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X189) | s(bool,'T_0') = s(bool,g(s(bool,X191))) | s(bool,X191) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X189)),s('type.num.num',X190)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24969,f134])).
fof(f24969,plain,(
  ( ! [X191,X189,X190] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X189),s('type.num.num','const.num.SUC_1'(s('type.num.num',X190)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X189) | s(bool,'T_0') = s(bool,g(s(bool,X191))) | s(bool,X191) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X189)),s('type.num.num',X190)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f454,f20682])).
fof(f25412,plain,(
  ( ! [X187,X185,X188,X186] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X186),s('type.num.num',X187)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X185),s('type.num.num',X187)))) | s(bool,'T_0') = s(bool,g(s(bool,X188))) | s(bool,X188) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X185),s('type.num.num',X186)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24968,f134])).
fof(f24968,plain,(
  ( ! [X187,X185,X188,X186] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X186),s('type.num.num',X187)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X185),s('type.num.num',X187)))) | s(bool,'T_0') = s(bool,g(s(bool,X188))) | s(bool,X188) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X185),s('type.num.num',X186)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f195,f20682])).
fof(f25411,plain,(
  ( ! [X182,X184,X183,X181] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X183),s('type.num.num',X181)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X183)))),s('type.num.num',X182)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X183) | s(bool,'T_0') = s(bool,g(s(bool,X184))) | s(bool,X184) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X181),s('type.num.num',X182)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24967,f134])).
fof(f24967,plain,(
  ( ! [X182,X184,X183,X181] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X183),s('type.num.num',X181)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X183)))),s('type.num.num',X182)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X183) | s(bool,'T_0') = s(bool,g(s(bool,X184))) | s(bool,X184) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X181),s('type.num.num',X182)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f512,f20682])).
fof(f25410,plain,(
  ( ! [X180,X179] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X179))) | s(bool,'T_0') = s(bool,g(s(bool,X180))) | s(bool,X180) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X179)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24966,f134])).
fof(f24966,plain,(
  ( ! [X180,X179] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X179))) | s(bool,'T_0') = s(bool,g(s(bool,X180))) | s(bool,X180) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X179)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f235,f20682])).
fof(f25405,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,'T_0') = s(bool,g(s(bool,X18))) | s(bool,X16) = s(bool,X18)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24899,f134])).
fof(f24899,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,'T_0') = s(bool,g(s(bool,X18))) | s(bool,X16) = s(bool,X18)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f193,f20682])).
fof(f25404,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') = s(bool,g(s(bool,X7))) | s(bool,X6) = s(bool,X7)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24895,f134])).
fof(f24895,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') = s(bool,g(s(bool,X7))) | s(bool,X6) = s(bool,X7)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f149,f20682])).
fof(f24862,plain,(
  ( ! [X43,X44,X42] : (s(bool,X43) != s(bool,f(s(bool,X42))) | s(bool,X43) = s(bool,f(s(bool,X43))) | s(bool,'T_0') = s(bool,g(s(bool,X44))) | s(bool,X42) = s(bool,X44)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f19994,f20682])).
fof(f24857,plain,(
  ( ! [X28,X29,X27] : (s(bool,X27) != s(bool,X28) | s(bool,X28) = s(bool,f(s(bool,X28))) | s(bool,'T_0') = s(bool,g(s(bool,X29))) | s(bool,X27) = s(bool,X29)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f4466,f20682])).
fof(f24856,plain,(
  ( ! [X26,X24,X25] : (s(bool,X25) != s(bool,f(s(bool,X24))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,g(s(bool,X26))) | s(bool,X24) = s(bool,X26)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f4243,f20682])).
fof(f24850,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,g(s(bool,X4))) | s(bool,X2) = s(bool,X4)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f145,f20682])).
fof(f24848,plain,(
  ( ! [X255,X257,X256,X258] : (s(bool,X258) != s(bool,f(s(X255,sK2(X256,X257,bool,X255)))) | s(fun(X255,bool),X256) = s(fun(X255,bool),X257) | s(bool,'T_0') = s(bool,g(s(bool,X258))) | s(bool,'T_0') = s(bool,g(s(X255,sK2(X256,X257,bool,X255))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f194,f20682])).
fof(f24847,plain,(
  ( ! [X253,X251,X254,X252] : (s(X253,f(s(bool,X254))) != s(X253,g(s(bool,X254))) | s(fun(bool,X253),X251) = s(fun(bool,X253),X252) | s(bool,'T_0') = s(bool,g(s(bool,X254))) | s(bool,'T_0') = s(bool,sK2(X251,X252,X253,bool))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f194,f20682])).
fof(f25380,plain,(
  ( ! [X140,X141,X139] : (s(bool,X139) != s(bool,X140) | s(bool,X140) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X141),s('type.num.num',X141))) | s(bool,'T_0') = s(bool,g(s(bool,X139)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24798,f240])).
fof(f24798,plain,(
  ( ! [X140,X141,X139] : (s(bool,X139) != s(bool,X140) | s(bool,X140) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X141),s('type.num.num',X141))) | s(bool,'T_0') = s(bool,g(s(bool,X139))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f11978,f20682])).
fof(f25374,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,g(s(bool,X0)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f25040,f2966])).
fof(f25040,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,X0) != s(bool,f(s(bool,'T_0')))) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f25039])).
fof(f25039,plain,(
  ( ! [X50,X49] : (s(bool,X50) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,g(s(bool,X50)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f24758])).
fof(f24758,plain,(
  ( ! [X50,X49] : (s(bool,X50) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,g(s(bool,X50))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f4243,f20682])).
fof(f25048,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,g(s(bool,X0)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f25047])).
fof(f25047,plain,(
  ( ! [X4,X5] : (~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,g(s(bool,X5)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f24742])).
fof(f24742,plain,(
  ( ! [X4,X5] : (~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f146,f20682])).
fof(f24739,plain,(
  ( ! [X255,X257,X256,X258] : (s(bool,X258) != s(bool,f(s(X255,sK2(X256,X257,bool,X255)))) | s(fun(X255,bool),X256) = s(fun(X255,bool),X257) | s(bool,'T_0') = s(bool,g(s(bool,g(s(X255,sK2(X256,X257,bool,X255)))))) | s(bool,'T_0') = s(bool,X258)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f194,f20682])).
fof(f24738,plain,(
  ( ! [X253,X251,X254,X252] : (s(X253,f(s(bool,X254))) != s(X253,g(s(bool,X254))) | s(fun(bool,X253),X251) = s(fun(bool,X253),X252) | s(bool,'T_0') = s(bool,g(s(bool,sK2(X251,X252,X253,bool)))) | s(bool,'T_0') = s(bool,X254)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f194,f20682])).
fof(f25290,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,X1)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f25289])).
fof(f25289,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,g(s(bool,X10)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24636,f145])).
fof(f24636,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,g(s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f190,f20682])).
fof(f25270,plain,(
  ( ! [X26,X27,X25] : (s(bool,X26) != s(bool,g(s(bool,X25))) | p(s(bool,X26)) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24601,f134])).
fof(f24601,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | s(bool,X26) != s(bool,g(s(bool,X25))) | p(s(bool,X26)) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f221,f20682])).
fof(f25269,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X23)) | s(bool,X23) = s(bool,g(s(bool,X22))) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24600,f134])).
fof(f24600,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | s(bool,X23) = s(bool,g(s(bool,X22))) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f193,f20682])).
fof(f25268,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,g(s(bool,X12))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24596,f134])).
fof(f24596,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,g(s(bool,X12))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f149,f20682])).
fof(f25111,plain,(
  ( ! [X8,X9] : (p(s(bool,g(s(bool,X8)))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl11 | $spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f24594])).
fof(f24594,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,g(s(bool,X8)))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f145,f20682])).
fof(f24558,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,X10))) | s(bool,X9) = s(bool,X10) | s(bool,X8) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f20682,f300])).
fof(f24557,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,X6))) | s(bool,X5) = s(bool,X6) | s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,X4) = s(bool,X7)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f20682,f2970])).
fof(f24556,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,X2))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,g(s(bool,X3))) | s(bool,X0) = s(bool,X3)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f20682,f20682])).
fof(f25146,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,sK2(X0,X1,X2,bool)))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0')))) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f25145])).
fof(f25145,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'T_0') = s(bool,X106) | s(bool,'T_0') = s(bool,g(s(bool,sK2(X103,X104,X105,bool)))) | s(fun(bool,X105),X103) = s(fun(bool,X105),X104) | s(X105,f(s(bool,'F_0'))) != s(X105,g(s(bool,'F_0')))) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f24509])).
fof(f24509,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'T_0') = s(bool,X106) | s(bool,'T_0') = s(bool,g(s(bool,sK2(X103,X104,X105,bool)))) | s(bool,'T_0') = s(bool,X106) | s(fun(bool,X105),X103) = s(fun(bool,X105),X104) | s(X105,f(s(bool,'F_0'))) != s(X105,g(s(bool,'F_0')))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f20682,f619])).
fof(f25190,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,g(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f25189])).
fof(f25189,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,g(s(bool,X35))) | s(bool,'T_0') = s(bool,g(s(bool,X37))) | s(bool,X35) = s(bool,X37)) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f24476])).
fof(f24476,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,g(s(bool,X35))) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,g(s(bool,X37))) | s(bool,X35) = s(bool,X37)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f20682,f20682])).
fof(f25193,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f25192])).
fof(f25192,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,g(s(bool,X20))) | s(bool,X22) = s(bool,f(s(bool,X22))) | s(bool,X20) = s(bool,X22)) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f24471])).
fof(f24471,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,g(s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,f(s(bool,X22))) | s(bool,X20) = s(bool,X22)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f20682,f2970])).
fof(f25199,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,g(s(bool,X1)))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f25198])).
fof(f25198,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,g(s(bool,g(s(bool,X24)))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f24461])).
fof(f24461,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,g(s(bool,g(s(bool,X24))))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f20682,f20682])).
fof(f25224,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f25223])).
fof(f25223,plain,(
  ( ! [X61,X62,X63] : (s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,g(s(bool,X61))) | s(bool,X61) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num',X63)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f25222])).
fof(f25222,plain,(
  ( ! [X61,X62,X63] : (s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,g(s(bool,X61))) | s(bool,'T_0') = s(bool,X62) | s(bool,X61) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num',X63)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f24430,f20681])).
fof(f24430,plain,(
  ( ! [X61,X62,X63] : (s(bool,'T_0') = s(bool,g(s(bool,X61))) | s(bool,X62) = s(bool,f(s(bool,X61))) | s(bool,'T_0') = s(bool,X62) | s(bool,X61) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num',X63)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f20682,f4252])).
fof(f25212,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | ($spl11 | $spl24 | $spl30)),
  inference(condensation,[],[f25211])).
fof(f25211,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) = s(bool,g(s(bool,X40))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,g(s(bool,X42))) | s(bool,X40) = s(bool,X42)) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f24422])).
fof(f24422,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) = s(bool,g(s(bool,X40))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,g(s(bool,X42))) | s(bool,X40) = s(bool,X42)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f20682,f20682])).
fof(f20697,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'T_0') = s(bool,g(s(X120,sK2(X121,X122,bool,X120)))) | s(bool,X123) != s(bool,f(s(X120,sK2(X121,X122,bool,X120)))) | s(fun(X120,bool),X121) = s(fun(X120,bool),X122) | s(bool,'T_0') = s(bool,X123)) ) | ($spl11 | $spl24 | $spl30)),
  inference(backward_demodulation,[],[f20681,f4108])).
fof(f20976,plain,(
  ( ! [X220,X219,X221] : (s(bool,'T_0') = s(bool,g(s(X219,sK2(X220,X221,bool,X219)))) | s(bool,'F_0') != s(bool,f(s(X219,sK2(X220,X221,bool,X219)))) | s(fun(X219,bool),X220) = s(fun(X219,bool),X221)) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f20975,f20681])).
fof(f20975,plain,(
  ( ! [X220,X219,X221] : (s(bool,'F_0') != s(bool,f(s(X219,sK2(X220,X221,bool,X219)))) | s(fun(X219,bool),X220) = s(fun(X219,bool),X221) | s(bool,g(s(X219,sK2(X220,X221,bool,X219)))) = s(bool,f(s(bool,g(s(X219,sK2(X220,X221,bool,X219))))))) ) | ($spl24 | $spl30)),
  inference(forward_demodulation,[],[f10770,f261])).
fof(f20956,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,f(s(X56,sK2(X57,X58,bool,X56)))) | s(fun(X56,bool),X57) = s(fun(X56,bool),X58) | s(bool,'T_0') = s(bool,g(s(X56,sK2(X57,X58,bool,X56))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f20955,f20681])).
fof(f20955,plain,(
  ( ! [X57,X58,X56] : (s(fun(X56,bool),X57) = s(fun(X56,bool),X58) | s(bool,'T_0') = s(bool,g(s(X56,sK2(X57,X58,bool,X56)))) | s(bool,f(s(X56,sK2(X57,X58,bool,X56)))) = s(bool,f(s(bool,f(s(X56,sK2(X57,X58,bool,X56))))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f14957,f261])).
fof(f20971,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,X28) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s(bool,X28) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X26)),s('type.num.num',X27)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f20054,f20681])).
fof(f20967,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,X27) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num',X26)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num',sK1(X26))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f20059,f20681])).
fof(f20963,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') = s(bool,X30) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28) | s(bool,X30) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X28)))),s('type.num.num',X29)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f20066,f20681])).
fof(f20959,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X33) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X32)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X32) | s(bool,X33) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X31),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X32))))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f20072,f20681])).
fof(f20952,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,X58) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X57),s('type.num.num',X55)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X57)))),s('type.num.num',X56)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X57) | s(bool,X58) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num',X56)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f20106,f20681])).
fof(f20948,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'T_0') = s(bool,X43) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X42),s('type.num.num',X40)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X42),s('type.num.num',X41)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40) | s(bool,X43) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X40)))),s('type.num.num',X41)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f20130,f20681])).
fof(f20941,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f19492,f20681])).
fof(f20682,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl11 | $spl24 | $spl30)),
  inference(backward_demodulation,[],[f20681,f2917])).
fof(f27392,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2066,f1908,f1907,f1889,f1906,f1669,f4360,f4373,f4343,f4334,f4324,f4318,f4010,f4011,f4309,f4410,f4417,f4429,f4450,f4456,f4269,f4267,f4266,f4466,f4160,f4166,f4167,f4492,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f12066,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f17054,f19490,f19492,f621,f4412,f12026,f6127,f631,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f1145,f471,f195,f828,f14974,f7500,f619,f11978,f10889,f7516,f7515,f7514,f7508,f6133,f20681,f20682,f261,f150,f149,f137,f5080,f5051,f5114,f5115,f19994,f20047,f20941,f19489,f20948,f20952,f20959,f20963,f20967,f20971,f478,f1502,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f1414,f1979,f1400,f1716,f458,f1695,f1715,f1712,f189,f477,f1269,f1840,f1841,f20956,f20976,f20697,f4161,f4226,f908,f402,f142,f140,f20042,f1299,f1298,f480,f389,f280,f274,f662,f136,f1627,f1274,f1816,f2104,f650,f1144,f235,f4498,f454,f459,f534,f535,f1415,f1515,f20139,f21894,f21895,f25212,f25224,f25199,f25193,f25190,f25146,f24556,f24557,f24558,f25111,f25268,f25269,f25270,f25290,f25352,f24738,f24739,f25365,f25048,f25374,f25380,f24847,f24848,f24850,f24856,f24857,f24862,f25404,f25405,f25410,f25411,f25412,f25413,f25414,f25415,f25416,f24979,f24980,f25002,f24989,f25556,f25869,f25884,f25885,f25886,f26233,f26243,f26244,f26245,f26231,f26246,f26247,f26248,f26266,f26570,f26572,f26574,f26576,f26578,f26562,f26580,f26555,f26586,f26549,f26595,f26596,f26597,f26599,f26601,f26604,f26541,f26345,f26614,f26528,f26617,f26618,f26619,f26636,f26654,f26503,f26571,f26685,f26883,f26884,f26975,f26976,f27074,f27299,f27302,f27297,f27309,f27311,f27295,f27314,f27316,f27391])).
fof(f27317,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2066,f1908,f1907,f1889,f1906,f1669,f4360,f4373,f4343,f4334,f4324,f4318,f4010,f4011,f4309,f4410,f4417,f4429,f4450,f4456,f4269,f4267,f4266,f4466,f4160,f4166,f4167,f4492,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f12066,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f17054,f19490,f19492,f621,f4412,f12026,f6127,f631,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f1145,f471,f195,f828,f14974,f7500,f619,f11978,f10889,f7516,f7515,f7514,f7508,f6133,f20681,f20682,f261,f150,f149,f137,f5080,f5051,f5114,f5115,f19994,f20047,f20941,f19489,f20948,f20952,f20959,f20963,f20967,f20971,f478,f1502,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f1414,f1979,f1400,f1716,f458,f1695,f1715,f1712,f189,f477,f1269,f1840,f1841,f20956,f20976,f20697,f4161,f4226,f908,f402,f142,f140,f20042,f1299,f1298,f480,f389,f280,f274,f662,f136,f1627,f1274,f1816,f2104,f650,f1144,f235,f4498,f454,f459,f534,f535,f1415,f1515,f20139,f21894,f21895,f25212,f25224,f25199,f25193,f25190,f25146,f24556,f24557,f24558,f25111,f25268,f25269,f25270,f25290,f25352,f24738,f24739,f25365,f25048,f25374,f25380,f24847,f24848,f24850,f24856,f24857,f24862,f25404,f25405,f25410,f25411,f25412,f25413,f25414,f25415,f25416,f24979,f24980,f25002,f24989,f25556,f25869,f25884,f25885,f25886,f26233,f26243,f26244,f26245,f26231,f26246,f26247,f26248,f26266,f26570,f26572,f26574,f26576,f26578,f26562,f26580,f26555,f26586,f26549,f26595,f26596,f26597,f26599,f26601,f26604,f26541,f26345,f26614,f26528,f26617,f26618,f26619,f26636,f26654,f26503,f26571,f26685,f26883,f26884,f26975,f26976,f27074,f27299,f27302,f27297,f27309,f27311,f27295,f27314,f27316])).
fof(f27315,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2066,f1908,f1907,f1889,f1906,f1669,f4360,f4373,f4343,f4334,f4324,f4318,f4010,f4011,f4309,f4410,f4417,f4429,f4450,f4456,f4269,f4267,f4266,f4466,f4160,f4166,f4167,f4492,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f12066,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f17054,f19490,f19492,f621,f4412,f12026,f6127,f631,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f1145,f471,f195,f828,f14974,f7500,f619,f11978,f10889,f7516,f7515,f7514,f7508,f6133,f20681,f20682,f261,f150,f149,f137,f5080,f5051,f5114,f5115,f19994,f20047,f20941,f19489,f20948,f20952,f20959,f20963,f20967,f20971,f478,f1502,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f1414,f1979,f1400,f1716,f458,f1695,f1715,f1712,f189,f477,f1269,f1840,f1841,f20956,f20976,f20697,f4161,f4226,f908,f402,f142,f140,f20042,f1299,f1298,f480,f389,f280,f274,f662,f136,f1627,f1274,f1816,f2104,f650,f1144,f235,f4498,f454,f459,f534,f535,f1415,f1515,f20139,f21894,f21895,f25212,f25224,f25199,f25193,f25190,f25146,f24556,f24557,f24558,f25111,f25268,f25269,f25270,f25290,f25352,f24738,f24739,f25365,f25048,f25374,f25380,f24847,f24848,f24850,f24856,f24857,f24862,f25404,f25405,f25410,f25411,f25412,f25413,f25414,f25415,f25416,f24979,f24980,f25002,f24989,f25556,f25869,f25884,f25885,f25886,f26233,f26243,f26244,f26245,f26231,f26246,f26247,f26248,f26266,f26570,f26572,f26574,f26576,f26578,f26562,f26580,f26555,f26586,f26549,f26595,f26596,f26597,f26599,f26601,f26604,f26541,f26345,f26614,f26528,f26617,f26618,f26619,f26636,f26654,f26503,f26571,f26685,f26883,f26884,f26975,f26976,f27074,f27299,f27302,f27297,f27309,f27311,f27295,f27314])).
fof(f27313,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2066,f1908,f1907,f1889,f1906,f1669,f4360,f4373,f4343,f4334,f4324,f4318,f4010,f4011,f4309,f4410,f4417,f4429,f4450,f4456,f4269,f4267,f4266,f4466,f4160,f4166,f4167,f4492,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f12066,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f17054,f19490,f19492,f621,f4412,f12026,f6127,f631,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f1145,f471,f195,f828,f14974,f7500,f619,f11978,f10889,f7516,f7515,f7514,f7508,f6133,f20681,f20682,f261,f150,f149,f137,f5080,f5051,f5114,f5115,f19994,f20047,f20941,f19489,f20948,f20952,f20959,f20963,f20967,f20971,f478,f1502,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f1414,f1979,f1400,f1716,f458,f1695,f1715,f1712,f189,f477,f1269,f1840,f1841,f20956,f20976,f20697,f4161,f4226,f908,f402,f142,f140,f20042,f1299,f1298,f480,f389,f280,f274,f662,f136,f1627,f1274,f1816,f2104,f650,f1144,f235,f4498,f454,f459,f534,f535,f1415,f1515,f20139,f21894,f21895,f25212,f25224,f25199,f25193,f25190,f25146,f24556,f24557,f24558,f25111,f25268,f25269,f25270,f25290,f25352,f24738,f24739,f25365,f25048,f25374,f25380,f24847,f24848,f24850,f24856,f24857,f24862,f25404,f25405,f25410,f25411,f25412,f25413,f25414,f25415,f25416,f24979,f24980,f25002,f24989,f25556,f25869,f25884,f25885,f25886,f26233,f26243,f26244,f26245,f26231,f26246,f26247,f26248,f26266,f26570,f26572,f26574,f26576,f26578,f26562,f26580,f26555,f26586,f26549,f26595,f26596,f26597,f26599,f26601,f26604,f26541,f26345,f26614,f26528,f26617,f26618,f26619,f26636,f26654,f26503,f26571,f26685,f26883,f26884,f26975,f26976,f27074,f27299,f27302,f27297,f27309,f27311,f27295])).
fof(f27312,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2066,f1908,f1907,f1889,f1906,f1669,f4360,f4373,f4343,f4334,f4324,f4318,f4010,f4011,f4309,f4410,f4417,f4429,f4450,f4456,f4269,f4267,f4266,f4466,f4160,f4166,f4167,f4492,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f12066,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f17054,f19490,f19492,f621,f4412,f12026,f6127,f631,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f1145,f471,f195,f828,f14974,f7500,f619,f11978,f10889,f7516,f7515,f7514,f7508,f6133,f20681,f20682,f261,f150,f149,f137,f5080,f5051,f5114,f5115,f19994,f20047,f20941,f19489,f20948,f20952,f20959,f20963,f20967,f20971,f478,f1502,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f1414,f1979,f1400,f1716,f458,f1695,f1715,f1712,f189,f477,f1269,f1840,f1841,f20956,f20976,f20697,f4161,f4226,f908,f402,f142,f140,f20042,f1299,f1298,f480,f389,f280,f274,f662,f136,f1627,f1274,f1816,f2104,f650,f1144,f235,f4498,f454,f459,f534,f535,f1415,f1515,f20139,f21894,f21895,f25212,f25224,f25199,f25193,f25190,f25146,f24556,f24557,f24558,f25111,f25268,f25269,f25270,f25290,f25352,f24738,f24739,f25365,f25048,f25374,f25380,f24847,f24848,f24850,f24856,f24857,f24862,f25404,f25405,f25410,f25411,f25412,f25413,f25414,f25415,f25416,f24979,f24980,f25002,f24989,f25556,f25869,f25884,f25885,f25886,f26233,f26243,f26244,f26245,f26231,f26246,f26247,f26248,f26266,f26570,f26572,f26574,f26576,f26578,f26562,f26580,f26555,f26586,f26549,f26595,f26596,f26597,f26599,f26601,f26604,f26541,f26345,f26614,f26528,f26617,f26618,f26619,f26636,f26654,f26503,f26571,f26685,f26883,f26884,f26975,f26976,f27074,f27299,f27302,f27297,f27309,f27311])).
fof(f27310,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2066,f1908,f1907,f1889,f1906,f1669,f4360,f4373,f4343,f4334,f4324,f4318,f4010,f4011,f4309,f4410,f4417,f4429,f4450,f4456,f4269,f4267,f4266,f4466,f4160,f4166,f4167,f4492,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f12066,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f17054,f19490,f19492,f621,f4412,f12026,f6127,f631,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f1145,f471,f195,f828,f14974,f7500,f619,f11978,f10889,f7516,f7515,f7514,f7508,f6133,f20681,f20682,f261,f150,f149,f137,f5080,f5051,f5114,f5115,f19994,f20047,f20941,f19489,f20948,f20952,f20959,f20963,f20967,f20971,f478,f1502,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f1414,f1979,f1400,f1716,f458,f1695,f1715,f1712,f189,f477,f1269,f1840,f1841,f20956,f20976,f20697,f4161,f4226,f908,f402,f142,f140,f20042,f1299,f1298,f480,f389,f280,f274,f662,f136,f1627,f1274,f1816,f2104,f650,f1144,f235,f4498,f454,f459,f534,f535,f1415,f1515,f20139,f21894,f21895,f25212,f25224,f25199,f25193,f25190,f25146,f24556,f24557,f24558,f25111,f25268,f25269,f25270,f25290,f25352,f24738,f24739,f25365,f25048,f25374,f25380,f24847,f24848,f24850,f24856,f24857,f24862,f25404,f25405,f25410,f25411,f25412,f25413,f25414,f25415,f25416,f24979,f24980,f25002,f24989,f25556,f25869,f25884,f25885,f25886,f26233,f26243,f26244,f26245,f26231,f26246,f26247,f26248,f26266,f26570,f26572,f26574,f26576,f26578,f26562,f26580,f26555,f26586,f26549,f26595,f26596,f26597,f26599,f26601,f26604,f26541,f26345,f26614,f26528,f26617,f26618,f26619,f26636,f26654,f26503,f26571,f26685,f26883,f26884,f26975,f26976,f27074,f27299,f27302,f27297,f27309])).
fof(f27304,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2066,f1908,f1907,f1889,f1906,f1669,f4360,f4373,f4343,f4334,f4324,f4318,f4010,f4011,f4309,f4410,f4417,f4429,f4450,f4456,f4269,f4267,f4266,f4466,f4160,f4166,f4167,f4492,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f12066,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f17054,f19490,f19492,f621,f4412,f12026,f6127,f631,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f1145,f471,f195,f828,f14974,f7500,f619,f11978,f10889,f7516,f7515,f7514,f7508,f6133,f20681,f20682,f261,f150,f149,f137,f5080,f5051,f5114,f5115,f19994,f20047,f20941,f19489,f20948,f20952,f20959,f20963,f20967,f20971,f478,f1502,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f1414,f1979,f1400,f1716,f458,f1695,f1715,f1712,f189,f477,f1269,f1840,f1841,f20956,f20976,f20697,f4161,f4226,f908,f402,f142,f140,f20042,f1299,f1298,f480,f389,f280,f274,f662,f136,f1627,f1274,f1816,f2104,f650,f1144,f235,f4498,f454,f459,f534,f535,f1415,f1515,f20139,f21894,f21895,f25212,f25224,f25199,f25193,f25190,f25146,f24556,f24557,f24558,f25111,f25268,f25269,f25270,f25290,f25352,f24738,f24739,f25365,f25048,f25374,f25380,f24847,f24848,f24850,f24856,f24857,f24862,f25404,f25405,f25410,f25411,f25412,f25413,f25414,f25415,f25416,f24979,f24980,f25002,f24989,f25556,f25869,f25884,f25885,f25886,f26233,f26243,f26244,f26245,f26231,f26246,f26247,f26248,f26266,f26570,f26572,f26574,f26576,f26578,f26562,f26580,f26555,f26586,f26549,f26595,f26596,f26597,f26599,f26601,f26604,f26541,f26345,f26614,f26528,f26617,f26618,f26619,f26636,f26654,f26503,f26571,f26685,f26883,f26884,f26975,f26976,f27074,f27299,f27302,f27297])).
fof(f27303,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2066,f1908,f1907,f1889,f1906,f1669,f4360,f4373,f4343,f4334,f4324,f4318,f4010,f4011,f4309,f4410,f4417,f4429,f4450,f4456,f4269,f4267,f4266,f4466,f4160,f4166,f4167,f4492,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f12066,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f17054,f19490,f19492,f621,f4412,f12026,f6127,f631,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f1145,f471,f195,f828,f14974,f7500,f619,f11978,f10889,f7516,f7515,f7514,f7508,f6133,f20681,f20682,f261,f150,f149,f137,f5080,f5051,f5114,f5115,f19994,f20047,f20941,f19489,f20948,f20952,f20959,f20963,f20967,f20971,f478,f1502,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f1414,f1979,f1400,f1716,f458,f1695,f1715,f1712,f189,f477,f1269,f1840,f1841,f20956,f20976,f20697,f4161,f4226,f908,f402,f142,f140,f20042,f1299,f1298,f480,f389,f280,f274,f662,f136,f1627,f1274,f1816,f2104,f650,f1144,f235,f4498,f454,f459,f534,f535,f1415,f1515,f20139,f21894,f21895,f25212,f25224,f25199,f25193,f25190,f25146,f24556,f24557,f24558,f25111,f25268,f25269,f25270,f25290,f25352,f24738,f24739,f25365,f25048,f25374,f25380,f24847,f24848,f24850,f24856,f24857,f24862,f25404,f25405,f25410,f25411,f25412,f25413,f25414,f25415,f25416,f24979,f24980,f25002,f24989,f25556,f25869,f25884,f25885,f25886,f26233,f26243,f26244,f26245,f26231,f26246,f26247,f26248,f26266,f26570,f26572,f26574,f26576,f26578,f26562,f26580,f26555,f26586,f26549,f26595,f26596,f26597,f26599,f26601,f26604,f26541,f26345,f26614,f26528,f26617,f26618,f26619,f26636,f26654,f26503,f26571,f26685,f26883,f26884,f26975,f26976,f27074,f27299,f27302])).
fof(f27300,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2066,f1908,f1907,f1889,f1906,f1669,f4360,f4373,f4343,f4334,f4324,f4318,f4010,f4011,f4309,f4410,f4417,f4429,f4450,f4456,f4269,f4267,f4266,f4466,f4160,f4166,f4167,f4492,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f18460,f4149,f19680,f19682,f19685,f19687,f4276,f19710,f19712,f4253,f19715,f19717,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f12066,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f17054,f19490,f19492,f621,f4412,f12026,f6127,f631,f4153,f20677,f4150,f20680,f4277,f4418,f4308,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f1145,f471,f195,f828,f14974,f7500,f619,f11978,f10889,f7516,f7515,f7514,f7508,f6133,f20681,f20682,f261,f150,f149,f137,f5080,f5051,f5114,f5115,f19994,f20047,f20941,f19489,f20948,f20952,f20959,f20963,f20967,f20971,f478,f1502,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f1414,f1979,f1400,f1716,f458,f1695,f1715,f1712,f189,f477,f1269,f1840,f1841,f20956,f20976,f20697,f4161,f4226,f908,f402,f142,f140,f20042,f1299,f1298,f480,f389,f280,f274,f662,f136,f1627,f1274,f1816,f2104,f650,f1144,f235,f4498,f454,f459,f534,f535,f1415,f1515,f20139,f21894,f21895,f25212,f25224,f25199,f25193,f25190,f25146,f24556,f24557,f24558,f25111,f25268,f25269,f25270,f25290,f25352,f24738,f24739,f25365,f25048,f25374,f25380,f24847,f24848,f24850,f24856,f24857,f24862,f25404,f25405,f25410,f25411,f25412,f25413,f25414,f25415,f25416,f24979,f24980,f25002,f24989,f25556,f25869,f25884,f25885,f25886,f26233,f26243,f26244,f26245,f26231,f26246,f26247,f26248,f26266,f26570,f26572,f26574,f26576,f26578,f26562,f26580,f26555,f26586,f26549,f26595,f26596,f26597,f26599,f26601,f26604,f26541,f26345,f26614,f26528,f26617,f26618,f26619,f26636,f26654,f26503,f26571,f26685,f26883,f26884,f26975,f26976,f27074,f27299])).
fof(f27298,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(subsumption_resolution,[],[f27173,f134])).
fof(f27173,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(backward_demodulation,[],[f27074,f137])).
fof(f27169,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(subsumption_resolution,[],[f27168,f134])).
fof(f27168,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(forward_demodulation,[],[f27008,f26571])).
fof(f27008,plain,(
  ~p(s(bool,g(s(bool,'F_0')))) | ($spl11 | $spl24 | $spl30 | $spl124)),
  inference(unit_resulting_resolution,[],[f26685,f149])).
fof(f27166,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(subsumption_resolution,[],[f27165,f134])).
fof(f27165,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl24 | $spl30 | $spl124 | $spl126)),
  inference(forward_demodulation,[],[f27010,f26571])).
fof(f27010,plain,(
  ~p(s(bool,g(s(bool,'F_0')))) | ($spl11 | $spl24 | $spl30 | $spl124)),
  inference(unit_resulting_resolution,[],[f137,f26685,f221])).
fof(f26569,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl124),
  inference(cnf_transformation,[],[f26569_D])).
fof(f26569_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f628,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl6),
  inference(cnf_transformation,[],[f628_D])).
fof(f628_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f26615,plain,(
  ( ! [X59] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X59) = s(bool,g(s(bool,X59))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81)),
  inference(forward_demodulation,[],[f26367,f21376])).
fof(f26367,plain,(
  ( ! [X59] : (s(bool,X59) = s(bool,g(s(bool,X59))) | s(bool,'T_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81)),
  inference(superposition,[],[f25002,f21378])).
fof(f26614,plain,(
  ( ! [X55] : (s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X55) = s(bool,g(s(bool,X55)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f26613,f20681])).
fof(f26613,plain,(
  ( ! [X55] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X55) = s(bool,g(s(bool,X55))) | s(bool,X55) = s(bool,f(s(bool,X55)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f26362,f261])).
fof(f26362,plain,(
  ( ! [X54,X55] : (s(bool,X55) = s(bool,g(s(bool,X55))) | s(bool,'T_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num',X54),s('type.num.num',X54))))) | s(bool,X55) = s(bool,f(s(bool,X55)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f4252])).
fof(f26345,plain,(
  ( ! [X34] : (s(bool,X34) = s(bool,g(s(bool,X34))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f141])).
fof(f26571,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,g(s(bool,X4)))) ) | $spl126),
  inference(cnf_transformation,[],[f26571_D])).
fof(f26571_D,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,g(s(bool,X4)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f26582,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl128),
  inference(cnf_transformation,[],[f26582_D])).
fof(f26582_D,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f26606,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X25)))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81)),
  inference(forward_demodulation,[],[f26605,f21376])).
fof(f26605,plain,(
  ( ! [X25] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,g(s(bool,X25)))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81)),
  inference(subsumption_resolution,[],[f26340,f25374])).
fof(f26340,plain,(
  ( ! [X25] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,g(s(bool,X25))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81)),
  inference(superposition,[],[f25002,f21378])).
fof(f26541,plain,(
  ( ! [X19,X20] : (s(bool,X20) = s(bool,g(s(bool,X20))) | s(bool,'T_0') = s(bool,g(s(bool,X19))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f26337])).
fof(f26337,plain,(
  ( ! [X19,X20] : (s(bool,X20) = s(bool,g(s(bool,X20))) | s(bool,'T_0') = s(bool,g(s(bool,X19))) | s(bool,'T_0') = s(bool,g(s(bool,X19))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f20682])).
fof(f26604,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X15)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26603,f25374])).
fof(f26603,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X15)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f26602,f20681])).
fof(f26602,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X15))) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f26335,f261])).
fof(f26335,plain,(
  ( ! [X15,X16] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num',X16))) = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num',X16))))) | s(bool,'T_0') = s(bool,g(s(bool,X15))) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f4252])).
fof(f26601,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') = s(bool,X12) | s(bool,X12) = s(bool,g(s(bool,X12))) | s(bool,'T_0') = s(bool,g(s(bool,X11)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f26600,f20681])).
fof(f26600,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,g(s(bool,X12))) | s(bool,'T_0') = s(bool,g(s(bool,X11))) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26333,f25374])).
fof(f26333,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,g(s(bool,X12))) | s(bool,'T_0') = s(bool,g(s(bool,X11))) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f2970])).
fof(f26599,plain,(
  ( ! [X10,X9] : (s(bool,X10) = s(bool,g(s(bool,X10))) | s(bool,'T_0') = s(bool,g(s(bool,X9))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26598,f25374])).
fof(f26598,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,X10) = s(bool,g(s(bool,X10))) | s(bool,'T_0') = s(bool,g(s(bool,X9))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f26332,f20681])).
fof(f26332,plain,(
  ( ! [X10,X9] : (s(bool,X10) = s(bool,g(s(bool,X10))) | s(bool,'T_0') = s(bool,g(s(bool,X9))) | s(bool,X9) = s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f2970])).
fof(f26597,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,g(s(bool,X6))) | s(bool,'T_0') = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26330,f25374])).
fof(f26330,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,g(s(bool,X6))) | s(bool,'T_0') = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f300])).
fof(f26596,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,g(s(bool,X4))) | s(bool,'T_0') = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26329,f25374])).
fof(f26329,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,g(s(bool,X4))) | s(bool,'T_0') = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f300])).
fof(f26595,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X0)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26327,f25374])).
fof(f26327,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f141])).
fof(f26588,plain,(
  ( ! [X59] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X59)))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81)),
  inference(forward_demodulation,[],[f26587,f21376])).
fof(f26587,plain,(
  ( ! [X59] : (s(bool,'T_0') = s(bool,g(s(bool,X59))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81)),
  inference(subsumption_resolution,[],[f26306,f25374])).
fof(f26306,plain,(
  ( ! [X59] : (s(bool,'T_0') = s(bool,g(s(bool,X59))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81)),
  inference(superposition,[],[f25002,f21378])).
fof(f26586,plain,(
  ( ! [X55] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X55)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26585,f25374])).
fof(f26585,plain,(
  ( ! [X55] : (s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X55)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f26584,f20681])).
fof(f26584,plain,(
  ( ! [X55] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X55))) | s(bool,X55) = s(bool,f(s(bool,X55)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f26301,f261])).
fof(f26301,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') = s(bool,g(s(bool,X55))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X54),s('type.num.num',X54))) = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num',X54),s('type.num.num',X54))))) | s(bool,X55) = s(bool,f(s(bool,X55)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f4252])).
fof(f26580,plain,(
  ( ! [X34] : (s(bool,'T_0') = s(bool,g(s(bool,X34))) | s(bool,'F_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26284,f25374])).
fof(f26284,plain,(
  ( ! [X34] : (s(bool,'T_0') = s(bool,g(s(bool,X34))) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f141])).
fof(f26579,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X25) = s(bool,g(s(bool,X25))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81)),
  inference(forward_demodulation,[],[f26279,f21376])).
fof(f26279,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))))) | s(bool,X25) = s(bool,g(s(bool,X25))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81)),
  inference(superposition,[],[f25002,f21378])).
fof(f26562,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') = s(bool,g(s(bool,X22))) | s(bool,X21) = s(bool,g(s(bool,X21))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f26277])).
fof(f26277,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') = s(bool,g(s(bool,X22))) | s(bool,X21) = s(bool,g(s(bool,X21))) | s(bool,'T_0') = s(bool,g(s(bool,X22))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f20682])).
fof(f26578,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X15) = s(bool,g(s(bool,X15)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f26577,f20681])).
fof(f26577,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X15) = s(bool,g(s(bool,X15))) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f26274,f261])).
fof(f26274,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num',X16))))) | s(bool,X15) = s(bool,g(s(bool,X15))) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f4252])).
fof(f26576,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') = s(bool,g(s(bool,X12))) | s(bool,X11) = s(bool,g(s(bool,X11))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26575,f25374])).
fof(f26575,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,g(s(bool,X12))) | s(bool,X11) = s(bool,g(s(bool,X11))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f26272,f20681])).
fof(f26272,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') = s(bool,g(s(bool,X12))) | s(bool,X11) = s(bool,g(s(bool,X11))) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f2970])).
fof(f26574,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,X10))) | s(bool,X9) = s(bool,g(s(bool,X9)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f26573,f20681])).
fof(f26573,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') = s(bool,g(s(bool,X10))) | s(bool,X9) = s(bool,g(s(bool,X9))) | s(bool,X9) = s(bool,f(s(bool,X9)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26271,f25374])).
fof(f26271,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') = s(bool,g(s(bool,X10))) | s(bool,X9) = s(bool,g(s(bool,X9))) | s(bool,X9) = s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f2970])).
fof(f26572,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,g(s(bool,X6))) | s(bool,X5) = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26269,f25374])).
fof(f26269,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,g(s(bool,X6))) | s(bool,X5) = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f300])).
fof(f26570,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,g(s(bool,X4))) | s(bool,X3) = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f26268,f25374])).
fof(f26268,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,g(s(bool,X4))) | s(bool,X3) = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f300])).
fof(f26266,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f25002,f141])).
fof(f1488,plain,(
  ( ! [X7] : (p(s(bool,X7))) ) | $spl16),
  inference(cnf_transformation,[],[f1488_D])).
fof(f1488_D,plain,(
  ( ! [X7] : (p(s(bool,X7))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f25366,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,X3)))) ) | $spl122),
  inference(cnf_transformation,[],[f25366_D])).
fof(f25366_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,X3)))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f25365,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,g(s(bool,X3)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24741,f145])).
fof(f24741,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f145,f20682])).
fof(f25353,plain,(
  ( ! [X222,X221] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X221),s('type.num.num',X222)))) | s(bool,'T_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X221))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X222)))))))) ) | $spl120),
  inference(cnf_transformation,[],[f25353_D])).
fof(f25353_D,plain,(
  ( ! [X222,X221] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X221),s('type.num.num',X222)))) | s(bool,'T_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X221))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X222)))))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f25352,plain,(
  ( ! [X222,X223,X221] : (p(s(bool,X223)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X221),s('type.num.num',X222)))) | s(bool,'T_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X221))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X222)))))))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f24723,f145])).
fof(f24723,plain,(
  ( ! [X222,X223,X221] : (p(s(bool,X223)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X221),s('type.num.num',X222)))) | s(bool,'T_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X221))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X222))))))) | s(bool,'T_0') = s(bool,X223)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f189,f20682])).
fof(f23899,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) | $spl118),
  inference(cnf_transformation,[],[f23899_D])).
fof(f23899_D,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f23889,plain,(
  ( ! [X19,X17,X18] : (s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) | ($spl24 | $spl30 | $spl70)),
  inference(duplicate_literal_removal,[],[f23842])).
fof(f23842,plain,(
  ( ! [X19,X17,X18] : (s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl24 | $spl30 | $spl70)),
  inference(superposition,[],[f14361,f2970])).
fof(f23539,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | $spl116),
  inference(cnf_transformation,[],[f23539_D])).
fof(f23539_D,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f23528,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | ($spl24 | $spl30 | $spl66 | $spl81)),
  inference(duplicate_literal_removal,[],[f23491])).
fof(f23491,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl24 | $spl30 | $spl66 | $spl81)),
  inference(superposition,[],[f22697,f2970])).
fof(f22721,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) | $spl114),
  inference(cnf_transformation,[],[f22721_D])).
fof(f22721_D,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f22711,plain,(
  ( ! [X19,X17,X18] : (s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) | ($spl24 | $spl30 | $spl66)),
  inference(duplicate_literal_removal,[],[f22675])).
fof(f22675,plain,(
  ( ! [X19,X17,X18] : (s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl24 | $spl30 | $spl66)),
  inference(superposition,[],[f13889,f2970])).
fof(f21723,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,f(s(fun(bool,bool),X19))) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) | $spl112),
  inference(cnf_transformation,[],[f21723_D])).
fof(f21723_D,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,f(s(fun(bool,bool),X19))) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f21711,plain,(
  ( ! [X19,X17,X18] : (s(bool,X18) != s(bool,f(s(fun(bool,bool),X19))) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) | ($spl24 | $spl30 | $spl78)),
  inference(duplicate_literal_removal,[],[f21682])).
fof(f21682,plain,(
  ( ! [X19,X17,X18] : (s(bool,X18) != s(bool,f(s(fun(bool,bool),X19))) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl24 | $spl30 | $spl78)),
  inference(superposition,[],[f15712,f2970])).
fof(f21072,plain,(
  $false | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(subsumption_resolution,[],[f21071,f134])).
fof(f21071,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(forward_demodulation,[],[f21047,f21030])).
fof(f21030,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(backward_demodulation,[],[f21026,f274])).
fof(f21026,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(subsumption_resolution,[],[f21025,f21001])).
fof(f21001,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0)))) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(forward_demodulation,[],[f19662,f20802])).
fof(f20802,plain,(
  ( ! [X44] : (s(bool,X44) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0)))) ) | ($spl11 | $spl24 | $spl30 | $spl90)),
  inference(subsumption_resolution,[],[f20695,f240])).
fof(f20695,plain,(
  ( ! [X44] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X44) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0)))) ) | ($spl11 | $spl24 | $spl30 | $spl90)),
  inference(backward_demodulation,[],[f20681,f19190])).
fof(f21025,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl11 | $spl24 | $spl30 | $spl90)),
  inference(forward_demodulation,[],[f1627,f20802])).
fof(f21047,plain,(
  ( ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0'))))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(backward_demodulation,[],[f21026,f21003])).
fof(f21003,plain,(
  ( ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',sK1(sK0)))))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(subsumption_resolution,[],[f21002,f21001])).
fof(f21002,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',sK1(sK0)))))) ) | ($spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(forward_demodulation,[],[f18176,f20802])).
fof(f18176,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',sK0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',sK1(sK0)))))) ) | $spl81),
  inference(superposition,[],[f189,f17790])).
fof(f21070,plain,(
  $false | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(subsumption_resolution,[],[f21069,f134])).
fof(f21069,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(forward_demodulation,[],[f21045,f21030])).
fof(f21045,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num','const.num.|48|_0')))) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(backward_demodulation,[],[f21026,f21001])).
fof(f21067,plain,(
  ( ! [X123,X121,X122,X120] : (s(fun(X120,bool),X121) = s(fun(X120,bool),X122) | s(bool,X123) = s(bool,f(s(bool,X123))) | s(bool,'T_0') = s(bool,g(s(X120,sK2(X121,X122,bool,X120))))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(subsumption_resolution,[],[f21063,f4252])).
fof(f21063,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,X123) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(fun(X120,bool),X121) = s(fun(X120,bool),X122) | s(bool,X123) = s(bool,f(s(bool,X123))) | s(bool,'T_0') = s(bool,g(s(X120,sK2(X121,X122,bool,X120))))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(backward_demodulation,[],[f21060,f4161])).
fof(f21060,plain,(
  ( ! [X44] : (s(bool,X44) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(forward_demodulation,[],[f21044,f21026])).
fof(f21044,plain,(
  ( ! [X44] : (s(bool,X44) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num','const.num.|48|_0')))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(backward_demodulation,[],[f21026,f20802])).
fof(f21055,plain,(
  $false | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(subsumption_resolution,[],[f21034,f242])).
fof(f21034,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl81 | $spl90)),
  inference(backward_demodulation,[],[f21026,f1298])).
fof(f21000,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,X28) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X26)),s('type.num.num',X27))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26)) ) | $spl110),
  inference(cnf_transformation,[],[f21000_D])).
fof(f21000_D,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,X28) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X26)),s('type.num.num',X27))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f20997,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0)))) | $spl106),
  inference(cnf_transformation,[],[f20997_D])).
fof(f20997_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0)))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f20994,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0)))) | s(bool,'T_0') = s(bool,X28) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s(bool,X28) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(X26)),s('type.num.num',X27)))) ) | ($spl11 | $spl24 | $spl30 | $spl90)),
  inference(backward_demodulation,[],[f20802,f20971])).
fof(f20998,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,X27) | s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num',sK1(X26)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26)) ) | $spl108),
  inference(cnf_transformation,[],[f20998_D])).
fof(f20998_D,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,X27) | s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num',sK1(X26)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f20991,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0)))) | s(bool,'T_0') = s(bool,X27) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num',sK1(X26))))) ) | ($spl11 | $spl24 | $spl30 | $spl90)),
  inference(backward_demodulation,[],[f20802,f20967])).
fof(f20995,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,X58) | s(bool,X58) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num',X56))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X57),s('type.num.num',X55)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X57)) ) | $spl104),
  inference(cnf_transformation,[],[f20995_D])).
fof(f20995_D,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,X58) | s(bool,X58) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num',X56))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X57),s('type.num.num',X55)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X57)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f20989,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0)))) | s(bool,'T_0') = s(bool,X58) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X57),s('type.num.num',X55)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X57) | s(bool,X58) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num',X56)))) ) | ($spl11 | $spl24 | $spl30 | $spl90)),
  inference(backward_demodulation,[],[f20802,f20952])).
fof(f20845,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,g(s(bool,'T_0'))) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | $spl102),
  inference(cnf_transformation,[],[f20845_D])).
fof(f20845_D,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,g(s(bool,'T_0'))) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f9632,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | ($spl24 | $spl30 | $spl58)),
  inference(duplicate_literal_removal,[],[f9561])).
fof(f9561,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl24 | $spl30 | $spl58)),
  inference(superposition,[],[f9165,f2970])).
fof(f20811,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,f(s(bool,'T_0'))) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | $spl100),
  inference(cnf_transformation,[],[f20811_D])).
fof(f20811_D,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,f(s(bool,'T_0'))) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f5080,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f5035])).
fof(f5035,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4243,f2970])).
fof(f20795,plain,(
  ( ! [X123,X121,X122,X120] : (s(fun(X120,bool),X121) = s(fun(X120,bool),X122) | s(bool,X123) = s(bool,f(s(bool,X123))) | s(bool,'T_0') = s(bool,g(s(X120,sK2(X121,X122,bool,X120))))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl60 | $spl81)),
  inference(subsumption_resolution,[],[f20794,f4252])).
fof(f20794,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,X123) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(fun(X120,bool),X121) = s(fun(X120,bool),X122) | s(bool,X123) = s(bool,f(s(bool,X123))) | s(bool,'T_0') = s(bool,g(s(X120,sK2(X121,X122,bool,X120))))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl60 | $spl81)),
  inference(backward_demodulation,[],[f20789,f4161])).
fof(f20789,plain,(
  ( ! [X41] : (s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl60 | $spl81)),
  inference(forward_demodulation,[],[f20771,f20744])).
fof(f20744,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl60 | $spl81)),
  inference(subsumption_resolution,[],[f20720,f19662])).
fof(f20720,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl11 | $spl24 | $spl30 | $spl60)),
  inference(backward_demodulation,[],[f20708,f1414])).
fof(f20708,plain,(
  ( ! [X41] : (s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) ) | ($spl11 | $spl24 | $spl30 | $spl60)),
  inference(subsumption_resolution,[],[f20694,f240])).
fof(f20694,plain,(
  ( ! [X41] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) ) | ($spl11 | $spl24 | $spl30 | $spl60)),
  inference(backward_demodulation,[],[f20681,f13449])).
fof(f20771,plain,(
  ( ! [X41] : (s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl60 | $spl81)),
  inference(backward_demodulation,[],[f20744,f20708])).
fof(f20788,plain,(
  $false | ($spl1 | $spl11 | $spl24 | $spl30 | $spl60 | $spl81)),
  inference(subsumption_resolution,[],[f20787,f134])).
fof(f20787,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl60 | $spl81)),
  inference(forward_demodulation,[],[f20786,f20753])).
fof(f20753,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl60 | $spl81)),
  inference(backward_demodulation,[],[f20744,f274])).
fof(f20786,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl60 | $spl81)),
  inference(forward_demodulation,[],[f20770,f20744])).
fof(f20770,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl60 | $spl81)),
  inference(backward_demodulation,[],[f20744,f19662])).
fof(f20781,plain,(
  $false | ($spl1 | $spl11 | $spl24 | $spl30 | $spl60 | $spl81)),
  inference(subsumption_resolution,[],[f20761,f242])).
fof(f20761,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl11 | $spl24 | $spl30 | $spl60 | $spl81)),
  inference(backward_demodulation,[],[f20744,f1298])).
fof(f4388,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) ) | $spl42),
  inference(cnf_transformation,[],[f4388_D])).
fof(f4388_D,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4308,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,f(s(bool,X23))) | s(bool,'F_0') = s(bool,X23)) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f4026])).
fof(f4026,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,f(s(bool,X23))) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f141])).
fof(f4418,plain,(
  ( ! [X8,X9] : (p(s(bool,X9)) | s(bool,'F_0') = s(bool,X8) | s(bool,X8) = s(bool,f(s(bool,X8)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4060,f145])).
fof(f4060,plain,(
  ( ! [X8,X9] : (p(s(bool,X9)) | s(bool,'F_0') = s(bool,X8) | s(bool,X8) = s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f150,f2970])).
fof(f4277,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,X1)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f4109])).
fof(f4109,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f141,f2970])).
fof(f20679,plain,(
  ( ! [X101,X100] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X100))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X100),s('type.num.num',X101))))) ) | $spl96),
  inference(cnf_transformation,[],[f20679_D])).
fof(f20679_D,plain,(
  ( ! [X101,X100] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X100))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X100),s('type.num.num',X101))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f20677,plain,(
  ( ! [X101,X102,X100] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X100))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101))))) | s(bool,X102) = s(bool,f(s(bool,X102))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X100),s('type.num.num',X101))))) ) | ($spl24 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f4360,f4373,f4343,f4324,f4318,f4010,f4011,f4309,f4450,f4456,f4466,f4160,f4166,f4167,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f18460,f4150,f4149,f19680,f19682,f19685,f19687,f4276,f4277,f4418,f4308,f4498,f4161,f4226,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f19710,f19712,f4253,f19715,f19717,f5080,f5051,f5114,f5115,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f19994,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f20042,f12066,f20047,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f20139,f17054,f19490,f19489,f19492,f621,f828,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f4412,f10889,f137,f149,f150,f261,f7508,f7514,f7515,f7516,f11978,f195,f1145,f1299,f12026,f6127,f6133,f619,f7500,f14974,f908,f4153])).
fof(f20676,plain,(
  $false | ($spl1 | $spl11 | $spl24 | $spl30 | $spl44 | $spl64 | $spl66 | $spl81)),
  inference(subsumption_resolution,[],[f19662,f20657])).
fof(f20657,plain,(
  ( ! [X25] : (p(s(bool,X25))) ) | ($spl11 | $spl24 | $spl30 | $spl44 | $spl64 | $spl66)),
  inference(global_subsumption,[],[f20656,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2066,f1908,f1907,f1889,f1906,f1669,f863,f402,f142,f140,f13889,f1298,f480,f389,f280,f274,f662,f136,f621,f828,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f137,f149,f150,f261,f7508,f7514,f7515,f7516,f195,f1145,f1299,f13847,f631,f619,f7500,f14974,f908])).
fof(f20656,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | s(bool,X25) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num',X26)))) ) | ($spl24 | $spl30 | $spl44 | $spl64 | $spl66)),
  inference(forward_demodulation,[],[f10851,f20648])).
fof(f20648,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(bool,X15)))) ) | ($spl24 | $spl30 | $spl44 | $spl64 | $spl66)),
  inference(global_subsumption,[],[f4389,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f4360,f4373,f4343,f4324,f4318,f4010,f4011,f4309,f4450,f4456,f4466,f4160,f4166,f4167,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f402,f142,f140,f13889,f1298,f480,f389,f280,f274,f662,f136,f18460,f4153,f4150,f4149,f19680,f19682,f19685,f19687,f4276,f4277,f4418,f4308,f4498,f4161,f4226,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f19710,f19712,f4253,f19715,f19717,f5080,f5051,f5114,f5115,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f19994,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f20042,f12066,f20047,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f20139,f17054,f19490,f19489,f19492,f621,f828,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f4412,f10889,f137,f149,f150,f261,f7508,f7514,f7515,f7516,f11978,f195,f1145,f1299,f13847,f12026,f6127,f6133,f619,f7500,f14974,f908])).
fof(f20665,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl44 | $spl64 | $spl66)),
  inference(subsumption_resolution,[],[f261,f20664])).
fof(f20664,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl11 | $spl24 | $spl30 | $spl44 | $spl64 | $spl66)),
  inference(subsumption_resolution,[],[f149,f20657])).
fof(f20663,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl44 | $spl64 | $spl66)),
  inference(subsumption_resolution,[],[f137,f20657])).
fof(f20662,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl44 | $spl64 | $spl66)),
  inference(subsumption_resolution,[],[f631,f20648])).
fof(f20647,plain,(
  $false | ($spl11 | $spl14 | $spl24 | $spl30 | $spl68 | $spl70)),
  inference(subsumption_resolution,[],[f631,f20533])).
fof(f20533,plain,(
  ( ! [X105] : (s(bool,X105) = s(bool,f(s(bool,X105)))) ) | ($spl11 | $spl14 | $spl24 | $spl30 | $spl68 | $spl70)),
  inference(global_subsumption,[],[f20532,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f4360,f4373,f4343,f4324,f4318,f4010,f4011,f4309,f4450,f4456,f4466,f4160,f4166,f4167,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f402,f142,f140,f14361,f1298,f480,f389,f280,f274,f662,f136,f18460,f4153,f4150,f4149,f19680,f19682,f19685,f19687,f4276,f4277,f4418,f4308,f4498,f4161,f4226,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f19710,f19712,f4253,f19715,f19717,f5080,f5051,f5114,f5115,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f19994,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f20042,f12066,f20047,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f20139,f17054,f19490,f19489,f19492,f621,f828,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f4412,f10889,f137,f149,f150,f261,f7508,f7514,f7515,f7516,f11978,f195,f1145,f1299,f14318,f12026,f6127,f6133,f619,f7500,f14974,f908])).
fof(f20532,plain,(
  ( ! [X103,X105,X102,X104] : (s(fun(X102,bool),X103) = s(fun(X102,bool),X104) | s(bool,X105) = s(bool,f(s(bool,X105))) | s(bool,X105) = s(bool,g(s(X102,sK2(X103,X104,bool,X102))))) ) | ($spl11 | $spl14 | $spl24 | $spl30 | $spl68 | $spl70)),
  inference(subsumption_resolution,[],[f4226,f20526])).
fof(f20526,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6)) ) | ($spl11 | $spl14 | $spl68 | $spl70)),
  inference(global_subsumption,[],[f1487,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2066,f1908,f1907,f1889,f1906,f1669,f863,f402,f142,f140,f14361,f1298,f480,f389,f280,f274,f662,f136,f621,f828,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f137,f149,f150,f261,f7508,f7514,f7515,f7516,f195,f1145,f1299,f14318,f631,f619,f7500,f14974,f908])).
fof(f20646,plain,(
  $false | ($spl11 | $spl14 | $spl24 | $spl30 | $spl68 | $spl70)),
  inference(global_subsumption,[],[f20645,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2066,f1908,f1907,f1889,f1906,f1669,f863,f402,f142,f140,f14361,f1298,f480,f389,f280,f274,f662,f136,f621,f828,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f137,f149,f150,f261,f7508,f7514,f7515,f7516,f195,f1145,f1299,f14318,f631,f619,f7500,f14974,f908,f1487,f20526,f20622,f20627,f20640])).
fof(f20640,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl14 | $spl68 | $spl70)),
  inference(subsumption_resolution,[],[f20639,f20526])).
fof(f20639,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl11 | $spl14 | $spl68 | $spl70)),
  inference(forward_demodulation,[],[f149,f20526])).
fof(f20627,plain,(
  ( ! [X127,X128,X126] : (s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(fun(bool,X128),X126) = s(fun(bool,X128),X127)) ) | ($spl11 | $spl14 | $spl68 | $spl70)),
  inference(duplicate_literal_removal,[],[f20626])).
fof(f20626,plain,(
  ( ! [X127,X128,X126] : (s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(fun(bool,X128),X126) = s(fun(bool,X128),X127)) ) | ($spl11 | $spl14 | $spl68 | $spl70)),
  inference(forward_demodulation,[],[f7500,f20526])).
fof(f20622,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | ($spl11 | $spl14 | $spl68 | $spl70)),
  inference(subsumption_resolution,[],[f20621,f134])).
fof(f20621,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | ($spl11 | $spl14 | $spl68 | $spl70)),
  inference(forward_demodulation,[],[f235,f20526])).
fof(f20645,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) | ($spl11 | $spl14 | $spl24 | $spl30 | $spl68 | $spl70)),
  inference(forward_demodulation,[],[f20644,f20533])).
fof(f20644,plain,(
  ( ! [X8] : (p(s(bool,f(s(bool,X8))))) ) | ($spl11 | $spl14 | $spl24 | $spl30 | $spl68 | $spl70)),
  inference(subsumption_resolution,[],[f6170,f20526])).
fof(f20643,plain,(
  $false | ($spl11 | $spl14 | $spl68 | $spl70)),
  inference(subsumption_resolution,[],[f20642,f134])).
fof(f20642,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl14 | $spl68 | $spl70)),
  inference(forward_demodulation,[],[f137,f20526])).
fof(f20641,plain,(
  $false | ($spl11 | $spl14 | $spl68 | $spl70)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2066,f1908,f1907,f1889,f1906,f1669,f863,f402,f142,f140,f14361,f1298,f480,f389,f280,f274,f662,f136,f621,f828,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f137,f149,f150,f261,f7508,f7514,f7515,f7516,f195,f1145,f1299,f14318,f631,f619,f7500,f14974,f908,f1487,f20526,f20622,f20627,f20640])).
fof(f20536,plain,(
  $false | ($spl1 | $spl11 | $spl14 | $spl68 | $spl70 | $spl81)),
  inference(subsumption_resolution,[],[f20535,f134])).
fof(f20535,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl11 | $spl14 | $spl68 | $spl70 | $spl81)),
  inference(forward_demodulation,[],[f19662,f20526])).
fof(f20525,plain,(
  $false | ($spl1 | $spl11 | $spl24 | $spl30 | $spl44 | $spl76 | $spl78 | $spl81)),
  inference(subsumption_resolution,[],[f19662,f20503])).
fof(f20503,plain,(
  ( ! [X25] : (p(s(bool,X25))) ) | ($spl11 | $spl24 | $spl30 | $spl44 | $spl76 | $spl78)),
  inference(global_subsumption,[],[f20502,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f1670,f1671,f1720,f2054,f2066,f1908,f1907,f1889,f1906,f1669,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f15668,f15712,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f195,f1145,f1299,f631])).
fof(f20502,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | s(bool,X25) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num',X26)))) ) | ($spl24 | $spl30 | $spl44 | $spl76 | $spl78)),
  inference(forward_demodulation,[],[f10851,f20494])).
fof(f20494,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(bool,X15)))) ) | ($spl24 | $spl30 | $spl44 | $spl76 | $spl78)),
  inference(global_subsumption,[],[f4389,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f4360,f4373,f4343,f4324,f4318,f4010,f4011,f4309,f4450,f4456,f4466,f4160,f4166,f4167,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f18460,f4153,f4150,f4149,f19680,f19682,f19685,f19687,f4276,f4277,f4418,f4308,f4498,f4161,f4226,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f19710,f19712,f4253,f19715,f19717,f5080,f5051,f5114,f5115,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f19994,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f20042,f12066,f20047,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f20139,f17054,f19490,f19489,f19492,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f4412,f10889,f15668,f15712,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f11978,f195,f1145,f1299,f12026,f6127,f6133])).
fof(f20511,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl44 | $spl76 | $spl78)),
  inference(subsumption_resolution,[],[f261,f20510])).
fof(f20510,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl11 | $spl24 | $spl30 | $spl44 | $spl76 | $spl78)),
  inference(subsumption_resolution,[],[f149,f20503])).
fof(f20509,plain,(
  $false | ($spl11 | $spl24 | $spl30 | $spl44 | $spl76 | $spl78)),
  inference(subsumption_resolution,[],[f137,f20503])).
fof(f20493,plain,(
  $false | ($spl11 | $spl14 | $spl24 | $spl30 | $spl90 | $spl92)),
  inference(subsumption_resolution,[],[f631,f20378])).
fof(f20378,plain,(
  ( ! [X105] : (s(bool,X105) = s(bool,f(s(bool,X105)))) ) | ($spl11 | $spl14 | $spl24 | $spl30 | $spl90 | $spl92)),
  inference(global_subsumption,[],[f20377,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f4360,f4373,f4343,f4324,f4318,f4010,f4011,f4309,f4450,f4456,f4466,f4160,f4166,f4167,f4252,f4225,f4229,f4243,f4245,f863,f5052,f6117,f6405,f12105,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f19226,f18460,f4153,f4150,f4149,f19680,f19682,f19685,f19687,f4276,f4277,f4418,f4308,f4498,f4161,f4226,f4107,f4108,f4502,f4499,f4491,f4420,f4414,f4413,f4350,f4163,f2917,f4490,f19710,f19712,f4253,f19715,f19717,f5080,f5051,f5114,f5115,f5116,f19761,f19767,f19769,f5388,f19772,f5662,f5649,f5926,f19777,f19779,f19782,f19785,f19787,f19789,f6156,f6153,f19794,f19797,f19800,f19802,f19804,f19806,f6134,f6170,f6132,f6173,f6176,f6178,f19819,f6128,f6180,f6125,f19828,f19830,f19832,f6120,f6119,f6198,f19837,f19839,f19843,f6331,f6332,f6404,f6462,f6452,f6479,f6483,f6514,f19853,f6744,f19856,f6939,f19859,f7509,f19862,f7503,f19865,f7502,f7501,f8084,f19870,f9975,f9954,f10205,f19896,f10451,f19899,f19903,f10839,f10546,f10560,f19909,f10566,f10838,f10836,f19916,f19918,f19920,f10835,f19923,f10618,f19926,f10647,f10851,f10649,f10650,f10651,f10652,f10653,f10654,f10655,f10660,f19938,f10664,f19942,f19944,f19946,f10871,f19950,f10699,f10700,f10701,f10702,f10703,f10704,f10833,f19965,f19967,f19969,f10757,f10759,f10760,f19974,f10769,f10770,f10772,f19980,f10777,f19983,f19985,f10781,f10783,f10784,f10785,f10786,f19994,f10829,f20025,f11929,f11930,f11943,f11975,f12024,f20035,f11960,f11962,f20039,f12064,f20042,f12066,f20047,f20049,f12156,f12841,f12850,f20054,f12852,f13146,f13154,f20059,f13156,f13974,f13984,f20066,f13986,f14575,f14585,f20072,f14587,f14987,f14957,f15277,f15296,f20106,f15298,f15954,f15964,f20130,f15966,f20136,f17052,f20139,f17054,f19490,f19489,f19492,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f4412,f10889,f19190,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f11978,f195,f1145,f1299,f12026,f6127,f6133])).
fof(f20377,plain,(
  ( ! [X103,X105,X102,X104] : (s(fun(X102,bool),X103) = s(fun(X102,bool),X104) | s(bool,X105) = s(bool,f(s(bool,X105))) | s(bool,X105) = s(bool,g(s(X102,sK2(X103,X104,bool,X102))))) ) | ($spl11 | $spl14 | $spl24 | $spl30 | $spl90 | $spl92)),
  inference(subsumption_resolution,[],[f4226,f20371])).
fof(f20371,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6)) ) | ($spl11 | $spl14 | $spl90 | $spl92)),
  inference(global_subsumption,[],[f1487,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f631,f1670,f1671,f1720,f2054,f2066,f1908,f1907,f1889,f1906,f1669,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f19226,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f19190,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f195,f1145,f1299])).
fof(f20492,plain,(
  $false | ($spl11 | $spl14 | $spl24 | $spl30 | $spl90 | $spl92)),
  inference(global_subsumption,[],[f20491,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f631,f1670,f1671,f1720,f2054,f2066,f1908,f1907,f1889,f1906,f1669,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f19226,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f19190,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f195,f1145,f1299,f1487,f20371,f20467,f20483,f20486])).
fof(f20486,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl14 | $spl90 | $spl92)),
  inference(subsumption_resolution,[],[f20485,f20371])).
fof(f20485,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl11 | $spl14 | $spl90 | $spl92)),
  inference(forward_demodulation,[],[f149,f20371])).
fof(f20483,plain,(
  ( ! [X127,X128,X126] : (s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(fun(bool,X128),X126) = s(fun(bool,X128),X127)) ) | ($spl11 | $spl14 | $spl90 | $spl92)),
  inference(duplicate_literal_removal,[],[f20482])).
fof(f20482,plain,(
  ( ! [X127,X128,X126] : (s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(fun(bool,X128),X126) = s(fun(bool,X128),X127)) ) | ($spl11 | $spl14 | $spl90 | $spl92)),
  inference(forward_demodulation,[],[f7500,f20371])).
fof(f20467,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | ($spl11 | $spl14 | $spl90 | $spl92)),
  inference(subsumption_resolution,[],[f20466,f134])).
fof(f20466,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | ($spl11 | $spl14 | $spl90 | $spl92)),
  inference(forward_demodulation,[],[f235,f20371])).
fof(f20491,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) | ($spl11 | $spl14 | $spl24 | $spl30 | $spl90 | $spl92)),
  inference(forward_demodulation,[],[f20490,f20378])).
fof(f20490,plain,(
  ( ! [X8] : (p(s(bool,f(s(bool,X8))))) ) | ($spl11 | $spl14 | $spl24 | $spl30 | $spl90 | $spl92)),
  inference(subsumption_resolution,[],[f6170,f20371])).
fof(f20489,plain,(
  $false | ($spl11 | $spl14 | $spl90 | $spl92)),
  inference(subsumption_resolution,[],[f20488,f134])).
fof(f20488,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl14 | $spl90 | $spl92)),
  inference(forward_demodulation,[],[f137,f20371])).
fof(f20487,plain,(
  $false | ($spl11 | $spl14 | $spl90 | $spl92)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f637,f636,f638,f666,f766,f867,f1141,f1142,f907,f1196,f1213,f1214,f1215,f1440,f1450,f1486,f1489,f631,f1670,f1671,f1720,f2054,f2066,f1908,f1907,f1889,f1906,f1669,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f19226,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f19190,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f195,f1145,f1299,f1487,f20371,f20467,f20483,f20486])).
fof(f20381,plain,(
  $false | ($spl1 | $spl11 | $spl14 | $spl81 | $spl90 | $spl92)),
  inference(subsumption_resolution,[],[f20380,f134])).
fof(f20380,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl11 | $spl14 | $spl81 | $spl90 | $spl92)),
  inference(forward_demodulation,[],[f19662,f20371])).
fof(f20370,plain,(
  $false | ($spl1 | $spl14 | $spl18 | $spl81)),
  inference(subsumption_resolution,[],[f20296,f134])).
fof(f20296,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl14 | $spl18 | $spl81)),
  inference(backward_demodulation,[],[f20224,f19662])).
fof(f20224,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl14 | $spl18)),
  inference(condensation,[],[f20223])).
fof(f20223,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X23)) ) | ($spl14 | $spl18)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f1487,f5019,f4959,f5018,f19728,f19730,f5020,f5022,f19738,f19740,f5007,f19841,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f1650,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f195,f1145,f1299,f20222])).
fof(f20324,plain,(
  $false | ($spl14 | $spl18)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f1487,f5019,f4959,f5018,f19728,f19730,f5020,f5022,f19738,f19740,f5007,f19841,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f1650,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f195,f1145,f1299,f20222,f20224,f20310,f20308,f20321,f20323])).
fof(f20323,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | ($spl14 | $spl18)),
  inference(subsumption_resolution,[],[f20256,f134])).
fof(f20256,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | ($spl14 | $spl18)),
  inference(backward_demodulation,[],[f20224,f235])).
fof(f20321,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl14 | $spl18)),
  inference(subsumption_resolution,[],[f20255,f134])).
fof(f20255,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl14 | $spl18)),
  inference(backward_demodulation,[],[f20224,f650])).
fof(f20308,plain,(
  ( ! [X127,X128,X126] : (s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(fun(bool,X128),X126) = s(fun(bool,X128),X127)) ) | ($spl14 | $spl18)),
  inference(duplicate_literal_removal,[],[f20234])).
fof(f20234,plain,(
  ( ! [X127,X128,X126] : (s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(X128,f(s(bool,'T_0'))) != s(X128,g(s(bool,'T_0'))) | s(fun(bool,X128),X126) = s(fun(bool,X128),X127)) ) | ($spl14 | $spl18)),
  inference(backward_demodulation,[],[f20224,f7500])).
fof(f20310,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl14 | $spl18)),
  inference(subsumption_resolution,[],[f20226,f20224])).
fof(f20226,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl14 | $spl18)),
  inference(backward_demodulation,[],[f20224,f149])).
fof(f20322,plain,(
  $false | ($spl14 | $spl18)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f1487,f5019,f4959,f5018,f19728,f19730,f5020,f5022,f19738,f19740,f5007,f19841,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f1650,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f195,f1145,f1299,f20222,f20224,f20310,f20308,f20321])).
fof(f20320,plain,(
  $false | ($spl14 | $spl18 | $spl24 | $spl30)),
  inference(global_subsumption,[],[f20319,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f1487,f5019,f4959,f5018,f19728,f19730,f5020,f5022,f19738,f19740,f5007,f19841,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f1650,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f195,f1145,f1299,f20222,f20224,f20310,f20308])).
fof(f20319,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl14 | $spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f20240,f20224])).
fof(f20240,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl14 | $spl18 | $spl24 | $spl30)),
  inference(backward_demodulation,[],[f20224,f11978])).
fof(f20318,plain,(
  $false | ($spl14 | $spl18 | $spl24 | $spl30)),
  inference(global_subsumption,[],[f20317,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f1487,f5019,f4959,f5018,f19728,f19730,f5020,f5022,f19738,f19740,f5007,f19841,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f1650,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f195,f1145,f1299,f20222,f20224,f20310,f20308])).
fof(f20317,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl14 | $spl18 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f20239,f12026])).
fof(f20239,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl14 | $spl18 | $spl24 | $spl30)),
  inference(backward_demodulation,[],[f20224,f10889])).
fof(f20312,plain,(
  $false | ($spl14 | $spl18)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f1487,f5019,f4959,f5018,f19728,f19730,f5020,f5022,f19738,f19740,f5007,f19841,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f1650,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f195,f1145,f1299,f20222,f20224,f20310,f20308])).
fof(f20311,plain,(
  $false | ($spl14 | $spl18)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f300,f340,f429,f507,f766,f867,f1141,f1142,f907,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f1487,f5019,f4959,f5018,f19728,f19730,f5020,f5022,f19738,f19740,f5007,f19841,f621,f828,f908,f14974,f235,f650,f1144,f2104,f453,f512,f517,f1612,f1628,f1811,f1843,f1964,f1980,f459,f534,f535,f1515,f1502,f1400,f454,f1415,f478,f1979,f1414,f1695,f458,f1716,f1712,f1715,f189,f471,f477,f1840,f1841,f1269,f1274,f1627,f1816,f1650,f137,f149,f150,f261,f619,f7500,f7508,f7514,f7515,f7516,f195,f1145,f1299,f20222,f20224,f20310])).
fof(f20309,plain,(
  $false | ($spl14 | $spl18)),
  inference(subsumption_resolution,[],[f20225,f134])).
fof(f20225,plain,(
  ~p(s(bool,'T_0')) | ($spl14 | $spl18)),
  inference(backward_demodulation,[],[f20224,f137])).
fof(f20195,plain,(
  $false | ($spl1 | $spl44 | $spl62)),
  inference(global_subsumption,[],[f4389,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f2104,f1144,f650,f235,f14974,f908,f828,f621])).
fof(f14772,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f242,f13488])).
fof(f14771,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl62),
  inference(unit_resulting_resolution,[],[f240,f13488])).
fof(f14859,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl62),
  inference(duplicate_literal_removal,[],[f14776])).
fof(f14776,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl62),
  inference(superposition,[],[f13488,f141])).
fof(f14883,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl62),
  inference(subsumption_resolution,[],[f14812,f242])).
fof(f14812,plain,(
  ( ! [X66] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X66),s('type.num.num',X66)))) ) | $spl62),
  inference(superposition,[],[f13488,f261])).
fof(f14886,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl62),
  inference(condensation,[],[f14885])).
fof(f14885,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) ) | $spl62),
  inference(subsumption_resolution,[],[f14825,f141])).
fof(f14825,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) ) | $spl62),
  inference(superposition,[],[f13488,f141])).
fof(f14889,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl62),
  inference(condensation,[],[f14888])).
fof(f14888,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) ) | $spl62),
  inference(condensation,[],[f14887])).
fof(f14887,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) ) | $spl62),
  inference(subsumption_resolution,[],[f14827,f300])).
fof(f14827,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) ) | $spl62),
  inference(superposition,[],[f13488,f300])).
fof(f14892,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl62),
  inference(condensation,[],[f14891])).
fof(f14891,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl62),
  inference(condensation,[],[f14890])).
fof(f14890,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X8)) ) | $spl62),
  inference(subsumption_resolution,[],[f14828,f300])).
fof(f14828,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X8)) ) | $spl62),
  inference(superposition,[],[f13488,f300])).
fof(f14848,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl62),
  inference(equality_resolution,[],[f13488])).
fof(f20194,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19492,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f2104,f1144,f650,f235,f14974,f908,f828,f621])).
fof(f20193,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f2104,f1144,f650,f235,f14974,f908,f828,f621])).
fof(f20192,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f621,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f2104,f1144,f650,f235,f14974,f908,f828])).
fof(f20191,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f828,f621,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f2104,f1144,f650,f235,f14974,f908])).
fof(f20190,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f908,f828,f621,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f2104,f1144,f650,f235,f14974])).
fof(f20189,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f2104,f1144,f650,f235])).
fof(f20188,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f140,f1298,f480,f389,f280,f274,f662,f235,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f2104,f1144,f650])).
fof(f20187,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f140,f1298,f480,f389,f280,f274,f650,f662,f235,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f2104,f1144])).
fof(f20186,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f140,f1298,f480,f389,f280,f274,f650,f662,f1144,f235,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f2104])).
fof(f20185,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19489,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f140,f1298,f480,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453])).
fof(f20184,plain,(
  $false | ($spl1 | $spl62 | $spl81)),
  inference(global_subsumption,[],[f18885,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f140,f1298,f480,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453])).
fof(f18885,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',X0))))) ) | $spl81),
  inference(subsumption_resolution,[],[f18804,f134])).
fof(f18804,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',X0))))) ) | $spl81),
  inference(superposition,[],[f195,f18651])).
fof(f18651,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | $spl81),
  inference(superposition,[],[f274,f17790])).
fof(f20183,plain,(
  $false | ($spl1 | $spl62 | $spl81)),
  inference(global_subsumption,[],[f18169,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f140,f1298,f480,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453])).
fof(f18169,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',X2))))) ) | $spl81),
  inference(superposition,[],[f189,f17790])).
fof(f20182,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f140,f1298,f480,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453])).
fof(f20181,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f140,f1298,f480,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512])).
fof(f20180,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f512,f140,f1298,f480,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612,f517])).
fof(f20179,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f517,f512,f140,f1298,f480,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628,f1612])).
fof(f20178,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f517,f512,f140,f1298,f480,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811,f1628])).
fof(f20177,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f517,f512,f140,f1628,f1298,f480,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843,f1811])).
fof(f20176,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1811,f517,f512,f140,f1628,f1298,f480,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964,f1843])).
fof(f20175,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1843,f1811,f517,f512,f140,f1628,f1298,f480,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980,f1964])).
fof(f20174,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1964,f1843,f1811,f517,f512,f140,f1628,f1298,f480,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459,f1980])).
fof(f20173,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1298,f480,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534,f459])).
fof(f20172,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1298,f480,f459,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535,f534])).
fof(f20171,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1298,f534,f480,f459,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515,f535])).
fof(f20170,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1298,f535,f534,f480,f459,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502,f1515])).
fof(f20169,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1515,f1298,f535,f534,f480,f459,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400,f1502])).
fof(f20168,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1515,f1298,f535,f534,f480,f459,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454,f1400])).
fof(f20167,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1515,f1298,f535,f534,f480,f459,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f1400,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415,f454])).
fof(f20166,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1515,f1298,f535,f534,f480,f459,f454,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f1400,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478,f1415])).
fof(f20165,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1515,f1415,f1298,f535,f534,f480,f459,f454,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f1400,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979,f478])).
fof(f20164,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f459,f454,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f1400,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414,f1979])).
fof(f20163,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f459,f454,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f1400,f1979,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695,f1414])).
fof(f20162,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f459,f454,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458,f1695])).
fof(f20161,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f459,f454,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f1695,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716,f458])).
fof(f20160,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f459,f458,f454,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f1695,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712,f1716])).
fof(f20159,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f459,f458,f454,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f1695,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715,f1712])).
fof(f20158,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f459,f458,f454,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f1695,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189,f1715])).
fof(f20157,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f459,f458,f454,f453,f389,f280,f274,f2104,f650,f662,f1144,f235,f136,f1612,f1695,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471,f189])).
fof(f20156,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1612,f1695,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477,f471])).
fof(f20155,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1612,f1695,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840,f477])).
fof(f20154,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1612,f1695,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841,f1840])).
fof(f20153,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1612,f1840,f1695,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269,f1841])).
fof(f20152,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1612,f1841,f1840,f1695,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274,f1269])).
fof(f20151,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1612,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627,f1274])).
fof(f20150,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1612,f1274,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816,f1627])).
fof(f20149,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137,f1816])).
fof(f20148,plain,(
  $false | ($spl1 | $spl20 | $spl62)),
  inference(global_subsumption,[],[f2191,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20147,plain,(
  $false | ($spl1 | $spl62 | $spl81)),
  inference(global_subsumption,[],[f18176,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20146,plain,(
  $false | ($spl1 | $spl18 | $spl62)),
  inference(global_subsumption,[],[f1650,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20145,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4412,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20144,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10889,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20143,plain,(
  $false | ($spl1 | $spl62 | $spl90)),
  inference(global_subsumption,[],[f19190,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20142,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19490,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20141,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f17054,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20140,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20139,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20138,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f17052,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20137,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20136,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20132,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f15966,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20131,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20130,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20129,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f15964,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20128,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f15954,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20127,plain,(
  $false | ($spl1 | $spl62 | $spl76)),
  inference(global_subsumption,[],[f15668,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20126,plain,(
  $false | ($spl1 | $spl62 | $spl78)),
  inference(global_subsumption,[],[f15712,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20125,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149,f137])).
fof(f20124,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f137,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150,f149])).
fof(f20123,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f149,f137,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261,f150])).
fof(f20122,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f150,f149,f137,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619,f261])).
fof(f20121,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f261,f150,f149,f137,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500,f619])).
fof(f20120,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508,f7500])).
fof(f20119,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7500,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514,f7508])).
fof(f20118,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7500,f7508,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515,f7514])).
fof(f20117,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7508,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516,f7515])).
fof(f20116,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7515,f7508,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195,f7516])).
fof(f20115,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f11978,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195])).
fof(f20114,plain,(
  $false | ($spl1 | $spl22 | $spl62)),
  inference(global_subsumption,[],[f14907,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195])).
fof(f14907,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,f(s(fun(X6,fun(bool,bool)),X9))) | s(fun(fun(X6,fun(bool,bool)),bool),X7) = s(fun(fun(X6,fun(bool,bool)),bool),X8) | s(bool,'T_0') = s(bool,g(s(fun(X6,fun(bool,bool)),sK2(X7,X8,bool,fun(X6,fun(bool,bool))))))) ) | $spl22),
  inference(superposition,[],[f621,f2854])).
fof(f2854,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(bool,bool)),X1) = s(fun(X0,fun(bool,bool)),X2)) ) | $spl22),
  inference(unit_resulting_resolution,[],[f2807,f194])).
fof(f20113,plain,(
  $false | ($spl1 | $spl22 | $spl62)),
  inference(global_subsumption,[],[f14909,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195])).
fof(f14909,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'F_0') != s(bool,f(s(fun(X14,fun(X15,fun(bool,bool))),X18))) | s(fun(fun(X14,fun(X15,fun(bool,bool))),bool),X16) = s(fun(fun(X14,fun(X15,fun(bool,bool))),bool),X17) | s(bool,'T_0') = s(bool,g(s(fun(X14,fun(X15,fun(bool,bool))),sK2(X16,X17,bool,fun(X14,fun(X15,fun(bool,bool)))))))) ) | $spl22),
  inference(superposition,[],[f621,f3120])).
fof(f3120,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X1,fun(bool,bool))),X2) = s(fun(X0,fun(X1,fun(bool,bool))),X3)) ) | $spl22),
  inference(unit_resulting_resolution,[],[f2854,f194])).
fof(f20112,plain,(
  $false | ($spl1 | $spl22 | $spl62)),
  inference(global_subsumption,[],[f14911,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195])).
fof(f14911,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'F_0') != s(bool,f(s(fun(X24,fun(X25,fun(X26,fun(bool,bool)))),X29))) | s(fun(fun(X24,fun(X25,fun(X26,fun(bool,bool)))),bool),X27) = s(fun(fun(X24,fun(X25,fun(X26,fun(bool,bool)))),bool),X28) | s(bool,'T_0') = s(bool,g(s(fun(X24,fun(X25,fun(X26,fun(bool,bool)))),sK2(X27,X28,bool,fun(X24,fun(X25,fun(X26,fun(bool,bool))))))))) ) | $spl22),
  inference(superposition,[],[f621,f3739])).
fof(f3739,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X1,fun(X2,fun(bool,bool)))),X3) = s(fun(X0,fun(X1,fun(X2,fun(bool,bool)))),X4)) ) | $spl22),
  inference(unit_resulting_resolution,[],[f3120,f194])).
fof(f20111,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145,f195])).
fof(f20110,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299,f1145])).
fof(f20109,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772,f1299])).
fof(f20108,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f15298,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20107,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20106,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20105,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f15296,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20104,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f15277,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20086,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f14957,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20085,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f14987,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20084,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20083,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20083,plain,(
  ( ! [X19] : (s(bool,'T_0') = s(bool,X19) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))))) ) | ($spl24 | $spl30 | $spl62)),
  inference(subsumption_resolution,[],[f14833,f507])).
fof(f14833,plain,(
  ( ! [X19,X18] : (s(bool,X19) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num',X18))) | s(bool,'T_0') = s(bool,X19) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))))) ) | ($spl24 | $spl30 | $spl62)),
  inference(superposition,[],[f13488,f4252])).
fof(f20082,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20081,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20081,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))))) ) | ($spl24 | $spl30 | $spl62)),
  inference(condensation,[],[f20080])).
fof(f20080,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl24 | $spl30 | $spl62)),
  inference(subsumption_resolution,[],[f14830,f300])).
fof(f14830,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl24 | $spl30 | $spl62)),
  inference(superposition,[],[f13488,f2970])).
fof(f20079,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20078,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20078,plain,(
  ( ! [X65] : (s(bool,X65) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,X65) = s(bool,f(s(bool,X65)))) ) | ($spl24 | $spl30 | $spl62)),
  inference(subsumption_resolution,[],[f14811,f242])).
fof(f14811,plain,(
  ( ! [X64,X65] : (s(bool,X65) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X64),s('type.num.num',X64))) | s(bool,X65) = s(bool,f(s(bool,X65)))) ) | ($spl24 | $spl30 | $spl62)),
  inference(superposition,[],[f13488,f4252])).
fof(f20077,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20076,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20076,plain,(
  ( ! [X15,X16] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num',X16))) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | ($spl24 | $spl30 | $spl62)),
  inference(subsumption_resolution,[],[f14784,f4466])).
fof(f14784,plain,(
  ( ! [X15,X16] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num',X16))) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | ($spl24 | $spl30 | $spl62)),
  inference(superposition,[],[f13488,f4252])).
fof(f20075,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f14854,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f14854,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | ($spl24 | $spl30 | $spl62)),
  inference(duplicate_literal_removal,[],[f14782])).
fof(f14782,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl24 | $spl30 | $spl62)),
  inference(superposition,[],[f13488,f2970])).
fof(f20074,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f14587,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20073,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20072,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20071,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f14585,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20070,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f14575,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20069,plain,(
  $false | ($spl1 | $spl62 | $spl68)),
  inference(global_subsumption,[],[f14318,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20068,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f13986,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20067,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20066,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20065,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f13984,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20064,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f13974,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20063,plain,(
  $false | ($spl1 | $spl62 | $spl64)),
  inference(global_subsumption,[],[f13847,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20062,plain,(
  $false | ($spl1 | $spl60 | $spl62)),
  inference(global_subsumption,[],[f13449,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20061,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f13156,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20060,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20059,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20058,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f13154,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20057,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f13146,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20056,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f12852,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20055,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20054,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20053,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f12850,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20052,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f12841,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20051,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f12156,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20050,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20049,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20048,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20047,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20044,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f12066,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20043,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20042,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20041,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f12064,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20040,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20039,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20038,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f11962,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20037,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f11960,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20036,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20035,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20034,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f12026,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20033,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f12024,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20032,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f11975,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20031,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f11943,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20030,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f11930,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20029,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f11929,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20027,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f11913,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f20026,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f20025,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19998,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl58 | $spl62)),
  inference(global_subsumption,[],[f19997,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19997,plain,(
  ( ! [X132] : (s(bool,X132) != s(bool,g(s(bool,'T_0'))) | s(bool,X132) = s(bool,f(s(bool,X132)))) ) | ($spl24 | $spl30 | $spl58)),
  inference(subsumption_resolution,[],[f10825,f242])).
fof(f10825,plain,(
  ( ! [X132,X131] : (s(bool,X132) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X131),s('type.num.num',X131))) | s(bool,X132) = s(bool,f(s(bool,X132)))) ) | ($spl24 | $spl30 | $spl58)),
  inference(superposition,[],[f9165,f4252])).
fof(f19996,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10829,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19995,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19994,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19993,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19992,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19991,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10786,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19990,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10785,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19989,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10784,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19988,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10783,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19987,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10781,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19986,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19985,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19984,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19983,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19982,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10777,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19981,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19980,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19978,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10772,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19977,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10770,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19976,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10769,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19975,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19974,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19973,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10760,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19972,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10759,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19971,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10757,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19970,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19969,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19968,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19967,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19966,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19965,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19962,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10833,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19959,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19958,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19957,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10704,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19956,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10703,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19955,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10702,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19954,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10701,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19953,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10700,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19952,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10699,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19951,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19950,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19949,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl58 | $spl62)),
  inference(global_subsumption,[],[f10875,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f10875,plain,(
  ( ! [X142,X143] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X142))))) | s(bool,X142) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X143),s('type.num.num',X143)))) ) | ($spl24 | $spl30 | $spl58)),
  inference(subsumption_resolution,[],[f10874,f507])).
fof(f10874,plain,(
  ( ! [X142,X143] : (s(bool,'T_0') != s(bool,X142) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X142))))) | s(bool,X142) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X143),s('type.num.num',X143)))) ) | ($spl24 | $spl30 | $spl58)),
  inference(forward_demodulation,[],[f10695,f2849])).
fof(f10695,plain,(
  ( ! [X142,X143] : (s(bool,X142) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X142))))) | s(bool,X142) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X143),s('type.num.num',X143)))) ) | ($spl24 | $spl30 | $spl58)),
  inference(superposition,[],[f9629,f4252])).
fof(f9629,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X36)))) ) | ($spl24 | $spl30 | $spl58)),
  inference(duplicate_literal_removal,[],[f9576])).
fof(f9576,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X36))) | s(bool,'T_0') = s(bool,f(s(bool,X36)))) ) | ($spl24 | $spl30 | $spl58)),
  inference(superposition,[],[f9165,f4245])).
fof(f19948,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10871,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19947,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19946,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19945,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19944,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19943,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19942,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19940,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10664,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19939,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19938,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19937,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10660,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19936,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10655,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19935,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10654,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19934,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10653,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19933,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10652,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19932,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10651,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19931,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10650,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19930,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10649,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19929,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10851,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19928,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10647,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19927,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19926,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19925,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10618,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19924,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19923,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19922,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10835,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19921,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19920,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19919,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19918,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19917,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19916,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19915,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10836,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19912,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10838,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19911,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10566,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19910,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19909,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19908,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10560,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19907,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10546,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19906,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10839,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19905,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10530,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19904,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19903,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19900,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19899,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19898,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10451,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19897,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19896,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19895,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f10205,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19890,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl58 | $spl62)),
  inference(global_subsumption,[],[f10101,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f10101,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,f(s(bool,X1))) | s(bool,X2) = s(bool,f(s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | ($spl24 | $spl30 | $spl58)),
  inference(superposition,[],[f9629,f2970])).
fof(f19889,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl58 | $spl62)),
  inference(global_subsumption,[],[f19888,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19888,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') = s(bool,f(s(bool,sK2(X54,X55,X56,bool)))) | s(fun(bool,X56),X54) = s(fun(bool,X56),X55) | s(X56,f(s(bool,'F_0'))) != s(X56,g(s(bool,'F_0')))) ) | ($spl24 | $spl30 | $spl58)),
  inference(subsumption_resolution,[],[f10100,f2849])).
fof(f10100,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,sK2(X54,X55,X56,bool)))) | s(fun(bool,X56),X54) = s(fun(bool,X56),X55) | s(X56,f(s(bool,'F_0'))) != s(X56,g(s(bool,'F_0')))) ) | ($spl24 | $spl30 | $spl58)),
  inference(superposition,[],[f9629,f619])).
fof(f19887,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl58 | $spl62)),
  inference(global_subsumption,[],[f19886,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19886,plain,(
  ( ! [X41] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X41))))) | s(bool,X41) = s(bool,f(s(bool,X41)))) ) | ($spl24 | $spl30 | $spl58)),
  inference(subsumption_resolution,[],[f10087,f2849])).
fof(f10087,plain,(
  ( ! [X41] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X41))))) | s(bool,X41) = s(bool,f(s(bool,X41)))) ) | ($spl24 | $spl30 | $spl58)),
  inference(superposition,[],[f9629,f4245])).
fof(f19885,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl58 | $spl62)),
  inference(global_subsumption,[],[f19884,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19884,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X38))))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | ($spl24 | $spl30 | $spl58)),
  inference(subsumption_resolution,[],[f19883,f300])).
fof(f19883,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X38) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X38))))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | ($spl24 | $spl30 | $spl58)),
  inference(forward_demodulation,[],[f10085,f2849])).
fof(f10085,plain,(
  ( ! [X39,X38] : (s(bool,X38) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X38))))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | ($spl24 | $spl30 | $spl58)),
  inference(superposition,[],[f9629,f2970])).
fof(f19882,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl58 | $spl62)),
  inference(global_subsumption,[],[f19881,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19881,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') = s(bool,f(s(bool,X13))) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X13) = s(bool,X14)) ) | ($spl24 | $spl30 | $spl58)),
  inference(subsumption_resolution,[],[f10071,f2849])).
fof(f10071,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X13))) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X13) = s(bool,X14)) ) | ($spl24 | $spl30 | $spl58)),
  inference(superposition,[],[f9629,f2970])).
fof(f19880,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl58 | $spl62)),
  inference(global_subsumption,[],[f10146,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f10146,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X11))) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | ($spl24 | $spl30 | $spl58)),
  inference(subsumption_resolution,[],[f10070,f6117])).
fof(f10070,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X11))) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl24 | $spl30 | $spl58)),
  inference(superposition,[],[f9629,f2970])).
fof(f19879,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f9954,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19878,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f9975,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19873,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl58 | $spl62)),
  inference(global_subsumption,[],[f9590,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f9590,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | ($spl24 | $spl30 | $spl58)),
  inference(superposition,[],[f9165,f2970])).
fof(f19872,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl58 | $spl62)),
  inference(global_subsumption,[],[f9632,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19871,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19870,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19869,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f8084,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19868,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f7501,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19867,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f7502,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19866,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19865,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19864,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f7503,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19863,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19862,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19861,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f7509,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19860,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19859,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19858,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6939,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19857,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19856,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19855,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6744,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19854,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19853,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19852,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6514,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19851,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6483,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19850,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6479,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19849,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6452,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19848,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6462,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19847,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6404,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19846,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6332,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19845,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6331,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19844,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19843,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19842,plain,(
  $false | ($spl1 | $spl18 | $spl62)),
  inference(global_subsumption,[],[f19841,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19840,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19839,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19838,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19837,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19836,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6198,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19835,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6119,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19834,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6120,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19833,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19832,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19831,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19830,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19829,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19828,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19827,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6125,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19826,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19825,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19823,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6180,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19822,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6127,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19821,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6128,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19820,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19819,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19818,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6178,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19817,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6176,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19816,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6173,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19815,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6132,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19814,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6170,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19813,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6133,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19812,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6134,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19811,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19810,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19808,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6027,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19807,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19806,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19805,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19804,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19803,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19802,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19801,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19800,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19798,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19797,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19795,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19794,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19793,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6153,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19792,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f6156,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19791,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f5983,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19790,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19789,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19788,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19787,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19786,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19785,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19783,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19782,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19780,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19779,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19778,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19777,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19776,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f5926,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19775,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f5649,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19774,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f5662,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19773,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19772,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19771,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f5388,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19770,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19769,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19768,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19767,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19762,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19761,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19760,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f5116,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19759,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f5115,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19758,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f5114,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19749,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f5051,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19748,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f5080,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19747,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19746,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19745,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19744,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19743,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f5008,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19742,plain,(
  $false | ($spl1 | $spl18 | $spl62)),
  inference(global_subsumption,[],[f5007,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19741,plain,(
  $false | ($spl1 | $spl18 | $spl62)),
  inference(global_subsumption,[],[f19740,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19739,plain,(
  $false | ($spl1 | $spl18 | $spl62)),
  inference(global_subsumption,[],[f19738,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19737,plain,(
  $false | ($spl1 | $spl18 | $spl62)),
  inference(global_subsumption,[],[f5022,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19736,plain,(
  $false | ($spl1 | $spl18 | $spl62)),
  inference(global_subsumption,[],[f5020,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19735,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19734,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19733,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19732,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19731,plain,(
  $false | ($spl1 | $spl18 | $spl62)),
  inference(global_subsumption,[],[f19730,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19729,plain,(
  $false | ($spl1 | $spl18 | $spl62)),
  inference(global_subsumption,[],[f19728,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19727,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19726,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19725,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19724,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19722,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f5013,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19721,plain,(
  $false | ($spl1 | $spl18 | $spl62)),
  inference(global_subsumption,[],[f5018,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19720,plain,(
  $false | ($spl1 | $spl18 | $spl62)),
  inference(global_subsumption,[],[f4959,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19719,plain,(
  $false | ($spl1 | $spl18 | $spl62)),
  inference(global_subsumption,[],[f5019,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19718,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19717,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19716,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19715,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19714,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4253,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19713,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19712,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19711,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19710,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19709,plain,(
  $false | ($spl1 | $spl42 | $spl62)),
  inference(global_subsumption,[],[f4388,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19708,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4490,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19707,plain,(
  $false | ($spl1 | $spl24 | $spl62)),
  inference(global_subsumption,[],[f2917,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19706,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4163,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19705,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4350,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19704,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4413,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19703,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4414,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19702,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4420,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19701,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4491,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19700,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4499,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19699,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4502,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19698,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4108,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19697,plain,(
  $false | ($spl1 | $spl48 | $spl62)),
  inference(global_subsumption,[],[f4451,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19696,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4107,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19695,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4226,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19694,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4161,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19693,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4498,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19692,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4308,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19691,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4418,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19690,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4277,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19689,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4276,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19688,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19687,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19686,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19685,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19683,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19682,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19681,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f19680,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19679,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4149,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19678,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4150,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19677,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f4153,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19676,plain,(
  $false | ($spl1 | $spl54 | $spl62)),
  inference(global_subsumption,[],[f4463,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19675,plain,(
  $false | ($spl1 | $spl52 | $spl62)),
  inference(global_subsumption,[],[f4462,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19674,plain,(
  $false | ($spl1 | $spl46 | $spl62)),
  inference(global_subsumption,[],[f4430,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19673,plain,(
  $false | ($spl1 | $spl14 | $spl62)),
  inference(global_subsumption,[],[f1487,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19672,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771,f14772])).
fof(f19671,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859,f14771])).
fof(f19670,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883,f14859])).
fof(f19669,plain,(
  $false | ($spl1 | $spl18 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f14876,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883])).
fof(f14876,plain,(
  ( ! [X44] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X44)))) ) | ($spl18 | $spl24 | $spl30 | $spl62)),
  inference(subsumption_resolution,[],[f14799,f6117])).
fof(f14799,plain,(
  ( ! [X44] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X44))) | s(bool,'T_0') = s(bool,X44)) ) | ($spl18 | $spl62)),
  inference(superposition,[],[f13488,f1650])).
fof(f19668,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886,f14883])).
fof(f19667,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889,f14886])).
fof(f19666,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892,f14889])).
fof(f19665,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848,f14892])).
fof(f19664,plain,(
  $false | ($spl1 | $spl24 | $spl30 | $spl62)),
  inference(global_subsumption,[],[f14894,f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848])).
fof(f14894,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,X0) = s(bool,f(s(bool,X0)))) ) | ($spl24 | $spl30 | $spl62)),
  inference(condensation,[],[f14893])).
fof(f14893,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) ) | ($spl24 | $spl30 | $spl62)),
  inference(subsumption_resolution,[],[f14831,f2970])).
fof(f14831,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) ) | ($spl24 | $spl30 | $spl62)),
  inference(superposition,[],[f13488,f2970])).
fof(f19663,plain,(
  $false | ($spl1 | $spl62)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f138,f139,f141,f146,f145,f183,f184,f193,f190,f194,f221,f237,f242,f241,f240,f300,f340,f429,f506,f507,f624,f626,f766,f867,f1141,f1142,f907,f623,f622,f260,f863,f619,f261,f150,f149,f137,f7514,f7500,f7516,f7515,f7508,f1145,f195,f14974,f908,f828,f621,f402,f142,f1980,f1964,f1843,f1811,f517,f512,f140,f1716,f1712,f1628,f1515,f1415,f1299,f1298,f535,f534,f480,f478,f477,f471,f459,f458,f454,f453,f389,f280,f274,f189,f2104,f650,f662,f1144,f235,f136,f1816,f1612,f1274,f1627,f1841,f1840,f1695,f1269,f1715,f1400,f1979,f1414,f1502,f13488,f14848])).
fof(f19660,plain,(
  $false | ($spl3 | $spl24 | $spl30 | $spl76 | $spl78 | $spl81)),
  inference(subsumption_resolution,[],[f6219,f19624])).
fof(f19624,plain,(
  ( ! [X109] : (s(bool,'T_0') = s(bool,X109)) ) | ($spl3 | $spl24 | $spl30 | $spl76 | $spl78 | $spl81)),
  inference(subsumption_resolution,[],[f15712,f19623])).
fof(f19623,plain,(
  ( ! [X78,X79] : (s(bool,X78) = s(bool,f(s(fun(bool,bool),X79)))) ) | ($spl3 | $spl24 | $spl30 | $spl76 | $spl81)),
  inference(subsumption_resolution,[],[f19622,f240])).
fof(f19622,plain,(
  ( ! [X78,X79] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X78) = s(bool,f(s(fun(bool,bool),X79)))) ) | ($spl3 | $spl24 | $spl30 | $spl76 | $spl81)),
  inference(forward_demodulation,[],[f15668,f19607])).
fof(f19607,plain,(
  ( ! [X66] : (s(bool,'T_0') = s(bool,f(s(bool,X66)))) ) | ($spl3 | $spl24 | $spl30 | $spl81)),
  inference(subsumption_resolution,[],[f19606,f19313])).
fof(f19606,plain,(
  ( ! [X66] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X66)))) ) | ($spl3 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f17928,f6117])).
fof(f17928,plain,(
  ( ! [X66] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X66))) | s(bool,'T_0') = s(bool,X66)) ) | ($spl3 | $spl24 | $spl30)),
  inference(superposition,[],[f6117,f6240])).
fof(f6240,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f6226,f134])).
fof(f6226,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f231,f300])).
fof(f6219,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | $spl3),
  inference(unit_resulting_resolution,[],[f134,f231,f221])).
fof(f19659,plain,(
  $false | ($spl3 | $spl24 | $spl30 | $spl76 | $spl78 | $spl81)),
  inference(subsumption_resolution,[],[f6213,f19624])).
fof(f6213,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | $spl3),
  inference(unit_resulting_resolution,[],[f231,f145])).
fof(f19650,plain,(
  $false | ($spl3 | $spl24 | $spl30 | $spl76 | $spl78 | $spl81)),
  inference(subsumption_resolution,[],[f19649,f19621])).
fof(f19621,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl3 | $spl81)),
  inference(subsumption_resolution,[],[f19620,f19313])).
fof(f19620,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | p(s(bool,X1))) ) | $spl3),
  inference(subsumption_resolution,[],[f17886,f145])).
fof(f17886,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f145,f6240])).
fof(f19649,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl24 | $spl30 | $spl76 | $spl78 | $spl81)),
  inference(subsumption_resolution,[],[f19626,f19624])).
fof(f19626,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl3 | $spl24 | $spl30 | $spl76 | $spl78 | $spl81)),
  inference(backward_demodulation,[],[f19624,f149])).
fof(f19648,plain,(
  $false | ($spl3 | $spl24 | $spl30 | $spl76 | $spl78 | $spl81)),
  inference(subsumption_resolution,[],[f19625,f134])).
fof(f19625,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl24 | $spl30 | $spl76 | $spl78 | $spl81)),
  inference(backward_demodulation,[],[f19624,f137])).
fof(f19616,plain,(
  $false | ($spl3 | $spl18 | $spl24 | $spl30 | $spl81)),
  inference(subsumption_resolution,[],[f19615,f240])).
fof(f19615,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl18 | $spl24 | $spl30 | $spl81)),
  inference(condensation,[],[f19608])).
fof(f19608,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X4)) ) | ($spl3 | $spl18 | $spl24 | $spl30 | $spl81)),
  inference(backward_demodulation,[],[f19607,f1650])).
fof(f19604,plain,(
  $false | ($spl3 | $spl81)),
  inference(subsumption_resolution,[],[f19603,f242])).
fof(f19603,plain,(
  ( ! [X130] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X130),s('type.num.num',X130)))) ) | ($spl3 | $spl81)),
  inference(subsumption_resolution,[],[f17967,f19313])).
fof(f17967,plain,(
  ( ! [X130] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X130),s('type.num.num',X130)))) ) | $spl3),
  inference(superposition,[],[f242,f6240])).
fof(f19454,plain,(
  $false | ($spl3 | $spl81)),
  inference(subsumption_resolution,[],[f19453,f240])).
fof(f19453,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl81)),
  inference(backward_demodulation,[],[f19313,f6235])).
fof(f6235,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | $spl3),
  inference(forward_demodulation,[],[f6222,f261])).
fof(f6222,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f242,f231,f340])).
fof(f19412,plain,(
  $false | ($spl3 | $spl10 | $spl24 | $spl30 | $spl81)),
  inference(subsumption_resolution,[],[f19313,f6223])).
fof(f6223,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | ($spl3 | $spl10 | $spl24 | $spl30)),
  inference(unit_resulting_resolution,[],[f240,f231,f5298])).
fof(f5298,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15)) ) | ($spl10 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f5297,f4660])).
fof(f4660,plain,(
  ( ! [X102] : (s(bool,X102) = s(bool,f(s(bool,X102)))) ) | ($spl10 | $spl24 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f630,f1519,f1544,f1594,f1604,f1605,f1606,f1649,f1651,f1656,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f1414,f1627,f1695,f1716,f477,f458,f1712,f1715,f1841,f1979,f1840,f1816,f1299,f1274,f1269,f1145,f471,f195,f189,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f459,f1515,f534,f1502,f478,f454,f1400,f1415,f1144,f2104,f650,f235,f535,f828,f908,f2917,f261,f150,f149,f137,f4360,f4373,f4350,f4343,f4324,f4318,f4010,f4011,f4309,f4412,f4413,f4414,f4420,f4450,f4107,f4108,f4456,f4276,f4466,f4160,f4161,f4163,f4166,f4167,f4490,f4491,f4252,f4498,f4499,f4225,f4226,f4229,f4243,f4245,f4502,f4253,f4153])).
fof(f1656,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X15)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f1641,f145])).
fof(f1641,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X15))))) | s(bool,'T_0') = s(bool,X15)) ) | $spl10),
  inference(superposition,[],[f1604,f1519])).
fof(f1606,plain,(
  ( ! [X21,X22] : (p(s(bool,X21)) | s(bool,X22) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,X22)) ) | $spl10),
  inference(subsumption_resolution,[],[f1589,f145])).
fof(f1589,plain,(
  ( ! [X21,X22] : (p(s(bool,X21)) | s(bool,X22) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X21)) ) | $spl10),
  inference(superposition,[],[f340,f1519])).
fof(f1605,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(bool,X5)))) ) | $spl10),
  inference(subsumption_resolution,[],[f1581,f145])).
fof(f1581,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl10),
  inference(superposition,[],[f190,f1519])).
fof(f1604,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'F_0') = s(bool,f(s(bool,X4)))) ) | $spl10),
  inference(subsumption_resolution,[],[f1580,f145])).
fof(f1580,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl10),
  inference(superposition,[],[f150,f1519])).
fof(f1594,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f1570])).
fof(f1570,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f1519,f300])).
fof(f1544,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(subsumption_resolution,[],[f1525,f240])).
fof(f1525,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f630,f300])).
fof(f1519,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f630,f141])).
fof(f630,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl10),
  inference(cnf_transformation,[],[f630_D])).
fof(f630_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5297,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,f(s(bool,X15))) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,X15)) ) | ($spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5263,f134])).
fof(f5263,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | s(bool,X16) != s(bool,f(s(bool,X15))) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,X15)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f221,f4717])).
fof(f4717,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f141,f4660])).
fof(f9164,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,X25)) ) | $spl56),
  inference(cnf_transformation,[],[f9164_D])).
fof(f9164_D,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,X25)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f19226,plain,(
  ( ! [X61] : (s(bool,X61) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X61)) ) | $spl92),
  inference(cnf_transformation,[],[f19226_D])).
fof(f19226_D,plain,(
  ( ! [X61] : (s(bool,X61) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X61)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f19152,plain,(
  ( ! [X61,X60] : (s(bool,X61) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X60)) ) | ($spl10 | $spl24 | $spl30 | $spl81)),
  inference(duplicate_literal_removal,[],[f19019])).
fof(f19019,plain,(
  ( ! [X61,X60] : (s(bool,X61) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X60) | s(bool,'F_0') = s(bool,X60)) ) | ($spl10 | $spl24 | $spl30 | $spl81)),
  inference(superposition,[],[f5528,f18882])).
fof(f18882,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,X5)) ) | ($spl10 | $spl24 | $spl30 | $spl81)),
  inference(subsumption_resolution,[],[f18795,f240])).
fof(f18795,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X5) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,X5)) ) | ($spl10 | $spl24 | $spl30 | $spl81)),
  inference(superposition,[],[f18651,f5471])).
fof(f5471,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5431,f5431])).
fof(f5431,plain,(
  ( ! [X43] : (s(bool,X43) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X43)) ) | ($spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5411,f134])).
fof(f5411,plain,(
  ( ! [X43] : (~p(s(bool,'T_0')) | s(bool,X43) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X43)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5296,f2966])).
fof(f5296,plain,(
  ( ! [X14,X13] : (~p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,'F_0') = s(bool,X13)) ) | ($spl10 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f5295,f4660])).
fof(f5295,plain,(
  ( ! [X14,X13] : (~p(s(bool,X14)) | s(bool,X14) = s(bool,f(s(bool,X13))) | s(bool,'F_0') = s(bool,X13)) ) | ($spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5262,f134])).
fof(f5262,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | s(bool,X14) = s(bool,f(s(bool,X13))) | s(bool,'F_0') = s(bool,X13)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f193,f4717])).
fof(f5528,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f4243,f5431])).
fof(f19190,plain,(
  ( ! [X44] : (s(bool,X44) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,f(s(bool,X44)))) ) | $spl90),
  inference(cnf_transformation,[],[f19190_D])).
fof(f19190_D,plain,(
  ( ! [X44] : (s(bool,X44) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,f(s(bool,X44)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f19189,plain,(
  ( ! [X45,X44] : (s(bool,X44) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,f(s(bool,X44))) | s(bool,'T_0') = s(bool,X45)) ) | ($spl10 | $spl24 | $spl30 | $spl81)),
  inference(subsumption_resolution,[],[f18934,f7190])).
fof(f7190,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,f(s(bool,X4)))) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f4243,f5452])).
fof(f5452,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X25)))) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5431,f4660])).
fof(f18934,plain,(
  ( ! [X45,X44] : (s(bool,X44) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,f(s(bool,X44))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | ($spl10 | $spl24 | $spl30 | $spl81)),
  inference(superposition,[],[f18882,f2970])).
fof(f17288,plain,(
  $false | ($spl0 | $spl80)),
  inference(subsumption_resolution,[],[f17138,f242])).
fof(f17138,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl80)),
  inference(backward_demodulation,[],[f17081,f11297])).
fof(f11297,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f11006,f146])).
fof(f11006,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f10987,f189])).
fof(f10987,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) | $spl0),
  inference(subsumption_resolution,[],[f10919,f134])).
fof(f10919,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) | $spl0),
  inference(superposition,[],[f5938,f274])).
fof(f5938,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f5916,f134])).
fof(f5916,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1))))) ) | $spl0),
  inference(superposition,[],[f195,f4903])).
fof(f4903,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl0),
  inference(unit_resulting_resolution,[],[f4666,f146])).
fof(f4666,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) | $spl0),
  inference(forward_demodulation,[],[f228,f136])).
fof(f17287,plain,(
  $false | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(subsumption_resolution,[],[f17137,f17257])).
fof(f17257,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(subsumption_resolution,[],[f17256,f137])).
fof(f17256,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(forward_demodulation,[],[f17255,f261])).
fof(f17255,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(forward_demodulation,[],[f17189,f17178])).
fof(f17178,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(subsumption_resolution,[],[f17162,f17176])).
fof(f17176,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(subsumption_resolution,[],[f17175,f137])).
fof(f17175,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(forward_demodulation,[],[f17160,f261])).
fof(f17160,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(backward_demodulation,[],[f17159,f17129])).
fof(f17129,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X0))))) ) | ($spl0 | $spl80)),
  inference(backward_demodulation,[],[f17081,f10421])).
fof(f10421,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X0))))) ) | $spl0),
  inference(subsumption_resolution,[],[f10355,f134])).
fof(f10355,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X0))))) ) | $spl0),
  inference(superposition,[],[f195,f6525])).
fof(f6525,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | $spl0),
  inference(subsumption_resolution,[],[f6503,f134])).
fof(f6503,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | $spl0),
  inference(superposition,[],[f471,f4903])).
fof(f17159,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(subsumption_resolution,[],[f17135,f11978])).
fof(f17135,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X8)) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(backward_demodulation,[],[f17081,f11046])).
fof(f11046,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,X8)) ) | ($spl0 | $spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f11028,f137])).
fof(f11028,plain,(
  ( ! [X8] : (p(s(bool,'F_0')) | s(bool,X8) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,X8)) ) | ($spl0 | $spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f10987,f5471])).
fof(f17162,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(backward_demodulation,[],[f17159,f1414])).
fof(f17189,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(backward_demodulation,[],[f17178,f189])).
fof(f17137,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) | ($spl0 | $spl80)),
  inference(backward_demodulation,[],[f17081,f11288])).
fof(f11288,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f11006,f189])).
fof(f17278,plain,(
  ( ! [X118,X116,X117] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(fun(X116,bool),X117) = s(fun(X116,bool),X118)) ) | ($spl0 | $spl3 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(forward_demodulation,[],[f17277,f261])).
fof(f17277,plain,(
  ( ! [X118,X116,X117] : (s(bool,'T_0') = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) | s(fun(X116,bool),X117) = s(fun(X116,bool),X118)) ) | ($spl0 | $spl3 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(forward_demodulation,[],[f17276,f17260])).
fof(f17260,plain,(
  ( ! [X50] : (s(bool,X50) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(subsumption_resolution,[],[f17211,f11978])).
fof(f17211,plain,(
  ( ! [X50] : (s(bool,X50) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X50)) ) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(backward_demodulation,[],[f17178,f5436])).
fof(f5436,plain,(
  ( ! [X50,X49] : (s(bool,X50) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))))) | s(bool,'F_0') = s(bool,X50)) ) | ($spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5416,f134])).
fof(f5416,plain,(
  ( ! [X50,X49] : (~p(s(bool,'T_0')) | s(bool,X50) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))))) | s(bool,'F_0') = s(bool,X50)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5296,f280])).
fof(f17276,plain,(
  ( ! [X118,X116,X117] : (s(fun(X116,bool),X117) = s(fun(X116,bool),X118) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X116,sK2(X117,X118,bool,X116))))))) ) | ($spl0 | $spl3 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(subsumption_resolution,[],[f17268,f17260])).
fof(f17268,plain,(
  ( ! [X118,X116,X117] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(fun(X116,bool),X117) = s(fun(X116,bool),X118) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X116,sK2(X117,X118,bool,X116))))))) ) | ($spl0 | $spl3 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(backward_demodulation,[],[f17260,f11834])).
fof(f11834,plain,(
  ( ! [X118,X116,X117] : (s(bool,'F_0') != s(bool,f(s(X116,sK2(X117,X118,bool,X116)))) | s(fun(X116,bool),X117) = s(fun(X116,bool),X118) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X116,sK2(X117,X118,bool,X116))))))) ) | ($spl3 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f11578,f6235])).
fof(f11578,plain,(
  ( ! [X118,X116,X117] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) != s(bool,f(s(X116,sK2(X117,X118,bool,X116)))) | s(fun(X116,bool),X117) = s(fun(X116,bool),X118) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X116,sK2(X117,X118,bool,X116))))))) ) | ($spl3 | $spl24 | $spl30)),
  inference(superposition,[],[f194,f6696])).
fof(f6696,plain,(
  ( ! [X23] : (s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X23)))) ) | ($spl3 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f6564])).
fof(f6564,plain,(
  ( ! [X23] : (s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X23))) | s(bool,'T_0') = s(bool,f(s(bool,X23)))) ) | ($spl3 | $spl24 | $spl30)),
  inference(superposition,[],[f6240,f4245])).
fof(f17258,plain,(
  $false | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(subsumption_resolution,[],[f17203,f242])).
fof(f17203,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl10 | $spl24 | $spl30 | $spl80)),
  inference(backward_demodulation,[],[f17178,f1298])).
fof(f17158,plain,(
  $false | ($spl0 | $spl80)),
  inference(subsumption_resolution,[],[f17134,f242])).
fof(f17134,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl80)),
  inference(backward_demodulation,[],[f17081,f11015])).
fof(f11015,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | $spl0),
  inference(unit_resulting_resolution,[],[f10987,f146])).
fof(f17157,plain,(
  $false | ($spl0 | $spl80)),
  inference(subsumption_resolution,[],[f17156,f183])).
fof(f17156,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl80)),
  inference(forward_demodulation,[],[f17133,f140])).
fof(f17133,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl80)),
  inference(backward_demodulation,[],[f17081,f11005])).
fof(f11005,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | $spl0),
  inference(unit_resulting_resolution,[],[f10987,f235])).
fof(f17081,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | $spl80),
  inference(cnf_transformation,[],[f17081_D])).
fof(f17081_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f17083,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(sK0)))))) | $spl82),
  inference(cnf_transformation,[],[f17083_D])).
fof(f17083_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(sK0)))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f17085,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(sK0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | $spl0),
  inference(subsumption_resolution,[],[f17039,f134])).
fof(f17039,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(sK0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | $spl0),
  inference(superposition,[],[f650,f4903])).
fof(f17084,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl84),
  inference(cnf_transformation,[],[f17084_D])).
fof(f17084_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f17079,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(sK0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0)) ) | ($spl0 | $spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f17038,f149])).
fof(f17038,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(sK0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | s(bool,'F_0') = s(bool,X0)) ) | ($spl0 | $spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f650,f5434])).
fof(f5434,plain,(
  ( ! [X46] : (s(bool,X46) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,X46)) ) | ($spl0 | $spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5414,f134])).
fof(f5414,plain,(
  ( ! [X46] : (~p(s(bool,'T_0')) | s(bool,X46) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,X46)) ) | ($spl0 | $spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5296,f4903])).
fof(f15712,plain,(
  ( ! [X109,X108] : (s(bool,X109) != s(bool,f(s(fun(bool,bool),X108))) | s(bool,'T_0') = s(bool,X109)) ) | $spl78),
  inference(cnf_transformation,[],[f15712_D])).
fof(f15712_D,plain,(
  ( ! [X109,X108] : (s(bool,X109) != s(bool,f(s(fun(bool,bool),X108))) | s(bool,'T_0') = s(bool,X109)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f15623,plain,(
  ( ! [X109,X107,X108] : (s(bool,X109) != s(bool,f(s(fun(bool,bool),X108))) | s(bool,'T_0') = s(bool,X109) | s(bool,'F_0') = s(bool,X107)) ) | ($spl10 | $spl24 | $spl30 | $spl74)),
  inference(duplicate_literal_removal,[],[f15472])).
fof(f15472,plain,(
  ( ! [X109,X107,X108] : (s(bool,X109) != s(bool,f(s(fun(bool,bool),X108))) | s(bool,'T_0') = s(bool,X109) | s(bool,'F_0') = s(bool,X107) | s(bool,'F_0') = s(bool,X107)) ) | ($spl10 | $spl24 | $spl30 | $spl74)),
  inference(superposition,[],[f5528,f15072])).
fof(f15072,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,f(s(fun(bool,bool),X25))) | s(bool,'F_0') = s(bool,X26)) ) | ($spl10 | $spl24 | $spl30 | $spl74)),
  inference(trivial_inequality_removal,[],[f15063])).
fof(f15063,plain,(
  ( ! [X26,X25] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X26) = s(bool,f(s(fun(bool,bool),X25))) | s(bool,'F_0') = s(bool,X26)) ) | ($spl10 | $spl24 | $spl30 | $spl74)),
  inference(superposition,[],[f14981,f5471])).
fof(f15668,plain,(
  ( ! [X78,X79] : (s(bool,X78) = s(bool,f(s(fun(bool,bool),X79))) | s(bool,'F_0') = s(bool,f(s(bool,X78)))) ) | $spl76),
  inference(cnf_transformation,[],[f15668_D])).
fof(f15668_D,plain,(
  ( ! [X78,X79] : (s(bool,X78) = s(bool,f(s(fun(bool,bool),X79))) | s(bool,'F_0') = s(bool,f(s(bool,X78)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f15667,plain,(
  ( ! [X80,X78,X79] : (s(bool,X78) = s(bool,f(s(fun(bool,bool),X79))) | s(bool,'F_0') = s(bool,f(s(bool,X78))) | s(bool,'T_0') = s(bool,X80)) ) | ($spl10 | $spl24 | $spl30 | $spl74)),
  inference(subsumption_resolution,[],[f15374,f7190])).
fof(f15374,plain,(
  ( ! [X80,X78,X79] : (s(bool,X78) = s(bool,f(s(fun(bool,bool),X79))) | s(bool,'F_0') = s(bool,f(s(bool,X78))) | s(bool,X78) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) ) | ($spl10 | $spl24 | $spl30 | $spl74)),
  inference(superposition,[],[f15072,f2970])).
fof(f14980,plain,(
  ( ! [X0,X1] : (s(fun(fun(bool,bool),bool),X0) = s(fun(fun(bool,bool),bool),X1) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK2(X0,X1,bool,fun(bool,bool)))))) ) | $spl72),
  inference(cnf_transformation,[],[f14980_D])).
fof(f14980_D,plain,(
  ( ! [X0,X1] : (s(fun(fun(bool,bool),bool),X0) = s(fun(fun(bool,bool),bool),X1) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK2(X0,X1,bool,fun(bool,bool)))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f14981,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X2)))) ) | $spl74),
  inference(cnf_transformation,[],[f14981_D])).
fof(f14981_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X2)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f14906,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X5))) | s(fun(fun(bool,bool),bool),X3) = s(fun(fun(bool,bool),bool),X4) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK2(X3,X4,bool,fun(bool,bool)))))) ) | $spl22),
  inference(superposition,[],[f621,f2807])).
fof(f14905,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X2))) | s(fun(fun(bool,bool),bool),X0) = s(fun(fun(bool,bool),bool),X1) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK2(X0,X1,bool,fun(bool,bool)))))) ) | $spl22),
  inference(superposition,[],[f621,f2807])).
fof(f14361,plain,(
  ( ! [X105,X106] : (s(bool,X106) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X105))))) | s(bool,'T_0') = s(bool,X106)) ) | $spl70),
  inference(cnf_transformation,[],[f14361_D])).
fof(f14361_D,plain,(
  ( ! [X105,X106] : (s(bool,X106) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X105))))) | s(bool,'T_0') = s(bool,X106)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f14276,plain,(
  ( ! [X105,X106,X104] : (s(bool,X106) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X105))))) | s(bool,'T_0') = s(bool,X106) | s(bool,'F_0') = s(bool,X104)) ) | ($spl10 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f14143])).
fof(f14143,plain,(
  ( ! [X105,X106,X104] : (s(bool,X106) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X105))))) | s(bool,'T_0') = s(bool,X106) | s(bool,'F_0') = s(bool,X104) | s(bool,'F_0') = s(bool,X104)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5528,f5436])).
fof(f14318,plain,(
  ( ! [X76,X75] : (s(bool,X75) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))))) | s(bool,'F_0') = s(bool,f(s(bool,X75)))) ) | $spl68),
  inference(cnf_transformation,[],[f14318_D])).
fof(f14318_D,plain,(
  ( ! [X76,X75] : (s(bool,X75) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))))) | s(bool,'F_0') = s(bool,f(s(bool,X75)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f14317,plain,(
  ( ! [X76,X77,X75] : (s(bool,X75) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))))) | s(bool,'F_0') = s(bool,f(s(bool,X75))) | s(bool,'T_0') = s(bool,X77)) ) | ($spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f14044,f7190])).
fof(f14044,plain,(
  ( ! [X76,X77,X75] : (s(bool,X75) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))))) | s(bool,'F_0') = s(bool,f(s(bool,X75))) | s(bool,X75) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5436,f2970])).
fof(f13889,plain,(
  ( ! [X103,X102] : (s(bool,X103) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X102),s('type.num.num','const.num.SUC_1'(s('type.num.num',X102))))) | s(bool,'T_0') = s(bool,X103)) ) | $spl66),
  inference(cnf_transformation,[],[f13889_D])).
fof(f13889_D,plain,(
  ( ! [X103,X102] : (s(bool,X103) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X102),s('type.num.num','const.num.SUC_1'(s('type.num.num',X102))))) | s(bool,'T_0') = s(bool,X103)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f13808,plain,(
  ( ! [X103,X101,X102] : (s(bool,X103) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X102),s('type.num.num','const.num.SUC_1'(s('type.num.num',X102))))) | s(bool,'T_0') = s(bool,X103) | s(bool,'F_0') = s(bool,X101)) ) | ($spl10 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f13675])).
fof(f13675,plain,(
  ( ! [X103,X101,X102] : (s(bool,X103) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X102),s('type.num.num','const.num.SUC_1'(s('type.num.num',X102))))) | s(bool,'T_0') = s(bool,X103) | s(bool,'F_0') = s(bool,X101) | s(bool,'F_0') = s(bool,X101)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5528,f5435])).
fof(f5435,plain,(
  ( ! [X47,X48] : (s(bool,X48) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X47),s('type.num.num','const.num.SUC_1'(s('type.num.num',X47))))) | s(bool,'F_0') = s(bool,X48)) ) | ($spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5415,f134])).
fof(f5415,plain,(
  ( ! [X47,X48] : (~p(s(bool,'T_0')) | s(bool,X48) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X47),s('type.num.num','const.num.SUC_1'(s('type.num.num',X47))))) | s(bool,'F_0') = s(bool,X48)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5296,f274])).
fof(f13847,plain,(
  ( ! [X72,X73] : (s(bool,X72) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X73),s('type.num.num','const.num.SUC_1'(s('type.num.num',X73))))) | s(bool,'F_0') = s(bool,f(s(bool,X72)))) ) | $spl64),
  inference(cnf_transformation,[],[f13847_D])).
fof(f13847_D,plain,(
  ( ! [X72,X73] : (s(bool,X72) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X73),s('type.num.num','const.num.SUC_1'(s('type.num.num',X73))))) | s(bool,'F_0') = s(bool,f(s(bool,X72)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f13846,plain,(
  ( ! [X74,X72,X73] : (s(bool,X72) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X73),s('type.num.num','const.num.SUC_1'(s('type.num.num',X73))))) | s(bool,'F_0') = s(bool,f(s(bool,X72))) | s(bool,'T_0') = s(bool,X74)) ) | ($spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f13580,f7190])).
fof(f13580,plain,(
  ( ! [X74,X72,X73] : (s(bool,X72) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X73),s('type.num.num','const.num.SUC_1'(s('type.num.num',X73))))) | s(bool,'F_0') = s(bool,f(s(bool,X72))) | s(bool,X72) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5435,f2970])).
fof(f13488,plain,(
  ( ! [X61] : (s(bool,X61) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X61)) ) | $spl62),
  inference(cnf_transformation,[],[f13488_D])).
fof(f13488_D,plain,(
  ( ! [X61] : (s(bool,X61) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X61)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f13415,plain,(
  ( ! [X61,X60] : (s(bool,X61) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X60)) ) | ($spl0 | $spl10 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f13288])).
fof(f13288,plain,(
  ( ! [X61,X60] : (s(bool,X61) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X60) | s(bool,'F_0') = s(bool,X60)) ) | ($spl0 | $spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5528,f5434])).
fof(f13449,plain,(
  ( ! [X41] : (s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,f(s(bool,X41)))) ) | $spl60),
  inference(cnf_transformation,[],[f13449_D])).
fof(f13449_D,plain,(
  ( ! [X41] : (s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,f(s(bool,X41)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f13448,plain,(
  ( ! [X41,X42] : (s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,f(s(bool,X41))) | s(bool,'T_0') = s(bool,X42)) ) | ($spl0 | $spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f13207,f7190])).
fof(f13207,plain,(
  ( ! [X41,X42] : (s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,f(s(bool,X41))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) ) | ($spl0 | $spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5434,f2970])).
fof(f9288,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9271,f242])).
fof(f9271,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl56),
  inference(backward_demodulation,[],[f9263,f1298])).
fof(f9263,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) ) | $spl56),
  inference(subsumption_resolution,[],[f9231,f137])).
fof(f9231,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) ) | $spl56),
  inference(backward_demodulation,[],[f9164,f1712])).
fof(f9165,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X26)) ) | $spl58),
  inference(cnf_transformation,[],[f9165_D])).
fof(f9165_D,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X26)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f9168,plain,(
  ( ! [X41,X42] : (s(bool,X42) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X41)) ) | ($spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f9060,f5294])).
fof(f5294,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,f(s(bool,X8))) | s(bool,'F_0') = s(bool,X8)) ) | ($spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5259,f134])).
fof(f5259,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,f(s(bool,X8))) | s(bool,'F_0') = s(bool,X8)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f149,f4717])).
fof(f9060,plain,(
  ( ! [X41,X42] : (s(bool,X42) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X41) | s(bool,'F_0') = s(bool,f(s(bool,X41)))) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5528,f5679])).
fof(f5679,plain,(
  ( ! [X26] : (s(bool,X26) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X26)))) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5432,f4660])).
fof(f5432,plain,(
  ( ! [X44] : (s(bool,X44) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X44)) ) | ($spl10 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5412,f134])).
fof(f5412,plain,(
  ( ! [X44] : (~p(s(bool,'T_0')) | s(bool,X44) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X44)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5296,f2849])).
fof(f9148,plain,(
  ( ! [X26,X25] : (s(bool,X26) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) ) | ($spl10 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f9054])).
fof(f9054,plain,(
  ( ! [X26,X25] : (s(bool,X26) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25) | s(bool,'F_0') = s(bool,X25)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f5528,f5432])).
fof(f6209,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f6208,f4665])).
fof(f4665,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl5),
  inference(forward_demodulation,[],[f233,f136])).
fof(f6208,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl0),
  inference(subsumption_resolution,[],[f6189,f134])).
fof(f6189,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f235,f4903])).
fof(f6203,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f6202,f134])).
fof(f6202,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f6183,f4903])).
fof(f6183,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) | $spl5),
  inference(unit_resulting_resolution,[],[f4665,f235])).
fof(f4869,plain,(
  $false | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f630,f1519,f1544,f1594,f1604,f1605,f1606,f1649,f1651,f1656,f1487,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2917,f4360,f4373,f4350,f4343,f4324,f4318,f4010,f4011,f4309,f4413,f4414,f4420,f4450,f4107,f4108,f4456,f4276,f4466,f4160,f4161,f4163,f4166,f4167,f4490,f4491,f4252,f4498,f4499,f4225,f4226,f4229,f4243,f4245,f4502,f4253,f4153,f4660,f863,f621,f619,f1414,f1627,f1716,f477,f458,f1695,f1715,f1712,f1841,f454,f1400,f1415,f478,f1502,f459,f1515,f1628,f1612,f453,f1980,f1979,f1964,f1843,f1840,f1816,f1811,f1299,f1274,f1269,f1145,f534,f517,f512,f471,f195,f189,f1144,f2104,f650,f235,f535,f828,f908,f4412,f261,f150,f149,f137,f4709,f4747,f4819,f4821,f4823,f4868])).
fof(f4868,plain,(
  ( ! [X2] : (p(s(bool,X2))) ) | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(forward_demodulation,[],[f4867,f4660])).
fof(f4867,plain,(
  ( ! [X2] : (p(s(bool,f(s(bool,X2))))) ) | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4718,f4747])).
fof(f4718,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,f(s(bool,X2))))) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f145,f4660])).
fof(f4823,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4764,f134])).
fof(f4764,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(backward_demodulation,[],[f4747,f650])).
fof(f4821,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4763,f134])).
fof(f4763,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(backward_demodulation,[],[f4747,f235])).
fof(f4819,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4749,f4747])).
fof(f4749,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(backward_demodulation,[],[f4747,f149])).
fof(f4747,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(condensation,[],[f4746])).
fof(f4746,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X4)) ) | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4710,f1487])).
fof(f4710,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f4660,f300])).
fof(f4709,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,f(s(bool,X2)))) ) | ($spl10 | $spl24 | $spl30)),
  inference(superposition,[],[f4660,f141])).
fof(f4824,plain,(
  $false | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f630,f1519,f1544,f1594,f1604,f1605,f1606,f1649,f1651,f1656,f1487,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2917,f4360,f4373,f4350,f4343,f4324,f4318,f4010,f4011,f4309,f4413,f4414,f4420,f4450,f4107,f4108,f4456,f4276,f4466,f4160,f4161,f4163,f4166,f4167,f4490,f4491,f4252,f4498,f4499,f4225,f4226,f4229,f4243,f4245,f4502,f4253,f4153,f4660,f863,f621,f619,f1414,f1627,f1716,f477,f458,f1695,f1715,f1712,f1841,f454,f1400,f1415,f478,f1502,f459,f1515,f1628,f1612,f453,f1980,f1979,f1964,f1843,f1840,f1816,f1811,f1299,f1274,f1269,f1145,f534,f517,f512,f471,f195,f189,f1144,f2104,f650,f235,f535,f828,f908,f4412,f261,f150,f149,f137,f4709,f4747,f4819,f4821,f4823])).
fof(f4822,plain,(
  $false | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f630,f1519,f1544,f1594,f1604,f1605,f1606,f1649,f1651,f1656,f1487,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2917,f4360,f4373,f4350,f4343,f4324,f4318,f4010,f4011,f4309,f4413,f4414,f4420,f4450,f4107,f4108,f4456,f4276,f4466,f4160,f4161,f4163,f4166,f4167,f4490,f4491,f4252,f4498,f4499,f4225,f4226,f4229,f4243,f4245,f4502,f4253,f4153,f4660,f863,f621,f619,f1414,f1627,f1716,f477,f458,f1695,f1715,f1712,f1841,f454,f1400,f1415,f478,f1502,f459,f1515,f1628,f1612,f453,f1980,f1979,f1964,f1843,f1840,f1816,f1811,f1299,f1274,f1269,f1145,f534,f517,f512,f471,f195,f189,f1144,f2104,f650,f235,f535,f828,f908,f4412,f261,f150,f149,f137,f4709,f4747,f4819,f4821])).
fof(f4820,plain,(
  $false | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f630,f1519,f1544,f1594,f1604,f1605,f1606,f1649,f1651,f1656,f1487,f2808,f2923,f2924,f2925,f2926,f2927,f2928,f2937,f2938,f2939,f2940,f2849,f2966,f2970,f3068,f2917,f4360,f4373,f4350,f4343,f4324,f4318,f4010,f4011,f4309,f4413,f4414,f4420,f4450,f4107,f4108,f4456,f4276,f4466,f4160,f4161,f4163,f4166,f4167,f4490,f4491,f4252,f4498,f4499,f4225,f4226,f4229,f4243,f4245,f4502,f4253,f4153,f4660,f863,f621,f619,f1414,f1627,f1716,f477,f458,f1695,f1715,f1712,f1841,f454,f1400,f1415,f478,f1502,f459,f1515,f1628,f1612,f453,f1980,f1979,f1964,f1843,f1840,f1816,f1811,f1299,f1274,f1269,f1145,f534,f517,f512,f471,f195,f189,f1144,f2104,f650,f235,f535,f828,f908,f4412,f261,f150,f149,f137,f4709,f4747,f4819])).
fof(f4818,plain,(
  $false | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4748,f134])).
fof(f4748,plain,(
  ~p(s(bool,'T_0')) | ($spl10 | $spl14 | $spl24 | $spl30)),
  inference(backward_demodulation,[],[f4747,f137])).
fof(f4609,plain,(
  $false | $spl38),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f1414,f1627,f1695,f1716,f477,f458,f1712,f1715,f1841,f1979,f1840,f1816,f1299,f1274,f1269,f1145,f471,f195,f189,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f459,f1515,f534,f1502,f478,f454,f1400,f1415,f1144,f2104,f650,f235,f535,f828,f908,f261,f150,f149,f137,f4368,f4600,f4606,f4608])).
fof(f4608,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | $spl38),
  inference(subsumption_resolution,[],[f4539,f134])).
fof(f4539,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | $spl38),
  inference(backward_demodulation,[],[f4368,f650])).
fof(f4606,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | $spl38),
  inference(subsumption_resolution,[],[f4538,f134])).
fof(f4538,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | $spl38),
  inference(backward_demodulation,[],[f4368,f235])).
fof(f4600,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl38),
  inference(subsumption_resolution,[],[f4507,f4368])).
fof(f4507,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl38),
  inference(backward_demodulation,[],[f4368,f149])).
fof(f4607,plain,(
  $false | $spl38),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f1414,f1627,f1695,f1716,f477,f458,f1712,f1715,f1841,f1979,f1840,f1816,f1299,f1274,f1269,f1145,f471,f195,f189,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f459,f1515,f534,f1502,f478,f454,f1400,f1415,f1144,f2104,f650,f235,f535,f828,f908,f261,f150,f149,f137,f4368,f4600,f4606])).
fof(f4605,plain,(
  $false | ($spl24 | $spl30 | $spl38)),
  inference(global_subsumption,[],[f4592,f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f1414,f1627,f1695,f1716,f477,f458,f1712,f1715,f1841,f1979,f1840,f1816,f1299,f1274,f1269,f1145,f471,f195,f189,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f459,f1515,f534,f1502,f478,f454,f1400,f1415,f1144,f2104,f650,f235,f535,f828,f908,f261,f150,f149,f137,f4368,f4600])).
fof(f4592,plain,(
  ( ! [X103,X105,X102,X104] : (s(fun(X102,bool),X103) = s(fun(X102,bool),X104) | s(bool,X105) = s(bool,f(s(bool,X105))) | s(bool,X105) = s(bool,g(s(X102,sK2(X103,X104,bool,X102))))) ) | ($spl24 | $spl30 | $spl38)),
  inference(trivial_inequality_removal,[],[f4530])).
fof(f4530,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X102,bool),X103) = s(fun(X102,bool),X104) | s(bool,X105) = s(bool,f(s(bool,X105))) | s(bool,X105) = s(bool,g(s(X102,sK2(X103,X104,bool,X102))))) ) | ($spl24 | $spl30 | $spl38)),
  inference(backward_demodulation,[],[f4368,f4226])).
fof(f4604,plain,(
  $false | ($spl24 | $spl30 | $spl38)),
  inference(global_subsumption,[],[f4603,f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f1414,f1627,f1695,f1716,f477,f458,f1712,f1715,f1841,f1979,f1840,f1816,f1299,f1274,f1269,f1145,f471,f195,f189,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f459,f1515,f534,f1502,f478,f454,f1400,f1415,f1144,f2104,f650,f235,f535,f828,f908,f261,f150,f149,f137,f4368,f4600])).
fof(f4603,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0)))) ) | ($spl24 | $spl30 | $spl38)),
  inference(condensation,[],[f4602])).
fof(f4602,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,X6) = s(bool,X7)) ) | ($spl24 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f4515,f4368])).
fof(f4515,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,X6) = s(bool,X7)) ) | ($spl24 | $spl30 | $spl38)),
  inference(backward_demodulation,[],[f4368,f4490])).
fof(f4601,plain,(
  $false | $spl38),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f1414,f1627,f1695,f1716,f477,f458,f1712,f1715,f1841,f1979,f1840,f1816,f1299,f1274,f1269,f1145,f471,f195,f189,f1980,f1964,f1843,f1811,f1628,f1612,f517,f512,f453,f459,f1515,f534,f1502,f478,f454,f1400,f1415,f1144,f2104,f650,f235,f535,f828,f908,f261,f150,f149,f137,f4368,f4600])).
fof(f4599,plain,(
  $false | $spl38),
  inference(subsumption_resolution,[],[f4506,f134])).
fof(f4506,plain,(
  ~p(s(bool,'T_0')) | $spl38),
  inference(backward_demodulation,[],[f4368,f137])).
fof(f4463,plain,(
  ( ! [X46] : (s(bool,'F_0') != s(bool,f(s(bool,X46))) | s(bool,X46) = s(bool,f(s(bool,X46)))) ) | $spl54),
  inference(cnf_transformation,[],[f4463_D])).
fof(f4463_D,plain,(
  ( ! [X46] : (s(bool,'F_0') != s(bool,f(s(bool,X46))) | s(bool,X46) = s(bool,f(s(bool,X46)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4266,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') != s(bool,f(s(bool,X46))) | s(bool,'T_0') = s(bool,X45) | s(bool,X46) = s(bool,f(s(bool,X46)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f4127])).
fof(f4127,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') != s(bool,f(s(bool,X46))) | s(bool,'T_0') = s(bool,X45) | s(bool,X46) = s(bool,f(s(bool,X46))) | s(bool,'T_0') = s(bool,X45)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f1906,f2970])).
fof(f4462,plain,(
  ( ! [X44] : (p(s(bool,f(s(bool,X44)))) | s(bool,X44) = s(bool,f(s(bool,X44)))) ) | $spl52),
  inference(cnf_transformation,[],[f4462_D])).
fof(f4462_D,plain,(
  ( ! [X44] : (p(s(bool,f(s(bool,X44)))) | s(bool,X44) = s(bool,f(s(bool,X44)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4267,plain,(
  ( ! [X43,X44] : (p(s(bool,f(s(bool,X44)))) | s(bool,'T_0') = s(bool,X43) | s(bool,X44) = s(bool,f(s(bool,X44)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f4126])).
fof(f4126,plain,(
  ( ! [X43,X44] : (p(s(bool,f(s(bool,X44)))) | s(bool,'T_0') = s(bool,X43) | s(bool,X44) = s(bool,f(s(bool,X44))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f1889,f2970])).
fof(f4457,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,X3) = s(bool,f(s(bool,X3)))) ) | $spl50),
  inference(cnf_transformation,[],[f4457_D])).
fof(f4457_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,X3) = s(bool,f(s(bool,X3)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4456,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X2)) | s(bool,X3) = s(bool,f(s(bool,X3)))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4110,f145])).
fof(f4110,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X2)) | s(bool,X3) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f145,f2970])).
fof(f4451,plain,(
  ( ! [X101,X100] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X100),s('type.num.num',X101)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X100))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X100))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))))) ) | $spl48),
  inference(cnf_transformation,[],[f4451_D])).
fof(f4451_D,plain,(
  ( ! [X101,X100] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X100),s('type.num.num',X101)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X100))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X100))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f4450,plain,(
  ( ! [X101,X102,X100] : (p(s(bool,X102)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X100),s('type.num.num',X101)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X100))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X100))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4100,f145])).
fof(f4100,plain,(
  ( ! [X101,X102,X100] : (p(s(bool,X102)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X100),s('type.num.num',X101)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X100))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101))))) = s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X100))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101))))))) | s(bool,'T_0') = s(bool,X102)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f189,f2970])).
fof(f4430,plain,(
  ( ! [X52,X50] : (s(bool,X50) != s(bool,f(s(bool,X52))) | s(bool,X50) = s(bool,f(s(bool,X50))) | s(bool,'T_0') = s(bool,X52)) ) | $spl46),
  inference(cnf_transformation,[],[f4430_D])).
fof(f4430_D,plain,(
  ( ! [X52,X50] : (s(bool,X50) != s(bool,f(s(bool,X52))) | s(bool,X50) = s(bool,f(s(bool,X50))) | s(bool,'T_0') = s(bool,X52)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4429,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X51)) | s(bool,X50) != s(bool,f(s(bool,X52))) | s(bool,'T_0') = s(bool,X52) | s(bool,X50) = s(bool,f(s(bool,X50)))) ) | ($spl11 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4076,f145])).
fof(f4076,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X51)) | s(bool,X50) != s(bool,f(s(bool,X52))) | s(bool,'T_0') = s(bool,X52) | s(bool,X50) = s(bool,f(s(bool,X50))) | s(bool,'T_0') = s(bool,X51)) ) | ($spl11 | $spl24 | $spl30)),
  inference(superposition,[],[f1908,f2970])).
fof(f4369,plain,(
  ( ! [X27,X25] : (s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,X25) = s(bool,X27) | s(bool,X27) = s(bool,f(s(bool,X27)))) ) | $spl40),
  inference(cnf_transformation,[],[f4369_D])).
fof(f4369_D,plain,(
  ( ! [X27,X25] : (s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,X25) = s(bool,X27) | s(bool,X27) = s(bool,f(s(bool,X27)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4309,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X12) = s(bool,X14) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f4024])).
fof(f4024,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X12) = s(bool,X14) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X13)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f2970])).
fof(f4318,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) = s(bool,X21) | s(bool,X21) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,f(s(bool,X20)))) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f3992])).
fof(f3992,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) = s(bool,X21) | s(bool,X21) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f2970])).
fof(f4389,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | $spl44),
  inference(cnf_transformation,[],[f4389_D])).
fof(f4389_D,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4319,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,X15) = s(bool,f(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f3990])).
fof(f3990,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,X15) = s(bool,f(s(bool,X15))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl24 | $spl30)),
  inference(superposition,[],[f2970,f506])).
fof(f4324,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f3985])).
fof(f3985,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f141])).
fof(f4343,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,f(s(bool,X22))) | s(bool,X24) = s(bool,f(s(bool,X24))) | s(bool,X22) = s(bool,X24)) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f3968])).
fof(f3968,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,f(s(bool,X24))) | s(bool,X22) = s(bool,X24)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f2970])).
fof(f4360,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,f(s(bool,X27))) | s(bool,X25) = s(bool,X27)) ) | ($spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f3933])).
fof(f3933,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,f(s(bool,X27))) | s(bool,X25) = s(bool,X27)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f2970,f2970])).
fof(f3448,plain,(
  $false | ($spl26 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f1414,f454,f1400,f1415,f1502,f459,f1515,f1980,f1979,f1964,f1843,f1811,f1299,f1145,f534,f517,f512,f195,f1144,f2104,f535,f828,f908,f261,f150,f149,f137,f1816,f1274,f1612,f1627,f1840,f1841,f1269,f1695,f1716,f1628,f478,f477,f471,f458,f453,f189,f1715,f1712,f650,f235,f2849,f2845,f3068,f3169,f3272,f3386,f3388,f3390,f3435,f3447])).
fof(f3447,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | ($spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3248,f3272])).
fof(f3248,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | $spl26),
  inference(equality_factoring,[],[f2845])).
fof(f3435,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3434,f3272])).
fof(f3434,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | p(s(bool,X1))) ) | $spl26),
  inference(subsumption_resolution,[],[f3184,f145])).
fof(f3184,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl26),
  inference(superposition,[],[f145,f2845])).
fof(f3390,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3326,f134])).
fof(f3326,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl26 | $spl30)),
  inference(backward_demodulation,[],[f3272,f650])).
fof(f3388,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | ($spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3325,f134])).
fof(f3325,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) ) | ($spl26 | $spl30)),
  inference(backward_demodulation,[],[f3272,f235])).
fof(f3386,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3305,f3272])).
fof(f3305,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl26 | $spl30)),
  inference(backward_demodulation,[],[f3272,f149])).
fof(f3272,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl26 | $spl30)),
  inference(duplicate_literal_removal,[],[f3176])).
fof(f3176,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl26 | $spl30)),
  inference(superposition,[],[f2845,f2849])).
fof(f3169,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl26),
  inference(superposition,[],[f2845,f300])).
fof(f3436,plain,(
  $false | ($spl26 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f1414,f454,f1400,f1415,f1502,f459,f1515,f1980,f1979,f1964,f1843,f1811,f1299,f1145,f534,f517,f512,f195,f1144,f2104,f535,f828,f908,f261,f150,f149,f137,f1816,f1274,f1612,f1627,f1840,f1841,f1269,f1695,f1716,f1628,f478,f477,f471,f458,f453,f189,f1715,f1712,f650,f235,f2849,f2845,f3068,f3169,f3272,f3386,f3388,f3390,f3435])).
fof(f3391,plain,(
  $false | ($spl26 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f1414,f454,f1400,f1415,f1502,f459,f1515,f1980,f1979,f1964,f1843,f1811,f1299,f1145,f534,f517,f512,f195,f1144,f2104,f535,f828,f908,f261,f150,f149,f137,f1816,f1274,f1612,f1627,f1840,f1841,f1269,f1695,f1716,f1628,f478,f477,f471,f458,f453,f189,f1715,f1712,f650,f235,f2849,f2845,f3068,f3169,f3272,f3386,f3388,f3390])).
fof(f3389,plain,(
  $false | ($spl26 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f1414,f454,f1400,f1415,f1502,f459,f1515,f1980,f1979,f1964,f1843,f1811,f1299,f1145,f534,f517,f512,f195,f1144,f2104,f535,f828,f908,f261,f150,f149,f137,f1816,f1274,f1612,f1627,f1840,f1841,f1269,f1695,f1716,f1628,f478,f477,f471,f458,f453,f189,f1715,f1712,f650,f235,f2849,f2845,f3068,f3169,f3272,f3386,f3388])).
fof(f3387,plain,(
  $false | ($spl26 | $spl30)),
  inference(global_subsumption,[],[f133,f132,f131,f134,f136,f138,f139,f140,f141,f142,f146,f145,f183,f184,f193,f190,f194,f221,f242,f240,f274,f280,f300,f340,f389,f402,f429,f480,f507,f662,f766,f867,f1141,f1142,f907,f1298,f1414,f454,f1400,f1415,f1502,f459,f1515,f1980,f1979,f1964,f1843,f1811,f1299,f1145,f534,f517,f512,f195,f1144,f2104,f535,f828,f908,f261,f150,f149,f137,f1816,f1274,f1612,f1627,f1840,f1841,f1269,f1695,f1716,f1628,f478,f477,f471,f458,f453,f189,f1715,f1712,f650,f235,f2849,f2845,f3068,f3169,f3272,f3386])).
fof(f3385,plain,(
  $false | ($spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3304,f134])).
fof(f3304,plain,(
  ~p(s(bool,'T_0')) | ($spl26 | $spl30)),
  inference(backward_demodulation,[],[f3272,f137])).
fof(f3303,plain,(
  $false | ($spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3302,f240])).
fof(f3302,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl26 | $spl30)),
  inference(forward_demodulation,[],[f3301,f2849])).
fof(f3301,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | $spl26),
  inference(subsumption_resolution,[],[f3167,f242])).
fof(f3167,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X19)))) ) | $spl26),
  inference(superposition,[],[f2845,f261])).
fof(f3300,plain,(
  $false | ($spl1 | $spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3299,f240])).
fof(f3299,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl26 | $spl30)),
  inference(forward_demodulation,[],[f3298,f2849])).
fof(f3298,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl1 | $spl26)),
  inference(subsumption_resolution,[],[f3162,f241])).
fof(f3162,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | ($spl1 | $spl26)),
  inference(superposition,[],[f2845,f260])).
fof(f3297,plain,(
  $false | ($spl1 | $spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3296,f240])).
fof(f3296,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl26 | $spl30)),
  inference(condensation,[],[f3295])).
fof(f3295,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl26 | $spl30)),
  inference(forward_demodulation,[],[f3294,f2849])).
fof(f3294,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl26)),
  inference(forward_demodulation,[],[f3277,f260])).
fof(f3277,plain,(
  ( ! [X9] : (s(bool,g(s(bool,'T_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl26)),
  inference(duplicate_literal_removal,[],[f3157])).
fof(f3157,plain,(
  ( ! [X9] : (s(bool,g(s(bool,'T_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl26)),
  inference(superposition,[],[f2845,f506])).
fof(f3293,plain,(
  $false | ($spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3292,f240])).
fof(f3292,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl26 | $spl30)),
  inference(condensation,[],[f3291])).
fof(f3291,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl26 | $spl30)),
  inference(forward_demodulation,[],[f3282,f2849])).
fof(f3282,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f3152])).
fof(f3152,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(superposition,[],[f2845,f141])).
fof(f3290,plain,(
  $false | ($spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3289,f240])).
fof(f3289,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl26 | $spl30)),
  inference(forward_demodulation,[],[f3149,f2849])).
fof(f3149,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | $spl26),
  inference(unit_resulting_resolution,[],[f240,f2845])).
fof(f3288,plain,(
  $false | ($spl1 | $spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3287,f240])).
fof(f3287,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl26 | $spl30)),
  inference(forward_demodulation,[],[f3286,f2849])).
fof(f3286,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl1 | $spl26)),
  inference(forward_demodulation,[],[f3150,f260])).
fof(f3150,plain,(
  s(bool,g(s(bool,'T_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0))) | ($spl1 | $spl26)),
  inference(unit_resulting_resolution,[],[f241,f2845])).
fof(f3285,plain,(
  $false | ($spl26 | $spl30)),
  inference(subsumption_resolution,[],[f3284,f240])).
fof(f3284,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl26 | $spl30)),
  inference(forward_demodulation,[],[f3283,f2849])).
fof(f3283,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | $spl26),
  inference(forward_demodulation,[],[f3151,f261])).
fof(f3151,plain,(
  ( ! [X0] : (s(bool,g(s(bool,'T_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0)))) ) | $spl26),
  inference(unit_resulting_resolution,[],[f242,f2845])).
fof(f2845,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(cnf_transformation,[],[f2845_D])).
fof(f2845_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2846,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl28),
  inference(cnf_transformation,[],[f2846_D])).
fof(f2846_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2940,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X7) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl24),
  inference(superposition,[],[f300,f2808])).
fof(f2849,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl30),
  inference(cnf_transformation,[],[f2849_D])).
fof(f2849_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2950,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl36),
  inference(cnf_transformation,[],[f2950_D])).
fof(f2950_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2939,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl24),
  inference(superposition,[],[f300,f2808])).
fof(f2938,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl24),
  inference(superposition,[],[f300,f2808])).
fof(f2937,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X4) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl24),
  inference(superposition,[],[f300,f2808])).
fof(f2851,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl34),
  inference(cnf_transformation,[],[f2851_D])).
fof(f2851_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2928,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl24),
  inference(superposition,[],[f141,f2808])).
fof(f2927,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl24),
  inference(superposition,[],[f2808,f141])).
fof(f2926,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl24),
  inference(superposition,[],[f2808,f300])).
fof(f2925,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl24),
  inference(superposition,[],[f2808,f300])).
fof(f2924,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(superposition,[],[f2808,f300])).
fof(f2923,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f2808,f300])).
fof(f2852,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl35),
  inference(cnf_transformation,[],[f2852_D])).
fof(f2852_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f2842,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl25),
  inference(superposition,[],[f2809,f141])).
fof(f2850,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl32),
  inference(cnf_transformation,[],[f2850_D])).
fof(f2850_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2841,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl25),
  inference(superposition,[],[f2809,f300])).
fof(f2840,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl25),
  inference(superposition,[],[f2809,f300])).
fof(f2847,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl29),
  inference(cnf_transformation,[],[f2847_D])).
fof(f2847_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f2839,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl25),
  inference(superposition,[],[f2809,f300])).
fof(f2838,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl25),
  inference(superposition,[],[f2809,f300])).
fof(f2807,plain,(
  ( ! [X52,X51] : (s(fun(bool,bool),X51) = s(fun(bool,bool),X52)) ) | $spl22),
  inference(cnf_transformation,[],[f2807_D])).
fof(f2807_D,plain,(
  ( ! [X52,X51] : (s(fun(bool,bool),X51) = s(fun(bool,bool),X52)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2809,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | $spl25),
  inference(cnf_transformation,[],[f2809_D])).
fof(f2809_D,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f2793,plain,(
  ( ! [X52,X51] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2792])).
fof(f2792,plain,(
  ( ! [X52,X51] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52)) ) | $spl8),
  inference(superposition,[],[f194,f629])).
fof(f2369,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2218,f134])).
fof(f2218,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f240,f1488,f193])).
fof(f2368,plain,(
  $false | ($spl1 | $spl16)),
  inference(subsumption_resolution,[],[f2220,f134])).
fof(f2220,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f241,f1488,f193])).
fof(f2367,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2222,f134])).
fof(f2222,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f242,f1488,f193])).
fof(f2366,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2224,f1488])).
fof(f2224,plain,(
  ~p(s(bool,'F_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f240,f1488,f193])).
fof(f2365,plain,(
  $false | ($spl1 | $spl16)),
  inference(subsumption_resolution,[],[f2226,f1488])).
fof(f2226,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f241,f1488,f193])).
fof(f2364,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2228,f1488])).
fof(f2228,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f242,f1488,f193])).
fof(f2363,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2232,f134])).
fof(f2232,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f240,f1488,f193])).
fof(f2362,plain,(
  $false | ($spl1 | $spl16)),
  inference(subsumption_resolution,[],[f2234,f134])).
fof(f2234,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f241,f1488,f193])).
fof(f2361,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2236,f134])).
fof(f2236,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f242,f1488,f193])).
fof(f2360,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2238,f1488])).
fof(f2238,plain,(
  ~p(s(bool,'F_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f240,f1488,f193])).
fof(f2359,plain,(
  $false | ($spl1 | $spl16)),
  inference(subsumption_resolution,[],[f2240,f1488])).
fof(f2240,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK0)))) | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f241,f1488,f193])).
fof(f2358,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2242,f1488])).
fof(f2242,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f242,f1488,f193])).
fof(f2357,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2356,f1488])).
fof(f2356,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl16),
  inference(subsumption_resolution,[],[f2304,f2246])).
fof(f2246,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(unit_resulting_resolution,[],[f1488,f146])).
fof(f2304,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl16),
  inference(backward_demodulation,[],[f2246,f149])).
fof(f2355,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2303,f134])).
fof(f2303,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f2246,f137])).
fof(f2252,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f242,f1488,f471])).
fof(f2247,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f240,f1488,f146])).
fof(f2245,plain,(
  $false | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f260,f1488,f149])).
fof(f2243,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f1488,f242,f1488,f193])).
fof(f2241,plain,(
  $false | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f1488,f241,f1488,f193])).
fof(f2239,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f1488,f240,f1488,f193])).
fof(f2237,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f134,f242,f1488,f193])).
fof(f2235,plain,(
  $false | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f134,f241,f1488,f193])).
fof(f2233,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f134,f240,f1488,f193])).
fof(f2229,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f1488,f242,f1488,f193])).
fof(f2227,plain,(
  $false | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f1488,f241,f1488,f193])).
fof(f2225,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f1488,f240,f1488,f193])).
fof(f2223,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f134,f242,f1488,f193])).
fof(f2221,plain,(
  $false | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f134,f241,f1488,f193])).
fof(f2219,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f134,f240,f1488,f193])).
fof(f2211,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f137,f261,f1488,f221])).
fof(f2209,plain,(
  $false | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f137,f260,f1488,f221])).
fof(f2195,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f137,f1488])).
fof(f2191,plain,(
  ( ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))),s('type.num.num',sK0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) ) | $spl20),
  inference(cnf_transformation,[],[f2191_D])).
fof(f2191_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))),s('type.num.num',sK0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2192,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X6)))),s('type.num.num',sK0)))) ) | $spl2),
  inference(subsumption_resolution,[],[f2186,f145])).
fof(f2186,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X6)))),s('type.num.num',sK0)))) ) | $spl2),
  inference(superposition,[],[f1839,f300])).
fof(f1839,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X0)))),s('type.num.num',sK0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl2),
  inference(subsumption_resolution,[],[f1818,f134])).
fof(f1818,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X0)))),s('type.num.num',sK0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl2),
  inference(superposition,[],[f512,f245])).
fof(f245,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | $spl2),
  inference(unit_resulting_resolution,[],[f230,f146])).
fof(f2190,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))),s('type.num.num',sK0)))) ) | $spl2),
  inference(subsumption_resolution,[],[f2185,f145])).
fof(f2185,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1(X4)))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,X5)) ) | $spl2),
  inference(superposition,[],[f1839,f300])).
fof(f1650,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl18),
  inference(cnf_transformation,[],[f1650_D])).
fof(f1650_D,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1651,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f1634,f145])).
fof(f1634,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(superposition,[],[f1604,f300])).
fof(f1649,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl10),
  inference(subsumption_resolution,[],[f1633,f145])).
fof(f1633,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X5)) ) | $spl10),
  inference(superposition,[],[f1604,f300])).
fof(f1487,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X8)) ) | $spl14),
  inference(cnf_transformation,[],[f1487_D])).
fof(f1487_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1489,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) != s(bool,f(s(bool,X11))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X9)) ) | $spl11),
  inference(subsumption_resolution,[],[f1466,f145])).
fof(f1466,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) != s(bool,f(s(bool,X11))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl11),
  inference(superposition,[],[f1215,f300])).
fof(f1486,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s(bool,X6) != s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X6)) ) | $spl11),
  inference(subsumption_resolution,[],[f1465,f145])).
fof(f1465,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s(bool,X6) != s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) ) | $spl11),
  inference(superposition,[],[f1215,f300])).
fof(f1147,plain,(
  ( ! [X78] : (s(bool,X78) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X78)) ) | $spl12),
  inference(cnf_transformation,[],[f1147_D])).
fof(f1147_D,plain,(
  ( ! [X78] : (s(bool,X78) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X78)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f629,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl8),
  inference(cnf_transformation,[],[f629_D])).
fof(f629_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1152,plain,(
  ( ! [X78,X76,X77] : (s(fun(bool,bool),X76) = s(fun(bool,bool),X77) | s(bool,'T_0') = s(bool,sK2(X76,X77,bool,bool)) | s(bool,X78) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X78)) ) | ($spl1 | $spl11)),
  inference(subsumption_resolution,[],[f953,f636])).
fof(f953,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X76) = s(fun(bool,bool),X77) | s(bool,'T_0') = s(bool,sK2(X76,X77,bool,bool)) | s(bool,X78) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X78)) ) | $spl1),
  inference(superposition,[],[f622,f300])).
fof(f1146,plain,(
  ( ! [X78,X76,X77] : (s(fun(bool,bool),X76) = s(fun(bool,bool),X77) | s(bool,'T_0') = s(bool,sK2(X76,X77,bool,bool)) | s(bool,X78) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X78)) ) | ($spl1 | $spl11)),
  inference(subsumption_resolution,[],[f909,f636])).
fof(f909,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X76) = s(fun(bool,bool),X77) | s(bool,'T_0') = s(bool,sK2(X76,X77,bool,bool)) | s(bool,X78) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X78)) ) | $spl1),
  inference(superposition,[],[f622,f300])).
fof(f631,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl11),
  inference(cnf_transformation,[],[f631_D])).
fof(f631_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f626,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl1),
  inference(superposition,[],[f622,f141])).
fof(f228,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0)))) | $spl0),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f230,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl2),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f131,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0))))),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) & s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f97])).
fof(f97,plain,(
  ? [X0] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ? [X0] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) | (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ? [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) <~> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ~! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) & s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) & s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))))),
  file('thm.arithmetic.ONE_MOD_IFF','thm.arithmetic.ONE_MOD_IFF')).
fof(f232,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0))) | $spl4),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f132,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0))))),
  inference(cnf_transformation,[],[f98])).
fof(f229,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0)))) | $spl1),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f231,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl3),
  inference(cnf_transformation,[],[f231_D])).
fof(f231_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f233,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0))) | $spl5),
  inference(cnf_transformation,[],[f233_D])).
fof(f233_D,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f133,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK0))))),
  inference(cnf_transformation,[],[f98])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 353
% Passive clauses: 1765
% Generated clauses: 48193
% Final active clauses: 73
% Final passive clauses: 151
% Input formulas: 36
% Initial clauses: 79
% 
% Duplicate literals: 2115
% Trivial inequalities: 620
% Fw subsumption resolutions: 4905
% Fw demodulations: 6449
% Bw demodulations: 1134
% Condensations: 2626
% Global subsumptions: 405
% 
% Simple tautologies: 553
% Equational tautologies: 228
% Forward subsumptions: 27955
% Bw demodulations to eq. taut.: 28
% 
% Unit resulting resolution: 1232
% Factoring: 1
% Forward superposition: 12211
% Backward superposition: 10961
% Self superposition: 334
% Equality factoring: 163
% Equality resolution: 47
% 
% Split clauses: 117
% Split components: 237
% SAT solver clauses: 2873
% SAT solver unit clauses: 101
% SAT solver binary clauses: 403
% 
% Sat splits: 117
% Sat splitting refutations: 467
% 
% Memory used [KB]: 16119
% Time elapsed: 1.438 s
% ------------------------------
% ------------------------------
% Success in time 2.757 s
