Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.semi_ring
% (104283)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 160
% Passive clauses: 3219
% Generated clauses: 21301
% Final active clauses: 133
% Final passive clauses: 2870
% Input formulas: 68
% Initial clauses: 85
% 
% Duplicate literals: 794
% Trivial inequalities: 107
% Fw subsumption resolutions: 1649
% Fw demodulations: 3248
% Bw demodulations: 24
% 
% Simple tautologies: 157
% Equational tautologies: 51
% Forward subsumptions: 11527
% 
% Binary resolution: 438
% Factoring: 2
% Forward superposition: 7873
% Backward superposition: 6653
% Self superposition: 216
% Equality factoring: 96
% Equality resolution: 14
% 
% Split clauses: 220
% Split components: 440
% SAT solver clauses: 636
% SAT solver unit clauses: 26
% SAT solver binary clauses: 121
% 
% Sat splits: 220
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 636
% TWLsolver calls for satisfiability: 142
% 
% Memory used [KB]: 20596
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.semi_ring
% (104292)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 119
% Passive clauses: 1169
% Generated clauses: 13676
% Final active clauses: 99
% Final passive clauses: 935
% Input formulas: 68
% Initial clauses: 85
% 
% Duplicate literals: 433
% Trivial inequalities: 83
% Fw subsumption resolutions: 705
% Fw demodulations: 1598
% Bw demodulations: 134
% Condensations: 964
% 
% Simple tautologies: 146
% Equational tautologies: 40
% Forward subsumptions: 6692
% 
% Binary resolution: 271
% Factoring: 1
% Forward superposition: 4557
% Backward superposition: 3532
% Self superposition: 171
% Equality factoring: 76
% Equality resolution: 11
% 
% Split clauses: 48
% Split components: 96
% SAT solver clauses: 1919
% SAT solver unit clauses: 125
% SAT solver binary clauses: 40
% 
% Sat splits: 48
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1919
% TWLsolver calls for satisfiability: 268
% 
% Memory used [KB]: 17782
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.semi_ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.semi_ring
% SZS output start Proof for thm.semi_ring
fof(f51414,plain,(
  $false),
  inference(subsumption_resolution,[],[f51413,f222])).
fof(f222,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f22])).
fof(f22,axiom,(
  p(s(bool,'T_0'))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|5','thm.bool.TRUTH')).
fof(f51413,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f51409,f4063])).
fof(f4063,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f3998,f313])).
fof(f313,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f224,f225])).
fof(f225,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|5',aHL_BOOL_CASES)).
fof(f224,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|5',aHL_FALSITY)).
fof(f3998,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(backward_demodulation,[],[f3873,f1955])).
fof(f1955,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f1876,f285])).
fof(f285,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6,X13,X7] : ! [X10,X9,X5] : s(X6,'const.combin.o_3'(s(fun(X13,X6),X10),s(fun(X7,X13),X9),s(X7,X5))) = s(X6,happ(s(fun(X13,X6),X10),s(X13,happ(s(fun(X7,X13),X9),s(X7,X5)))))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|5','thm.combin.o_DEF')).
fof(f1876,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f232,f283])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X35,X37,X38] : (p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),'=_0'),s(X35,X37))),s(X35,X38)))) <=> s(X35,X37) = s(X35,X38))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|5',arity_0_HL_EQ)).
fof(f232,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X31] : ! [X33] : s(X31,X33) = s(X31,happ(s(fun(X31,X31),'_i_0_0'),s(X31,X33)))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|5','_i0_ax')).
fof(f3873,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X1))) = s(X2,'const.combin.o_3'(s(fun(X0,X2),X3),s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(superposition,[],[f285,f232])).
fof(f51409,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))))),
  inference(backward_demodulation,[],[f51305,f14441])).
fof(f14441,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),'const.combin.o_2'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4),s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))),
  inference(unit_resulting_resolution,[],[f221,f282])).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f220])).
fof(f221,plain,(
  s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),'const.combin.o_2'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4),s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))) != s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),'const.combin.o_2'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4),s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))) != s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f192])).
fof(f192,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2))))) != s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),'const.combin.o_2'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4),s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2))))) = s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),'const.combin.o_2'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4),s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2))))) = s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),'const.combin.o_2'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4),s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8))))) = s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),'const.combin.o_2'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10),s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8))))) = s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),'const.combin.o_2'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10),s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8)))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|5','thm.semi_ring.semi_ring_fupdfupds_comp|split|5')).
fof(f51305,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(fun(X11,'type.semi_ring.semi_ring'(X8)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X8),'type.semi_ring.semi_ring'(X8)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X9))),s(fun(X11,'type.semi_ring.semi_ring'(X8)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X8),'type.semi_ring.semi_ring'(X8)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X10))),s(fun(X11,'type.semi_ring.semi_ring'(X8)),X12))))) = s(fun(X11,'type.semi_ring.semi_ring'(X8)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X8),'type.semi_ring.semi_ring'(X8)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),'const.combin.o_2'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X9),s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X10))))),s(fun(X11,'type.semi_ring.semi_ring'(X8)),X12)))) )),
  inference(superposition,[],[f289,f247])).
fof(f247,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.o_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1)))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.o_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1)))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.o_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),'const.combin.o_2'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10),s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))))) = s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10))),s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9)))))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|5','thm.semi_ring.semi_ring_fupdfupds_comp|split|4')).
fof(f289,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7,X17,X13] : ! [X10,X9,X8] : s(fun(X17,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X17,X7),'const.combin.o_2'(s(fun(X13,X7),X9),s(fun(X17,X13),X8))))) = s(fun(X17,X6),'const.combin.o_2'(s(fun(X13,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X13,X7),X9))),s(fun(X17,X13),X8)))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|5','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.semi_ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 302
% Passive clauses: 5010
% Generated clauses: 58037
% Final active clauses: 125
% Final passive clauses: 1762
% Input formulas: 68
% Initial clauses: 85
% 
% Duplicate literals: 1351
% Trivial inequalities: 153
% Fw subsumption resolutions: 2602
% Fw demodulations: 9072
% Bw demodulations: 2713
% Condensations: 2302
% Global subsumptions: 14
% 
% Simple tautologies: 178
% Equational tautologies: 282
% Forward subsumptions: 36179
% Fw demodulations to eq. taut.: 41
% Bw demodulations to eq. taut.: 79
% 
% Unit resulting resolution: 1650
% Factoring: 1
% Forward superposition: 16694
% Backward superposition: 13695
% Self superposition: 371
% Equality factoring: 142
% Equality resolution: 13
% 
% Split clauses: 93
% Split components: 188
% SAT solver clauses: 5879
% SAT solver unit clauses: 1073
% SAT solver binary clauses: 2568
% 
% Sat splits: 93
% Sat splitting refutations: 667
% 
% Memory used [KB]: 51939
% Time elapsed: 2.289 s
% ------------------------------
% ------------------------------
% Success in time 3.6 s
