Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (70518)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1034
% Passive clauses: 9798
% Generated clauses: 30779
% Final active clauses: 957
% Final passive clauses: 8095
% Input formulas: 62
% Initial clauses: 223
% 
% Duplicate literals: 435
% Trivial inequalities: 196
% Fw subsumption resolutions: 3457
% Fw demodulations: 2554
% Bw demodulations: 75
% 
% Simple tautologies: 291
% Equational tautologies: 41
% Forward subsumptions: 13143
% 
% Binary resolution: 9927
% Factoring: 1210
% Forward superposition: 4456
% Backward superposition: 7814
% Self superposition: 33
% Equality factoring: 26
% Equality resolution: 12
% 
% Split clauses: 1014
% Split components: 2034
% SAT solver clauses: 1445
% SAT solver unit clauses: 14
% SAT solver binary clauses: 194
% 
% Sat splits: 1014
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1443
% TWLsolver calls for satisfiability: 505
% 
% Memory used [KB]: 15991
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (70535)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 477
% Passive clauses: 2967
% Generated clauses: 15863
% Final active clauses: 411
% Final passive clauses: 2032
% Input formulas: 62
% Initial clauses: 223
% 
% Duplicate literals: 279
% Trivial inequalities: 68
% Fw subsumption resolutions: 1255
% Fw demodulations: 1008
% Bw demodulations: 492
% Condensations: 395
% Global subsumptions: 22
% 
% Simple tautologies: 284
% Equational tautologies: 19
% Forward subsumptions: 6359
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 3832
% Factoring: 207
% Forward superposition: 2697
% Backward superposition: 1818
% Self superposition: 39
% Equality factoring: 28
% Equality resolution: 11
% 
% Split clauses: 141
% Split components: 282
% SAT solver clauses: 3350
% SAT solver unit clauses: 223
% SAT solver binary clauses: 85
% SAT solver learnt clauses: 45
% SAT solver learnt literals: 11
% 
% Sat splits: 141
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 3261
% TWLsolver calls for satisfiability: 1781
% 
% Memory used [KB]: 13560
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f100982,plain,(
  $false),
  inference(subsumption_resolution,[],[f100702,f60741])).
fof(f60741,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK16),s(sK13,X0))),s(sK12,happ(s(fun(sK13,sK12),sK14),s(sK13,X0))))))) )),
  inference(forward_demodulation,[],[f60740,f1853])).
fof(f1853,plain,(
  ( ! [X0] : (s(sK12,happ(s(fun(sK13,sK12),sK14),s(sK13,X0))) = s(sK12,'const.relation.WFREC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK17),s(sK13,X0)))) )),
  inference(superposition,[],[f456,f315])).
fof(f315,plain,(
  s(fun(sK13,sK12),sK14) = s(fun(sK13,sK12),'const.relation.WFREC_2'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK17)))),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  s(fun(sK13,sK12),sK14) = s(fun(sK13,sK12),'const.relation.WFREC_2'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK17))) & p(s(bool,'const.relation.WF_1'(s(fun(sK13,fun(sK13,bool)),sK15)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(sK13,fun(sK12,bool)),sK16),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK17)))) & ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK16),s(sK13,sK18))),s(sK12,happ(s(fun(sK13,sK12),sK14),s(sK13,sK18))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16,sK17,sK18])],[f157])).
fof(f157,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (s(fun(X1,X0),X2) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X5))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,fun(X0,bool)),X4),s(fun(fun(X1,X0),fun(X1,X0)),X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((s(fun(X1,X0),X2) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X5))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,fun(X0,bool)),X4),s(fun(fun(X1,X0),fun(X1,X0)),X5))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : ((s(fun(X1,X0),X2) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X5))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,fun(X0,bool)),X4),s(fun(fun(X1,X0),fun(X1,X0)),X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5,X6] : ((s(fun(X1,X0),X2) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X5))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,fun(X0,bool)),X4),s(fun(fun(X1,X0),fun(X1,X0)),X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11,X12] : ((s(fun(X7,X6),X8) = s(fun(X7,X6),'const.relation.WFREC_2'(s(fun(X7,fun(X7,bool)),X9),s(fun(fun(X7,X6),fun(X7,X6)),X11))) & p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),X9)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X7,fun(X7,bool)),X9),s(fun(X7,fun(X6,bool)),X10),s(fun(fun(X7,X6),fun(X7,X6)),X11))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X10),s(X7,X12))),s(X6,happ(s(fun(X7,X6),X8),s(X7,X12)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11,X12] : ((s(fun(X7,X6),X8) = s(fun(X7,X6),'const.relation.WFREC_2'(s(fun(X7,fun(X7,bool)),X9),s(fun(fun(X7,X6),fun(X7,X6)),X11))) & p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),X9)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X7,fun(X7,bool)),X9),s(fun(X7,fun(X6,bool)),X10),s(fun(fun(X7,X6),fun(X7,X6)),X11))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X10),s(X7,X12))),s(X6,happ(s(fun(X7,X6),X8),s(X7,X12)))))))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC','thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC')).
fof(f456,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4)))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3))),s(X1,X4)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6,X7] : ! [X22,X23,X24] : s(X6,'const.relation.WFREC_3'(s(fun(X7,fun(X7,bool)),X22),s(fun(fun(X7,X6),fun(X7,X6)),X23),s(X7,X24))) = s(X6,happ(s(fun(X7,X6),'const.relation.WFREC_2'(s(fun(X7,fun(X7,bool)),X22),s(fun(fun(X7,X6),fun(X7,X6)),X23))),s(X7,X24)))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC','ARITY_2_3_const.relation.WFREC')).
fof(f60740,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK16),s(sK13,X0))),s(sK12,'const.relation.WFREC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK17),s(sK13,X0))))))) )),
  inference(subsumption_resolution,[],[f60719,f319])).
fof(f319,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC','thm.bool.TRUTH')).
fof(f60719,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK16),s(sK13,X0))),s(sK12,'const.relation.WFREC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK17),s(sK13,X0))))))) )),
  inference(superposition,[],[f2159,f791])).
fof(f791,plain,(
  s(bool,'T_0') = s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(sK13,fun(sK12,bool)),sK16),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK17)))),
  inference(unit_resulting_resolution,[],[f317,f340])).
fof(f340,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC','thm.bool.EQ_CLAUSES|split|0')).
fof(f317,plain,(
  p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(sK13,fun(sK12,bool)),sK16),s(fun(fun(sK13,sK12),fun(sK13,sK12)),sK17))))),
  inference(cnf_transformation,[],[f211])).
fof(f2159,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(sK13,fun(X0,bool)),X1),s(fun(fun(sK13,X0),fun(sK13,X0)),X2)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(sK13,fun(X0,bool)),X1),s(sK13,X3))),s(X0,'const.relation.WFREC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(fun(sK13,X0),fun(sK13,X0)),X2),s(sK13,X3))))))) )),
  inference(subsumption_resolution,[],[f2151,f319])).
fof(f2151,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(sK13,fun(X0,bool)),X1),s(fun(fun(sK13,X0),fun(sK13,X0)),X2)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(sK13,fun(X0,bool)),X1),s(sK13,X3))),s(X0,'const.relation.WFREC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(fun(fun(sK13,X0),fun(sK13,X0)),X2),s(sK13,X3))))))) )),
  inference(superposition,[],[f457,f597])).
fof(f597,plain,(
  s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(sK13,fun(sK13,bool)),sK15)))),
  inference(unit_resulting_resolution,[],[f316,f340])).
fof(f316,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun(sK13,fun(sK13,bool)),sK15))))),
  inference(cnf_transformation,[],[f211])).
fof(f457,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X5))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X5))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X5)))))))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) | ~p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4))))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X5))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X5)))))))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4))))) => ! [X5] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X5))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X5)))))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,fun(X0,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4))))) => ! [X5] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X3),s(X1,X5))),s(X0,'const.relation.WFREC_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X5)))))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6,X7] : ! [X9,X10,X11] : ((p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),X9)))) & p(s(bool,'const.relation.INDUCTIVE_INVARIANT_3'(s(fun(X7,fun(X7,bool)),X9),s(fun(X7,fun(X6,bool)),X10),s(fun(fun(X7,X6),fun(X7,X6)),X11))))) => ! [X12] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X10),s(X7,X12))),s(X6,'const.relation.WFREC_3'(s(fun(X7,fun(X7,bool)),X9),s(fun(fun(X7,X6),fun(X7,X6)),X11),s(X7,X12)))))))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC','thm.relation.INDUCTIVE_INVARIANT_WFREC')).
fof(f100702,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK16),s(sK13,sK18))),s(sK12,happ(s(fun(sK13,sK12),sK14),s(sK13,sK18))))))),
  inference(unit_resulting_resolution,[],[f960,f550,f9115,f9115,f7587])).
fof(f7587,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP11(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f7070,f319])).
fof(f7070,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP11(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f535,f866])).
fof(f866,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f327,f327])).
fof(f327,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X13))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC','thm.bool.BOOL_CASES_AX')).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f192,f209,f208])).
fof(f208,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f209,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X19,X20,X21] : ((p(s(bool,X21)) <=> (p(s(bool,X20)) => p(s(bool,X19)))) <=> ((p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X21)) | ~p(s(bool,X19))) & (~p(s(bool,X20)) | p(s(bool,X19)) | ~p(s(bool,X21)))))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC','thm.sat.sat_27_prove_ltb')).
fof(f9115,plain,(
  s(bool,'T_0') != s(bool,sK21('const.combin.I_0','F_0',bool))),
  inference(unit_resulting_resolution,[],[f2164,f339])).
fof(f339,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f2164,plain,(
  ~p(s(bool,sK21('const.combin.I_0','F_0',bool)))),
  inference(unit_resulting_resolution,[],[f550,f1817])).
fof(f1817,plain,(
  ( ! [X0] : (~p(s(bool,sK21('const.combin.I_0',X0,bool))) | p(s(bool,X0))) )),
  inference(condensation,[],[f1816])).
fof(f1816,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | ~p(s(bool,sK21('const.combin.I_0',X4,bool))) | p(s(bool,X4))) )),
  inference(forward_demodulation,[],[f1813,f551])).
fof(f551,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f397,f396])).
fof(f396,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X7] : ! [X12] : s(X7,X12) = s(X7,'const.combin.I_1'(s(X7,X12)))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC','thm.combin.I_THM')).
fof(f397,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X7] : ! [X22] : s(X7,'const.combin.I_1'(s(X7,X22))) = s(X7,happ(s(fun(X7,X7),'const.combin.I_0'),s(X7,X22)))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC','ARITY_0_1_const.combin.I')).
fof(f1813,plain,(
  ( ! [X4,X5] : (~p(s(bool,sK21('const.combin.I_0',X4,bool))) | p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X5))))) )),
  inference(superposition,[],[f407,f551])).
fof(f407,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0)))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0)))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22])],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X7] : ! [X10,X17] : (! [X12] : (p(s(bool,X10)) | p(s(bool,happ(s(fun(X7,bool),X17),s(X7,X12))))) <=> (p(s(bool,X10)) | ! [X12] : p(s(bool,happ(s(fun(X7,bool),X17),s(X7,X12))))))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC','thm.bool.RIGHT_FORALL_OR_THM')).
fof(f550,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f322,f319])).
fof(f322,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f31])).
fof(f31,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.relation.TFL_INDUCTIVE_INVARIANT_WFREC','thm.bool.NOT_CLAUSES|split|1')).
fof(f960,plain,(
  ( ! [X0] : (sP11('T_0',X0,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK16),s(sK13,sK18))),s(sK12,happ(s(fun(sK13,sK12),sK14),s(sK13,sK18)))))) )),
  inference(unit_resulting_resolution,[],[f319,f318,f533])).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f314])).
fof(f318,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK13,fun(sK12,bool)),sK16),s(sK13,sK18))),s(sK12,happ(s(fun(sK13,sK12),sK14),s(sK13,sK18))))))),
  inference(cnf_transformation,[],[f211])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1977
% Passive clauses: 7185
% Generated clauses: 99314
% Final active clauses: 1945
% Final passive clauses: 4380
% Input formulas: 62
% Initial clauses: 223
% 
% Duplicate literals: 3921
% Trivial inequalities: 168
% Fw subsumption resolutions: 6311
% Fw demodulations: 13165
% Bw demodulations: 199
% Condensations: 6657
% Global subsumptions: 68
% 
% Simple tautologies: 216
% Equational tautologies: 869
% Forward subsumptions: 57698
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 19054
% Factoring: 7
% Forward superposition: 23458
% Backward superposition: 25462
% Self superposition: 180
% Equality factoring: 144
% Equality resolution: 3
% 
% Split clauses: 365
% Split components: 761
% SAT solver clauses: 9746
% SAT solver unit clauses: 2672
% SAT solver binary clauses: 2926
% 
% Sat splits: 365
% Sat splitting refutations: 1761
% 
% Memory used [KB]: 45287
% Time elapsed: 6.219 s
% ------------------------------
% ------------------------------
% Success in time 7.528 s
