Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (57336)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 610
% Passive clauses: 9867
% Generated clauses: 17361
% Final active clauses: 545
% Final passive clauses: 8736
% Input formulas: 43
% Initial clauses: 179
% 
% Duplicate literals: 268
% Trivial inequalities: 88
% Fw subsumption resolutions: 1236
% Fw demodulations: 111
% Bw demodulations: 64
% 
% Simple tautologies: 209
% Equational tautologies: 66
% Forward subsumptions: 5364
% 
% Binary resolution: 10531
% Factoring: 1015
% Forward superposition: 1643
% Backward superposition: 1974
% Self superposition: 33
% Equality factoring: 26
% Equality resolution: 33
% 
% Split clauses: 289
% Split components: 579
% SAT solver clauses: 613
% SAT solver unit clauses: 15
% SAT solver binary clauses: 141
% 
% Sat splits: 289
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 611
% TWLsolver calls for satisfiability: 351
% 
% Memory used [KB]: 8571
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (57351)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 410
% Passive clauses: 2319
% Generated clauses: 11598
% Final active clauses: 367
% Final passive clauses: 1825
% Input formulas: 43
% Initial clauses: 179
% 
% Duplicate literals: 236
% Trivial inequalities: 113
% Fw subsumption resolutions: 775
% Fw demodulations: 734
% Bw demodulations: 91
% Condensations: 393
% Global subsumptions: 13
% 
% Simple tautologies: 287
% Equational tautologies: 18
% Forward subsumptions: 3747
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 3001
% Factoring: 80
% Forward superposition: 2321
% Backward superposition: 1189
% Self superposition: 32
% Equality factoring: 28
% Equality resolution: 25
% 
% Split clauses: 118
% Split components: 259
% SAT solver clauses: 3148
% SAT solver unit clauses: 186
% SAT solver binary clauses: 83
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 118
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 3096
% TWLsolver calls for satisfiability: 1886
% 
% Memory used [KB]: 13176
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f30364,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f246,f434,f433,f245,f436,f244,f438,f432,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2734,f2644,f2736,f2648,f2738,f2651,f2741,f2655,f2743,f2658,f2745,f2661,f2747,f2664,f2749,f2667,f2751,f2670,f2753,f2673,f2755,f2676,f2757,f2679,f2759,f2682,f2761,f2685,f2763,f2688,f2765,f2691,f2767,f2694,f2769,f2697,f2994,f2995,f2998,f3001,f3026,f3028,f3030,f3032,f3034,f3036,f3038,f3040,f3042,f3044,f3046,f3048,f3050,f3052,f3054,f3056,f3058,f3060,f3062,f3064,f3066,f3140,f3142,f3144,f3146,f3148,f3150,f3152,f3154,f3156,f3158,f3160,f3162,f3164,f3166,f3168,f3170,f3172,f3174,f3176,f3178,f3180,f3182,f3184,f3186,f3188,f3190,f3192,f3194,f3196,f3198,f3200,f3202,f3204,f3206,f3208,f3210,f3212,f3214,f3216,f3218,f3220,f3222,f4970,f4971,f4978,f4975,f4972,f4973,f4981,f4974,f4983,f4976,f4980,f5062,f5063,f5075,f5073,f7478,f7481,f7479,f7521,f7524,f7522,f7566,f7569,f7567,f7603,f7606,f7604,f7759,f7762,f7760,f7804,f7807,f7805,f7845,f7848,f7846,f7884,f7887,f7885,f7938,f7941,f7939,f8083,f8086,f8084,f9993,f9996,f9994,f10033,f10036,f10034,f10086,f10089,f10087,f10137,f10140,f10138,f10446,f10449,f10447,f10491,f10494,f10492,f10544,f10547,f10545,f10593,f10596,f10594,f10640,f10643,f10641,f10683,f10686,f10684,f10866,f10869,f10867,f10909,f10912,f10910,f12966,f12969,f12967,f13016,f13019,f2645,f13017,f18138,f18139,f18141,f18142,f18146,f18148,f18150,f18156,f18158,f18160,f18162,f18166,f18168,f18170,f18172,f18174,f18176,f18192,f18194,f18210,f18212,f18222,f18232,f18236,f18238,f18243,f18248,f18252,f18254,f18258,f18260,f18264,f18266,f18293,f18360,f18457,f18459,f18708,f18712,f18722,f18727,f18752,f18758,f18772,f18784,f18798,f18810,f18812,f18824,f18834,f18846,f18856,f18858,f18860,f18862,f18864,f21272,f21273,f21274,f18900,f18902,f18904,f18906,f18908,f18910,f18912,f18914,f18917,f18919,f18921,f18923,f18925,f18927,f18929,f18931,f18999,f19001,f19003,f19005,f19024,f19043,f19061,f19065,f19075,f19077,f19087,f19089,f19107,f19115,f19125,f19127,f19137,f19139,f19151,f19161,f19173,f19183,f19185,f19205,f19207,f19221,f19223,f19247,f19249,f19263,f19265,f19291,f19293,f19307,f19309,f19325,f19344,f19360,f19379,f19419,f19459,f19482,f19484,f19517,f19521,f19523,f19538,f19540,f19542,f19554,f19556,f19568,f19570,f19582,f19584,f19596,f19598,f19640,f19674,f19692,f19694,f19715,f19717,f19719,f19766,f19768,f19770,f19772,f19774,f19776,f19778,f19806,f19809,f19837,f19840,f19861,f19863,f19883,f19885,f19887,f19905,f19922,f19939,f19941,f19958,f19960,f19982,f20016,f20052,f20086,f20103,f20105,f20107,f20109,f20111,f20113,f20115,f20128,f20130,f20132,f20134,f20136,f20138,f20140,f20142,f20144,f20147,f20149,f20151,f20153,f20155,f20162,f20164,f20171,f20173,f20180,f20182,f20185,f20187,f20189,f20196,f20198,f20205,f20207,f20214,f20216,f20219,f20221,f20223,f20225,f20227,f20234,f20236,f20243,f20245,f20248,f20250,f20257,f20262,f20270,f20273,f20281,f20284,f20292,f20295,f20299,f20302,f20312,f20315,f20324,f20333,f20335,f20344,f20346,f20355,f20357,f20362,f20364,f20367,f20370,f20379,f20381,f20384,f20388,f20394,f20402,f20405,f20413,f20416,f20424,f20427,f20431,f20434,f20437,f20440,f20448,f20451,f20454,f20458,f20461,f20471,f20479,f20487,f20499,f20501,f20503,f20511,f20513,f20521,f20523,f20530,f20532,f20539,f20541,f20548,f20550,f20555,f20557,f20562,f20564,f20571,f20573,f20575,f20578,f20580,f20587,f20594,f20596,f20603,f20605,f20612,f20614,f20617,f20619,f20621,f20623,f20630,f20632,f20634,f20637,f20639,f20646,f20653,f20655,f20662,f20664,f20671,f20673,f20676,f20678,f20680,f20682,f20689,f20691,f20693,f20696,f20717,f20783,f20786,f20795,f20798,f20807,f20810,f20815,f20818,f20831,f20835,f20847,f20850,f20856,f20861,f20967,f21012,f21015,f21024,f21027,f21036,f21039,f21044,f21047,f21060,f21064,f21076,f21079,f21085,f21090,f21118,f21161,f21164,f21173,f21176,f21185,f21188,f21193,f21196,f21219,f21222,f21275,f21276,f21277,f21278,f21279,f21280,f21281,f21282,f21283,f21284,f21285,f21286,f21287,f21288,f21289,f21290,f21291,f21292,f21293,f21294,f21295,f21296,f21297,f21298,f21299,f21300,f21301,f21302,f21303,f21304,f21305,f21306,f21307,f21308,f21309,f21310,f21311,f21312,f21313,f21314,f21315,f21316,f21317,f21318,f21319,f21320,f21321,f21322,f21323,f21324,f21325,f21326,f21327,f21328,f21329,f21330,f21331,f21332,f21333,f21334,f21335,f21336,f21337,f21338,f21339,f21340,f21341,f21342,f21343,f21344,f21345,f21346,f21347,f21348,f21349,f21350,f21351,f21352,f21353,f21354,f21355,f21356,f21357,f21358,f21359,f21360,f21361,f21362,f21363,f21364,f21365,f21366,f21367,f21368,f21369,f21370,f21371,f21372,f21373,f21374,f21375,f21376,f21377,f21378,f21379,f21380,f21381,f21382,f21383,f21384,f21385,f21386,f21387,f21388,f21389,f21390,f21391,f21392,f21393,f21394,f21395,f21396,f21397,f21398,f21399,f21400,f21401,f21402,f21403,f21404,f21405,f21406,f21407,f21408,f21409,f21410,f21411,f21412,f21413,f21414,f21415,f21416,f21417,f21418,f21419,f21420,f21421,f21422,f21423,f21424,f21425,f21426,f21427,f21428,f21429,f21430,f21431,f21432,f21433,f21434,f21435,f21436,f21437,f21438,f21439,f21440,f21441,f21442,f21443,f21444,f21445,f21446,f21447,f21448,f21449,f21450,f21451,f21452,f21453,f21454,f21455,f21456,f21457,f21458,f21459,f21460,f21461,f21462,f21463,f21464,f21465,f21466,f21467,f21468,f21469,f21470,f21471,f21472,f21473,f21474,f21475,f21476,f21477,f21478,f21479,f21480,f21481,f21482,f21483,f21484,f21485,f21486,f21487,f21488,f21489,f21490,f21491,f21492,f21493,f21494,f21495,f21496,f21497,f21498,f21499,f21500,f21501,f21502,f21503,f21504,f21505,f21506,f21507,f21508,f21509,f21510,f21511,f21512,f21513,f21514,f21515,f21516,f21517,f21518,f21519,f21520,f21521,f21522,f21523,f21524,f21525,f21526,f21527,f21528,f21529,f21530,f21531,f21533,f21535,f21536,f21539,f21541,f21544,f21547,f21550,f21552,f21553,f21555,f21557,f21559,f21561,f21562,f21563,f21564,f21568,f21569,f21570,f21571,f21572,f21573,f21574,f21575,f21576,f21577,f21578,f21579,f21580,f21581,f21582,f21583,f21587,f21588,f21591,f21592,f21593,f21594,f21595,f21596,f21597,f21598,f21599,f21600,f21601,f21602,f21603,f21604,f21605,f21606,f21607,f21608,f21609,f21610,f21611,f21612,f21613,f21614,f21615,f21616,f21617,f21618,f21619,f21620,f21621,f21622,f21623,f21624,f21625,f21626,f21627,f21628,f21629,f21630,f21631,f21632,f21633,f21634,f21635,f21636,f21637,f21638,f21639,f21640,f21641,f21642,f21643,f21644,f21645,f21646,f21647,f21648,f21649,f21650,f21651,f21652,f21653,f21654,f21655,f21656,f21657,f21658,f21659,f21660,f21661,f21662,f21663,f21664,f21665,f21666,f21667,f21668,f21669,f21670,f21671,f21672,f21673,f21674,f21675,f21676,f21677,f21678,f21679,f21680,f21747,f21748,f21749,f18140,f21750,f22900,f22903,f22901,f23309,f23352,f23354,f23358,f23356,f30359])).
fof(f30359,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f29628,f23722])).
fof(f23722,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK23(sK17,sK16,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),sK15,sK14,sK13)))),s(sK13,sK14)))) | $spl0),
  inference(forward_demodulation,[],[f23384,f301])).
fof(f301,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1))) = s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1))) = s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X1)))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1))) = s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X8,X10] : s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X10),s(X6,X8))),s(X6,X7))) = s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X10),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X6,X8),s(X6,X7)))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.arithmetic.NRC_|49|')).
fof(f23384,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK23(sK17,sK16,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),sK15,sK14,sK13)),s(sK13,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f432,f1869])).
fof(f1869,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s(X1,X3),s(X1,X4)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,sK23(X2,X0,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X3,X4,X1)),s(X1,X4))))) )),
  inference(superposition,[],[f303,f254])).
fof(f254,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X12] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.arithmetic.ADD|49|')).
fof(f303,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X3),s(X0,sK23(X5,X4,X3,X2,X1,X0)),s(X0,X1))))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) | (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X4),s(X0,X2),s(X0,sK23(X5,X4,X3,X2,X1,X0))))) & p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X3),s(X0,sK23(X5,X4,X3,X2,X1,X0)),s(X0,X1)))))) & (! [X6] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X4),s(X0,X2),s(X0,X6)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X3),s(X0,X6),s(X0,X1))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s(X0,X2),s(X0,X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23])],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) | ? [X7] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X4),s(X0,X2),s(X0,X7)))) & p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X3),s(X0,X7),s(X0,X1)))))) & (! [X6] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X4),s(X0,X2),s(X0,X6)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X3),s(X0,X6),s(X0,X1))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s(X0,X2),s(X0,X1))))))),
  inference(rectify,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) | ? [X6] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X4),s(X0,X2),s(X0,X6)))) & p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X3),s(X0,X6),s(X0,X1)))))) & (! [X6] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X4),s(X0,X2),s(X0,X6)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X3),s(X0,X6),s(X0,X1))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s(X0,X2),s(X0,X1))))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) <=> ? [X6] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X4),s(X0,X2),s(X0,X6)))) & p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X3),s(X0,X6),s(X0,X1))))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) <=> ? [X6] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X4),s(X0,X2),s(X0,X6)))) & p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X3),s(X0,X6),s(X0,X1))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X11,X8,X9,X12,X10] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X9))),s(X6,X8),s(X6,X11)))) <=> ? [X7] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X10),s('type.num.num',X12),s(X6,X8),s(X6,X7)))) & p(s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X10),s('type.num.num',X9),s(X6,X7),s(X6,X11))))))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.arithmetic.NRC_ADD_EQN')).
fof(f29628,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK23(sK17,sK16,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),sK15,sK14,sK13)))),s(sK13,sK14)))) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f23385,f434])).
fof(f23385,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK23(sK17,sK16,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),sK15,sK14,sK13))))) | $spl0),
  inference(unit_resulting_resolution,[],[f432,f1822])).
fof(f1822,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s(X1,X3),s(X1,X4)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num',X0),s(X1,X3),s(X1,sK23(X2,X0,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X3,X4,X1)))))) )),
  inference(superposition,[],[f302,f254])).
fof(f302,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X4),s(X0,X2),s(X0,sK23(X5,X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f192])).
fof(f23356,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | $spl125),
  inference(cnf_transformation,[],[f23356_D])).
fof(f23356_D,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f23358,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,'const.min.@_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))))) | $spl126),
  inference(cnf_transformation,[],[f23358_D])).
fof(f23358_D,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,'const.min.@_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f23354,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,'const.min.@_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | $spl6),
  inference(forward_demodulation,[],[f23298,f254])).
fof(f23298,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(sK13,sK15),s(sK13,'const.min.@_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))))) | $spl6),
  inference(superposition,[],[f1931,f1712])).
fof(f1712,plain,(
  ( ! [X6,X4,X5] : (s(bool,'|dollar|exists_1'(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X6))))) = s(bool,'const.arithmetic.NRC_4'(s(fun(X4,fun(X4,bool)),X5),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X4,X6),s(X4,'const.min.@_1'(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X6)))))))) )),
  inference(superposition,[],[f301,f289])).
fof(f289,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X6,bool),X5))) = s(bool,happ(s(fun(X6,bool),X5),s(X6,'const.min.@_1'(s(fun(X6,bool),X5)))))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.bool.EXISTS_DEF')).
fof(f1931,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',X0),s(sK13,sK18),s(sK13,X1)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',X0))),s(sK13,sK15),s(sK13,X1))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1919,f247])).
fof(f247,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT',aHL_TRUTH)).
fof(f1919,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',X0),s(sK13,sK18),s(sK13,X1)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',X0))),s(sK13,sK15),s(sK13,X1))))) ) | $spl6),
  inference(superposition,[],[f304,f490])).
fof(f490,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | $spl6),
  inference(unit_resulting_resolution,[],[f438,f263])).
fof(f263,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.bool.EQ_CLAUSES|split|1')).
fof(f304,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X4),s(X0,X2),s(X0,X6)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num',X3),s(X0,X6),s(X0,X1)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s(X0,X2),s(X0,X1))))) )),
  inference(cnf_transformation,[],[f192])).
fof(f23352,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f23351,f247])).
fof(f23351,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f23350,f491])).
fof(f491,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f436,f263])).
fof(f23350,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f23344,f247])).
fof(f23344,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl2 | $spl6)),
  inference(superposition,[],[f434,f490])).
fof(f23309,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f23308,f439])).
fof(f439,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f250,f247])).
fof(f250,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f21])).
fof(f21,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.bool.NOT_CLAUSES|split|2')).
fof(f23308,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f23307,f524])).
fof(f524,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f433,f267])).
fof(f267,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.bool.EQ_CLAUSES|split|3')).
fof(f23307,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f23306,f254])).
fof(f23306,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(sK13,sK15),s(sK13,sK14)))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f23294,f247])).
fof(f23294,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(sK13,sK15),s(sK13,sK14)))) | ($spl4 | $spl6)),
  inference(superposition,[],[f1931,f7026])).
fof(f7026,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | $spl4),
  inference(subsumption_resolution,[],[f6962,f247])).
fof(f6962,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | $spl4),
  inference(superposition,[],[f1770,f491])).
fof(f1770,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X9),s(X8,X10))),s(X8,X11)))) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(X8,fun(X8,bool)),X9),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X8,X10),s(X8,X11)))) )),
  inference(superposition,[],[f263,f301])).
fof(f22901,plain,(
  ( ! [X685,X686] : (s(fun(bool,bool),X685) = s(fun(bool,bool),X686)) ) | $spl120),
  inference(cnf_transformation,[],[f22901_D])).
fof(f22901_D,plain,(
  ( ! [X685,X686] : (s(fun(bool,bool),X685) = s(fun(bool,bool),X686)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f22903,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl123),
  inference(cnf_transformation,[],[f22903_D])).
fof(f22903_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f22900,plain,(
  ( ! [X685,X686] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X685) = s(fun(bool,bool),X686)) ) | ($spl48 | $spl61)),
  inference(forward_demodulation,[],[f22839,f5088])).
fof(f5088,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f439,f5075,f769])).
fof(f769,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f765,f439])).
fof(f765,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f305,f255])).
fof(f255,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT',aHL_BOOL_CASES)).
fof(f305,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X14,X15] : ((p(s(bool,X14)) => p(s(bool,X15))) => ((p(s(bool,X15)) => p(s(bool,X14))) => s(bool,X14) = s(bool,X15)))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.bool.IMP_ANTISYM_AX')).
fof(f22839,plain,(
  ( ! [X685,X686] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X685) = s(fun(bool,bool),X686)) ) | $spl48),
  inference(duplicate_literal_removal,[],[f22714])).
fof(f22714,plain,(
  ( ! [X685,X686] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X685) = s(fun(bool,bool),X686) | s(fun(bool,bool),X685) = s(fun(bool,bool),X686)) ) | $spl48),
  inference(superposition,[],[f338,f4976])).
fof(f338,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK24(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK24(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK24(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK24(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24])],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT',aHL_EXT)).
fof(f21750,plain,(
  ( ! [X39,X38] : (sP1('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),X38) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X39)) ) | $spl118),
  inference(cnf_transformation,[],[f21750_D])).
fof(f21750_D,plain,(
  ( ! [X39,X38] : (sP1('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),X38) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X39)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f18140,plain,(
  ( ! [X30] : (p(s(bool,X30))) ) | $spl112),
  inference(cnf_transformation,[],[f18140_D])).
fof(f18140_D,plain,(
  ( ! [X30] : (p(s(bool,X30))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f21749,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | sP1('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),X38) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | $spl1),
  inference(subsumption_resolution,[],[f21726,f262])).
fof(f262,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f165])).
fof(f21726,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | sP1('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),X38) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X37)) ) | $spl1),
  inference(superposition,[],[f2703,f2047])).
fof(f2047,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(subsumption_resolution,[],[f2041,f439])).
fof(f2041,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(superposition,[],[f769,f524])).
fof(f2703,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f2332,f247])).
fof(f2332,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f333,f711])).
fof(f711,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f255,f255])).
fof(f333,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f200])).
fof(f200,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f21748,plain,(
  ( ! [X31,X32] : (sP1('F_0',X31) | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) | $spl116),
  inference(cnf_transformation,[],[f21748_D])).
fof(f21748_D,plain,(
  ( ! [X31,X32] : (sP1('F_0',X31) | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f21747,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | sP1('F_0',X31) | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f21723,f262])).
fof(f21723,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | sP1('F_0',X31) | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f2703,f255])).
fof(f21680,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18149,f18164])).
fof(f18164,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f18140,f308])).
fof(f308,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f125])).
fof(f18149,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f266])).
fof(f266,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f21679,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18155,f247])).
fof(f18155,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f487,f18140,f308])).
fof(f487,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f439,f262])).
fof(f21678,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18157,f247])).
fof(f18157,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f523,f18140,f308])).
fof(f523,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f433,f262])).
fof(f21677,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18161,f18140])).
fof(f18161,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f523,f18140,f308])).
fof(f21676,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18165,f247])).
fof(f18165,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f487,f18140,f308])).
fof(f21675,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18167,f247])).
fof(f18167,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f523,f18140,f308])).
fof(f21674,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18171,f18140])).
fof(f18171,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f523,f18140,f308])).
fof(f21673,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18178,f18140])).
fof(f18178,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f529,f18140,f329])).
fof(f329,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f201])).
fof(f529,plain,(
  sP1('T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f247,f433,f330])).
fof(f330,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f201])).
fof(f21672,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18179,f247])).
fof(f18179,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f507,f18140,f329])).
fof(f507,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f439,f247,f333])).
fof(f21671,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f18180,f18140])).
fof(f18180,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f5989,f18140,f329])).
fof(f5989,plain,(
  sP1('F_0',f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f439,f5394])).
fof(f5394,plain,(
  ( ! [X6] : (sP1(X6,f(s(bool,'T_0'))) | p(s(bool,X6))) ) | $spl61),
  inference(subsumption_resolution,[],[f5335,f247])).
fof(f5335,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1(X6,f(s(bool,'T_0')))) ) | $spl61),
  inference(superposition,[],[f333,f5088])).
fof(f21670,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f18181,f18140])).
fof(f18181,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f825,f18140,f329])).
fof(f825,plain,(
  sP1('F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | $spl6),
  inference(unit_resulting_resolution,[],[f439,f633])).
fof(f633,plain,(
  ( ! [X3] : (sP1(X3,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | p(s(bool,X3))) ) | $spl6),
  inference(subsumption_resolution,[],[f622,f247])).
fof(f622,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | sP1(X3,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) ) | $spl6),
  inference(superposition,[],[f333,f490])).
fof(f21669,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18182,f18140])).
fof(f18182,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10005,f18140,f329])).
fof(f10005,plain,(
  sP1('F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f7439])).
fof(f7439,plain,(
  ( ! [X0] : (sP1(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f7413,f247])).
fof(f7413,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) ) | $spl4),
  inference(superposition,[],[f1779,f491])).
fof(f1779,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,happ(s(fun(X50,bool),happ(s(fun(X50,fun(X50,bool)),X51),s(X50,X52))),s(X50,X53)))) | p(s(bool,X54)) | sP1(X54,'const.arithmetic.NRC_4'(s(fun(X50,fun(X50,bool)),X51),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X50,X52),s(X50,X53)))) )),
  inference(superposition,[],[f333,f301])).
fof(f21668,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18183,f18140])).
fof(f18183,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f4984,f18140,f329])).
fof(f4984,plain,(
  sP1('F_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f1090])).
fof(f1090,plain,(
  ( ! [X6] : (sP1(X6,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | p(s(bool,X6))) ) | $spl4),
  inference(subsumption_resolution,[],[f1066,f247])).
fof(f1066,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1(X6,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) ) | $spl4),
  inference(superposition,[],[f333,f491])).
fof(f21667,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18184,f18140])).
fof(f18184,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1323,f18140,f329])).
fof(f1323,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f1309,f333])).
fof(f1309,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | $spl4),
  inference(subsumption_resolution,[],[f1307,f247])).
fof(f1307,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | $spl4),
  inference(superposition,[],[f440,f491])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f290,f289])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X13,X8] : (p(s(bool,happ(s(fun(X6,bool),X13),s(X6,X8)))) => p(s(bool,happ(s(fun(X6,bool),X13),s(X6,'const.min.@_1'(s(fun(X6,bool),X13)))))))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.bool.SELECT_AX')).
fof(f21666,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18186,f247])).
fof(f18186,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f530,f18140,f329])).
fof(f530,plain,(
  sP1('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f247,f433,f333])).
fof(f21665,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18195,f247])).
fof(f18195,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f502,f18140,f329])).
fof(f502,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f439,f247,f330])).
fof(f21664,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18196,f247])).
fof(f18196,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f529,f18140,f329])).
fof(f21663,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f18203,f18140])).
fof(f18203,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f5903,f18140,f329])).
fof(f5903,plain,(
  sP1(f(s(bool,'T_0')),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f439,f5393])).
fof(f5393,plain,(
  ( ! [X5] : (sP1(f(s(bool,'T_0')),X5) | p(s(bool,X5))) ) | $spl61),
  inference(subsumption_resolution,[],[f5334,f247])).
fof(f5334,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP1(f(s(bool,'T_0')),X5)) ) | $spl61),
  inference(superposition,[],[f330,f5088])).
fof(f21662,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18204,f18140])).
fof(f18204,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f530,f18140,f329])).
fof(f21661,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f18205,f18140])).
fof(f18205,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f782,f18140,f329])).
fof(f782,plain,(
  sP1('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f439,f632])).
fof(f632,plain,(
  ( ! [X2] : (sP1('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X2) | p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f621,f247])).
fof(f621,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X2)) ) | $spl6),
  inference(superposition,[],[f330,f490])).
fof(f21660,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18206,f18140])).
fof(f18206,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f9965,f18140,f329])).
fof(f9965,plain,(
  sP1('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f439,f7358])).
fof(f7358,plain,(
  ( ! [X0] : (sP1('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X0) | p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f7332,f247])).
fof(f7332,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X0)) ) | $spl4),
  inference(superposition,[],[f1778,f491])).
fof(f1778,plain,(
  ( ! [X47,X45,X48,X46,X49] : (~p(s(bool,happ(s(fun(X45,bool),happ(s(fun(X45,fun(X45,bool)),X46),s(X45,X47))),s(X45,X48)))) | p(s(bool,X49)) | sP1('const.arithmetic.NRC_4'(s(fun(X45,fun(X45,bool)),X46),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X45,X47),s(X45,X48)),X49)) )),
  inference(superposition,[],[f330,f301])).
fof(f21659,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18207,f18140])).
fof(f18207,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1181,f18140,f329])).
fof(f1181,plain,(
  sP1(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f439,f1089])).
fof(f1089,plain,(
  ( ! [X5] : (sP1(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X5) | p(s(bool,X5))) ) | $spl4),
  inference(subsumption_resolution,[],[f1065,f247])).
fof(f1065,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP1(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X5)) ) | $spl4),
  inference(superposition,[],[f330,f491])).
fof(f21658,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18208,f18140])).
fof(f18208,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1322,f18140,f329])).
fof(f1322,plain,(
  sP1('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f439,f1309,f330])).
fof(f21657,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18234,f18140])).
fof(f18234,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f562,f18140,f342])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f205])).
fof(f205,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f204])).
fof(f204,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f562,plain,(
  ( ! [X0] : (~sP2('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f439,f439,f340])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f206])).
fof(f21656,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18249,f18140])).
fof(f18249,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f557,f18140,f343])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f206])).
fof(f557,plain,(
  ( ! [X0] : (~sP2(X0,'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f439,f439,f339])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f206])).
fof(f21655,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18256,f18140])).
fof(f18256,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f562,f18140,f344])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f206])).
fof(f21654,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18261,f18140])).
fof(f18261,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f557,f18140,f344])).
fof(f21653,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18268,f18140])).
fof(f18268,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1970,f439,f18140,f352])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP3(X2,X1,X0) & sP4(X1,X2,X0) & sP5(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP3(X2,X1,X0) | ~sP4(X1,X2,X0) | ~sP5(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f210])).
fof(f210,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP6(X2,X1,X0)))),
  inference(flattening,[],[f209])).
fof(f209,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP6(X2,X1,X0)))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X2,X1,X0] : (sP6(X2,X1,X0) <=> (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f1970,plain,(
  sP6('T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f247,f524,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP6(X2,X1,X0)) & (~sP6(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP6(X2,X1,X0))),
  inference(definition_folding,[],[f105,f149,f148,f147,f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : (sP3(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f147,plain,(
  ! [X1,X0,X2] : (sP4(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f148,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f105,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X22)) | p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X22)) | ~p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.sat.sat_2|49|_prove_ltb')).
fof(f21652,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18269,f247])).
fof(f18269,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f786,f439,f18140,f352])).
fof(f786,plain,(
  sP6('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f439,f487,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f220])).
fof(f21651,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f18270,f18140])).
fof(f18270,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f5085,f439,f18140,f352])).
fof(f5085,plain,(
  sP6('F_0',f(s(bool,'T_0')),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f439,f5075,f368])).
fof(f21650,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f18271,f18140])).
fof(f18271,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f1287,f439,f18140,f352])).
fof(f1287,plain,(
  sP6('F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f468,f448,f445,f439,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X0)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f211])).
fof(f445,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f439,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f213])).
fof(f213,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f148])).
fof(f448,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f439,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f216])).
fof(f216,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(nnf_transformation,[],[f147])).
fof(f468,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f438,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(flattening,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(nnf_transformation,[],[f146])).
fof(f21649,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18272,f18140])).
fof(f18272,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1288,f439,f18140,f352])).
fof(f1288,plain,(
  sP6('F_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f480,f448,f445,f439,f355])).
fof(f480,plain,(
  ( ! [X0,X1] : (sP3(X0,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f436,f366])).
fof(f21648,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18273,f18140])).
fof(f18273,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1847,f439,f18140,f352])).
fof(f1847,plain,(
  sP6('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f439,f446,f448,f1384,f355])).
fof(f1384,plain,(
  ( ! [X0,X1] : (sP3(X0,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1309,f366])).
fof(f446,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f439,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f21647,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18446,f247])).
fof(f18446,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f1227,f439,f18140,f352])).
fof(f1227,plain,(
  sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f450,f447,f445,f439,f354])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X1)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f211])).
fof(f447,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f217])).
fof(f450,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f219])).
fof(f21646,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18447,f247])).
fof(f18447,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1970,f439,f18140,f352])).
fof(f21645,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f18453,f18140])).
fof(f18453,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f1228,f439,f18140,f352])).
fof(f1228,plain,(
  sP6('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),'F_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f467,f466,f445,f439,f354])).
fof(f466,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f438,f361])).
fof(f467,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f438,f365])).
fof(f21644,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18454,f18140])).
fof(f18454,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1229,f439,f18140,f352])).
fof(f1229,plain,(
  sP6(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f479,f478,f445,f439,f354])).
fof(f478,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f436,f361])).
fof(f479,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f436,f365])).
fof(f21643,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18455,f18140])).
fof(f18455,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1812,f439,f18140,f352])).
fof(f1812,plain,(
  sP6('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f439,f445,f449,f1383,f354])).
fof(f1383,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1309,f365])).
fof(f449,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f439,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f21642,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18484,f18998])).
fof(f18998,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f487,f18140,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f220])).
fof(f18484,plain,(
  ( ! [X0] : (sP6(X0,'T_0','F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f451,f448,f444,f18140,f353])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2) | ~sP5(X0,X2,X1) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f211])).
fof(f444,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f214])).
fof(f451,plain,(
  ( ! [X0,X1] : (sP3(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f247,f366])).
fof(f21641,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18489,f19000])).
fof(f19000,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f523,f18140,f370])).
fof(f18489,plain,(
  ( ! [X0] : (sP6(X0,'T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f533,f444,f18140,f353])).
fof(f533,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f433,f362])).
fof(f21640,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18601,f18998])).
fof(f18601,plain,(
  ( ! [X0] : (sP6(X0,'T_0','F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f451,f448,f445,f18140,f353])).
fof(f21639,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18606,f806])).
fof(f806,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f247,f487,f370])).
fof(f18606,plain,(
  sP6('T_0','T_0','F_0') | $spl112),
  inference(unit_resulting_resolution,[],[f451,f447,f445,f18140,f353])).
fof(f21638,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18609,f1342])).
fof(f1342,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f1309,f352])).
fof(f18609,plain,(
  sP6('T_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0') | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1384,f447,f445,f18140,f353])).
fof(f21637,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f18611,f18998])).
fof(f18611,plain,(
  sP6('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),'T_0','F_0') | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f466,f445,f18140,f353])).
fof(f21636,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18616,f18998])).
fof(f18616,plain,(
  sP6(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),'T_0','F_0') | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f478,f445,f18140,f353])).
fof(f21635,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18621,f1390])).
fof(f1390,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f487,f1309,f370])).
fof(f18621,plain,(
  sP6('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'T_0','F_0') | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f1382,f445,f18140,f353])).
fof(f1382,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1309,f361])).
fof(f21634,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18624,f1343])).
fof(f1343,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1309,f439,f1309,f352])).
fof(f18624,plain,(
  sP6('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0') | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1384,f1382,f445,f18140,f353])).
fof(f21633,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18637,f19000])).
fof(f18637,plain,(
  ( ! [X0] : (sP6(X0,'T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f533,f531,f18140,f353])).
fof(f531,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f433,f358])).
fof(f21632,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18642,f1951])).
fof(f1951,plain,(
  ~sP6('T_0','T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f247,f523,f370])).
fof(f18642,plain,(
  sP6('T_0','T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f447,f531,f18140,f353])).
fof(f21631,plain,(
  $false | ($spl1 | $spl6 | $spl112)),
  inference(subsumption_resolution,[],[f18647,f19000])).
fof(f18647,plain,(
  sP6('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),'T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f466,f531,f18140,f353])).
fof(f21630,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18652,f19000])).
fof(f18652,plain,(
  sP6(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),'T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f478,f531,f18140,f353])).
fof(f21629,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18657,f19000])).
fof(f18657,plain,(
  sP6('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f1382,f531,f18140,f353])).
fof(f21628,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18689,f18998])).
fof(f18689,plain,(
  sP6('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'T_0','F_0') | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f448,f532,f18140,f353])).
fof(f532,plain,(
  ( ! [X0,X1] : (sP5('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f433,f359])).
fof(f21627,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18693,f19000])).
fof(f18693,plain,(
  sP6('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f533,f532,f18140,f353])).
fof(f21626,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18705,f18811])).
fof(f18811,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f357])).
fof(f18705,plain,(
  ~sP5('T_0','F_0','T_0') | $spl112),
  inference(unit_resulting_resolution,[],[f451,f448,f806,f18140,f353])).
fof(f21625,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18706,f18857])).
fof(f18857,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f361])).
fof(f18706,plain,(
  ~sP4('T_0','F_0','T_0') | $spl112),
  inference(unit_resulting_resolution,[],[f451,f445,f806,f18140,f353])).
fof(f21624,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18707,f18863])).
fof(f18863,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f367])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f219])).
fof(f18707,plain,(
  ~sP3('F_0','T_0','T_0') | $spl112),
  inference(unit_resulting_resolution,[],[f448,f445,f806,f18140,f353])).
fof(f21623,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18709,f18811])).
fof(f18709,plain,(
  ~sP5('T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'T_0') | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f533,f1951,f18140,f353])).
fof(f21622,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18710,f18857])).
fof(f18710,plain,(
  ~sP4('T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'T_0') | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f531,f1951,f18140,f353])).
fof(f21621,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18711,f18863])).
fof(f18711,plain,(
  ~sP3('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'T_0','T_0') | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f533,f531,f1951,f18140,f353])).
fof(f21620,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f18719,f18811])).
fof(f18719,plain,(
  ~sP5('T_0','F_0',f(s(bool,'T_0'))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f450,f448,f5084,f18140,f353])).
fof(f5084,plain,(
  ~sP6('T_0',f(s(bool,'T_0')),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f247,f5075,f370])).
fof(f21619,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f18720,f18857])).
fof(f18720,plain,(
  ~sP4(f(s(bool,'T_0')),'F_0','T_0') | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f450,f445,f5084,f18140,f353])).
fof(f21618,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f18721,f18863])).
fof(f18721,plain,(
  ~sP3('F_0',f(s(bool,'T_0')),'T_0') | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f448,f445,f5084,f18140,f353])).
fof(f21617,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18724,f18811])).
fof(f18724,plain,(
  ~sP5('T_0','F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1384,f448,f1342,f18140,f353])).
fof(f21616,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18725,f18857])).
fof(f18725,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0','T_0') | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1384,f445,f1342,f18140,f353])).
fof(f21615,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18726,f18863])).
fof(f18726,plain,(
  ~sP3('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'T_0') | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f448,f445,f1342,f18140,f353])).
fof(f21614,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18749,f18811])).
fof(f18749,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0','T_0') | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f448,f1390,f18140,f353])).
fof(f21613,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18750,f18857])).
fof(f18750,plain,(
  ~sP4('T_0','F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f445,f1390,f18140,f353])).
fof(f21612,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18751,f18863])).
fof(f18751,plain,(
  ~sP3('F_0','T_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f448,f445,f1390,f18140,f353])).
fof(f21611,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18755,f18811])).
fof(f18755,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1384,f448,f1343,f18140,f353])).
fof(f21610,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18756,f18857])).
fof(f18756,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1384,f445,f1343,f18140,f353])).
fof(f21609,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18757,f18863])).
fof(f18757,plain,(
  ~sP3('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f448,f445,f1343,f18140,f353])).
fof(f21608,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18760,f18140])).
fof(f18760,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f531,f439,f18140,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f21607,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18761,f18140])).
fof(f18761,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f446,f439,f18140,f356])).
fof(f21606,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18762,f18140])).
fof(f18762,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f532,f439,f18140,f356])).
fof(f21605,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18783,f18811])).
fof(f18783,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f18140,f356])).
fof(f21604,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18785,f18140])).
fof(f18785,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f445,f439,f18140,f356])).
fof(f21603,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18786,f18140])).
fof(f18786,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f531,f439,f18140,f356])).
fof(f21602,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18788,f18140])).
fof(f18788,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f532,f439,f18140,f356])).
fof(f21601,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18797,f18811])).
fof(f18797,plain,(
  ( ! [X0] : (~sP5('T_0',X0,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f247,f18140,f356])).
fof(f21600,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18809,f18811])).
fof(f18809,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f18140,f356])).
fof(f21599,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18814,f18140])).
fof(f18814,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f533,f439,f18140,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f217])).
fof(f21598,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18815,f18140])).
fof(f18815,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f449,f439,f18140,f360])).
fof(f21597,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18816,f18140])).
fof(f18816,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f534,f439,f18140,f360])).
fof(f534,plain,(
  ( ! [X0,X1] : (sP4('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f433,f363])).
fof(f21596,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18833,f18857])).
fof(f18833,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f18140,f360])).
fof(f21595,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18835,f18140])).
fof(f18835,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f448,f439,f18140,f360])).
fof(f21594,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18836,f18140])).
fof(f18836,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f533,f439,f18140,f360])).
fof(f21593,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18838,f18140])).
fof(f18838,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f534,f439,f18140,f360])).
fof(f21592,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18845,f18857])).
fof(f18845,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f247,f18140,f360])).
fof(f21591,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18855,f18857])).
fof(f18855,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f18140,f360])).
fof(f21588,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18870,f20957])).
fof(f20957,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f18140,f1775])).
fof(f1775,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,X34)) | ~p(s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,X32))),s(X30,X33)))) | s(bool,X34) = s(bool,'const.arithmetic.NRC_4'(s(fun(X30,fun(X30,bool)),X31),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X30,X32),s(X30,X33)))) )),
  inference(superposition,[],[f308,f301])).
fof(f18870,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13472,f18140,f369])).
fof(f13472,plain,(
  ~sP6('T_0','F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f13458,f351])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f211])).
fof(f13458,plain,(
  ~sP5('T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f9497])).
fof(f9497,plain,(
  ( ! [X24,X25] : (~sP5(X24,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X25) | p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl4),
  inference(subsumption_resolution,[],[f9376,f247])).
fof(f9376,plain,(
  ( ! [X24,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | p(s(bool,X25)) | ~sP5(X24,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X25)) ) | $spl4),
  inference(superposition,[],[f356,f7026])).
fof(f21587,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18871,f20707])).
fof(f20707,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f18140,f1744])).
fof(f1744,plain,(
  ( ! [X61,X59,X57,X60,X58] : (~p(s(bool,X61)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X57,fun(X57,bool)),X58),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X57,X59),s(X57,X60)))) | s(bool,X61) = s(bool,happ(s(fun(X57,bool),happ(s(fun(X57,fun(X57,bool)),X58),s(X57,X59))),s(X57,X60)))) )),
  inference(superposition,[],[f308,f301])).
fof(f18871,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10059,f18140,f369])).
fof(f10059,plain,(
  ~sP6('T_0','F_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f10045,f351])).
fof(f10045,plain,(
  ~sP5('T_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f1092])).
fof(f1092,plain,(
  ( ! [X14,X13] : (~sP5(X13,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X14) | p(s(bool,X14)) | ~p(s(bool,X13))) ) | $spl4),
  inference(subsumption_resolution,[],[f1070,f247])).
fof(f1070,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | p(s(bool,X14)) | ~sP5(X13,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X14)) ) | $spl4),
  inference(superposition,[],[f356,f491])).
fof(f21583,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f18877,f5088])).
fof(f18877,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f5313,f18140,f369])).
fof(f5313,plain,(
  ~sP6('F_0','T_0',f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f439,f5088,f371])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f220])).
fof(f21582,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f18878,f490])).
fof(f18878,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f815,f18140,f369])).
fof(f815,plain,(
  ~sP6('F_0','T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | $spl6),
  inference(unit_resulting_resolution,[],[f439,f490,f371])).
fof(f21581,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18879,f7026])).
fof(f18879,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f9338,f18140,f369])).
fof(f9338,plain,(
  ~sP6('F_0','T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f7026,f371])).
fof(f21580,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18880,f491])).
fof(f18880,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1052,f18140,f369])).
fof(f1052,plain,(
  ~sP6('F_0','T_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f491,f371])).
fof(f21579,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18881,f1314])).
fof(f1314,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f247,f1309,f308])).
fof(f18881,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1706,f18140,f369])).
fof(f1706,plain,(
  ~sP6('F_0','T_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1375,f350])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f211])).
fof(f1375,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f1309,f360])).
fof(f21578,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18883,f524])).
fof(f18883,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1971,f18140,f369])).
fof(f1971,plain,(
  ~sP6('F_0','F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f439,f524,f371])).
fof(f21577,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f18884,f5088])).
fof(f18884,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f5311,f18140,f369])).
fof(f5311,plain,(
  ~sP6('F_0',f(s(bool,'T_0')),'T_0') | $spl61),
  inference(unit_resulting_resolution,[],[f439,f5088,f371])).
fof(f21576,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f18885,f490])).
fof(f18885,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f813,f18140,f369])).
fof(f813,plain,(
  ~sP6('F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),'T_0') | $spl6),
  inference(unit_resulting_resolution,[],[f439,f490,f371])).
fof(f21575,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18886,f524])).
fof(f18886,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1969,f18140,f369])).
fof(f1969,plain,(
  ~sP6('F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f439,f524,f371])).
fof(f21574,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18887,f7026])).
fof(f18887,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f9336,f18140,f369])).
fof(f9336,plain,(
  ~sP6('F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f439,f7026,f371])).
fof(f21573,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18888,f301])).
fof(f18888,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) != s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3)))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f1708,f18140,f369])).
fof(f1708,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('F_0','const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3)),happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) )),
  inference(unit_resulting_resolution,[],[f439,f301,f371])).
fof(f21572,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18889,f491])).
fof(f18889,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1050,f18140,f369])).
fof(f1050,plain,(
  ~sP6('F_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f439,f491,f371])).
fof(f21571,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18890,f289])).
fof(f18890,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) != s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f1597,f18140,f369])).
fof(f1597,plain,(
  ( ! [X0,X1] : (~sP6('F_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f439,f289,f371])).
fof(f21570,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18891,f301])).
fof(f18891,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) != s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3)))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f1710,f18140,f369])).
fof(f1710,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('F_0',happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)),'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3)))) )),
  inference(unit_resulting_resolution,[],[f439,f301,f371])).
fof(f21569,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f18892,f1314])).
fof(f18892,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1701,f18140,f369])).
fof(f1701,plain,(
  ~sP6('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1368,f350])).
fof(f1368,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f1309,f360])).
fof(f21568,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18893,f289])).
fof(f18893,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) != s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f1599,f18140,f369])).
fof(f1599,plain,(
  ( ! [X0,X1] : (~sP6('F_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f439,f289,f371])).
fof(f21564,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18953,f487])).
fof(f18953,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl112),
  inference(unit_resulting_resolution,[],[f786,f18140,f370])).
fof(f21563,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18954,f523])).
fof(f18954,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1290,f18140,f370])).
fof(f1290,plain,(
  sP6('F_0','T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f451,f533,f531,f439,f355])).
fof(f21562,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f18955,f487])).
fof(f18955,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl112),
  inference(unit_resulting_resolution,[],[f785,f18140,f370])).
fof(f785,plain,(
  sP6('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f439,f487,f368])).
fof(f21561,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f21560,f487])).
fof(f21560,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl61 | $spl112)),
  inference(forward_demodulation,[],[f18956,f5088])).
fof(f18956,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f5090,f18140,f370])).
fof(f5090,plain,(
  sP6('F_0','F_0',f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f439,f5075,f368])).
fof(f21559,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f21558,f487])).
fof(f21558,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl112)),
  inference(forward_demodulation,[],[f18957,f490])).
fof(f18957,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f1241,f18140,f370])).
fof(f1241,plain,(
  sP6('F_0','F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | $spl6),
  inference(unit_resulting_resolution,[],[f469,f449,f446,f439,f354])).
fof(f469,plain,(
  ( ! [X0,X1] : (sP3('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f438,f367])).
fof(f21557,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f21556,f487])).
fof(f21556,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl112)),
  inference(forward_demodulation,[],[f18959,f1314])).
fof(f18959,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1863,f18140,f370])).
fof(f1863,plain,(
  sP6('F_0','F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f446,f449,f1385,f354])).
fof(f1385,plain,(
  ( ! [X0,X1] : (sP3('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1309,f367])).
fof(f21555,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f21554,f487])).
fof(f21554,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl61 | $spl112)),
  inference(forward_demodulation,[],[f18960,f5088])).
fof(f18960,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f5085,f18140,f370])).
fof(f21553,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18961,f523])).
fof(f18961,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1195,f18140,f370])).
fof(f1195,plain,(
  sP6('F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f452,f534,f446,f247,f353])).
fof(f452,plain,(
  ( ! [X0,X1] : (sP3('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f247,f367])).
fof(f21552,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f21551,f487])).
fof(f21551,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl112)),
  inference(forward_demodulation,[],[f18962,f490])).
fof(f18962,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f1287,f18140,f370])).
fof(f21550,plain,(
  $false | ($spl1 | $spl6 | $spl112)),
  inference(subsumption_resolution,[],[f21549,f487])).
fof(f21549,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl6 | $spl112)),
  inference(forward_demodulation,[],[f21548,f490])).
fof(f21548,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | ($spl1 | $spl6 | $spl112)),
  inference(forward_demodulation,[],[f18963,f524])).
fof(f18963,plain,(
  s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f1267,f18140,f370])).
fof(f1267,plain,(
  sP6('F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | ($spl1 | $spl6)),
  inference(unit_resulting_resolution,[],[f469,f534,f446,f439,f355])).
fof(f21547,plain,(
  $false | ($spl1 | $spl6 | $spl112)),
  inference(subsumption_resolution,[],[f21546,f487])).
fof(f21546,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl6 | $spl112)),
  inference(forward_demodulation,[],[f21545,f490])).
fof(f21545,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | ($spl1 | $spl6 | $spl112)),
  inference(forward_demodulation,[],[f18964,f524])).
fof(f18964,plain,(
  s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f1291,f18140,f370])).
fof(f1291,plain,(
  sP6('F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl6)),
  inference(unit_resulting_resolution,[],[f468,f533,f531,f439,f355])).
fof(f21544,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f21543,f487])).
fof(f21543,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl4 | $spl112)),
  inference(forward_demodulation,[],[f21542,f1314])).
fof(f21542,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl1 | $spl4 | $spl112)),
  inference(forward_demodulation,[],[f18966,f524])).
fof(f18966,plain,(
  s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1862,f18140,f370])).
fof(f1862,plain,(
  sP6('F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f439,f446,f534,f1385,f355])).
fof(f21541,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f21540,f487])).
fof(f21540,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl112)),
  inference(forward_demodulation,[],[f18969,f1314])).
fof(f18969,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1847,f18140,f370])).
fof(f21539,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f21538,f487])).
fof(f21538,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl4 | $spl112)),
  inference(forward_demodulation,[],[f21537,f1314])).
fof(f21537,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl1 | $spl4 | $spl112)),
  inference(forward_demodulation,[],[f18970,f524])).
fof(f18970,plain,(
  s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1850,f18140,f370])).
fof(f1850,plain,(
  sP6('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f439,f446,f533,f1384,f355])).
fof(f21536,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18978,f487])).
fof(f18978,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1196,f18140,f370])).
fof(f1196,plain,(
  sP6('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f452,f449,f532,f247,f353])).
fof(f21535,plain,(
  $false | ($spl1 | $spl6 | $spl112)),
  inference(subsumption_resolution,[],[f21534,f487])).
fof(f21534,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl6 | $spl112)),
  inference(forward_demodulation,[],[f18979,f490])).
fof(f18979,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | ($spl1 | $spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f1245,f18140,f370])).
fof(f1245,plain,(
  sP6('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | ($spl1 | $spl6)),
  inference(unit_resulting_resolution,[],[f469,f449,f532,f439,f354])).
fof(f21533,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f21532,f487])).
fof(f21532,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl4 | $spl112)),
  inference(forward_demodulation,[],[f18981,f1314])).
fof(f18981,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1864,f18140,f370])).
fof(f1864,plain,(
  sP6('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f439,f532,f449,f1385,f354])).
fof(f21531,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f18982,f523])).
fof(f18982,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1197,f18140,f370])).
fof(f1197,plain,(
  sP6('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f452,f534,f532,f247,f353])).
fof(f21530,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19023,f19308])).
fof(f19308,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f18140,f18140,f385])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f107,f152,f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f152,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) & p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.sat.sat_23_prove_ltb')).
fof(f19023,plain,(
  ( ! [X0,X1] : (~sP8(X0,X1,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (sP7(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f222])).
fof(f222,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f152])).
fof(f21529,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19042,f19308])).
fof(f19042,plain,(
  ( ! [X0,X1] : (~sP8(X0,'F_0',X1)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f223])).
fof(f21528,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19045,f18140])).
fof(f19045,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f455,f581,f18140,f375])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f223])).
fof(f581,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f247,f439,f374])).
fof(f455,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f439,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f151])).
fof(f21527,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19046,f18140])).
fof(f19046,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f453,f581,f18140,f375])).
fof(f453,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f380])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f21526,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19048,f18140])).
fof(f19048,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f535,f581,f18140,f375])).
fof(f535,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f433,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f225])).
fof(f21525,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19050,f18140])).
fof(f19050,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f455,f1393,f18140,f375])).
fof(f1393,plain,(
  ( ! [X0] : (~sP8('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f1309,f374])).
fof(f21524,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19051,f18140])).
fof(f19051,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1422,f1393,f18140,f375])).
fof(f1422,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1309,f380])).
fof(f21523,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19053,f18140])).
fof(f19053,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f535,f1393,f18140,f375])).
fof(f21522,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19060,f19184])).
fof(f19184,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f380])).
fof(f19060,plain,(
  ~sP7('F_0','T_0','T_0') | $spl112),
  inference(unit_resulting_resolution,[],[f581,f247,f18140,f375])).
fof(f21521,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19064,f19184])).
fof(f19064,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1393,f247,f18140,f375])).
fof(f21520,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19074,f19184])).
fof(f19074,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f578,f18140,f18140,f375])).
fof(f578,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f439,f373])).
fof(f21519,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19076,f19184])).
fof(f19076,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f581,f18140,f18140,f375])).
fof(f21518,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19086,f19184])).
fof(f19086,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1392,f18140,f18140,f375])).
fof(f1392,plain,(
  ( ! [X0] : (~sP8('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f1309,f373])).
fof(f21517,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19088,f19184])).
fof(f19088,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1393,f18140,f18140,f375])).
fof(f21516,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19090,f18140])).
fof(f19090,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f454,f578,f18140,f375])).
fof(f454,plain,(
  ( ! [X0,X1] : (sP7(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f439,f381])).
fof(f21515,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19091,f18140])).
fof(f19091,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f453,f578,f18140,f375])).
fof(f21514,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19093,f18140])).
fof(f19093,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f536,f578,f18140,f375])).
fof(f536,plain,(
  ( ! [X0,X1] : (sP7('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f433,f382])).
fof(f21513,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19095,f18140])).
fof(f19095,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f454,f1392,f18140,f375])).
fof(f21512,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19096,f18140])).
fof(f19096,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1422,f1392,f18140,f375])).
fof(f21511,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19098,f18140])).
fof(f19098,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f536,f1392,f18140,f375])).
fof(f21510,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19106,f19184])).
fof(f19106,plain,(
  ~sP7('T_0','F_0','T_0') | $spl112),
  inference(unit_resulting_resolution,[],[f578,f247,f18140,f375])).
fof(f21509,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19114,f19184])).
fof(f19114,plain,(
  ~sP7('T_0','F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1392,f247,f18140,f375])).
fof(f21508,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19124,f19184])).
fof(f19124,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f578,f18140,f18140,f375])).
fof(f21507,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19126,f19184])).
fof(f19126,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f581,f18140,f18140,f375])).
fof(f21506,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19136,f19184])).
fof(f19136,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1392,f18140,f18140,f375])).
fof(f21505,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19138,f19184])).
fof(f19138,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1393,f18140,f18140,f375])).
fof(f21504,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19141,f18140])).
fof(f19141,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f535,f439,f18140,f379])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f225])).
fof(f21503,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19142,f18140])).
fof(f19142,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f455,f439,f18140,f379])).
fof(f21502,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19143,f18140])).
fof(f19143,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f536,f439,f18140,f379])).
fof(f21501,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19160,f19184])).
fof(f19160,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f18140,f379])).
fof(f21500,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19162,f18140])).
fof(f19162,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f454,f439,f18140,f379])).
fof(f21499,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19163,f18140])).
fof(f19163,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f535,f439,f18140,f379])).
fof(f21498,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19165,f18140])).
fof(f19165,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f536,f439,f18140,f379])).
fof(f21497,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19172,f19184])).
fof(f19172,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f247,f18140,f379])).
fof(f21496,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19182,f19184])).
fof(f19182,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f18140,f379])).
fof(f21495,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19187,f18140])).
fof(f19187,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f581,f247,f18140,f385])).
fof(f21494,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19189,f18140])).
fof(f19189,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f581,f18140,f18140,f385])).
fof(f21493,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19190,f247])).
fof(f19190,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f920,f18140,f18140,f385])).
fof(f920,plain,(
  ~sP8('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f908,f372])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f223])).
fof(f908,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f439,f247,f247,f379])).
fof(f21492,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19191,f247])).
fof(f19191,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1437,f18140,f18140,f385])).
fof(f1437,plain,(
  ~sP8('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f1309,f388])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP8(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f227])).
fof(f21491,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f19192,f18140])).
fof(f19192,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f11345,f18140,f18140,f385])).
fof(f11345,plain,(
  ~sP8('F_0','T_0',f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f11323,f372])).
fof(f11323,plain,(
  ~sP7('T_0',f(s(bool,'T_0')),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f247,f439,f5400])).
fof(f5400,plain,(
  ( ! [X19,X20] : (~sP7(X19,f(s(bool,'T_0')),X20) | p(s(bool,X20)) | ~p(s(bool,X19))) ) | $spl61),
  inference(subsumption_resolution,[],[f5342,f247])).
fof(f5342,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | p(s(bool,X20)) | ~sP7(X19,f(s(bool,'T_0')),X20)) ) | $spl61),
  inference(superposition,[],[f379,f5088])).
fof(f21490,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f19193,f18140])).
fof(f19193,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f7544,f18140,f18140,f385])).
fof(f7544,plain,(
  ~sP8('F_0','T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | $spl6),
  inference(unit_resulting_resolution,[],[f7530,f372])).
fof(f7530,plain,(
  ~sP7('T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f247,f439,f916])).
fof(f916,plain,(
  ( ! [X0,X1] : (~sP7(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f909,f247])).
fof(f909,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP7(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X1)) ) | $spl6),
  inference(superposition,[],[f379,f490])).
fof(f21489,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19194,f18140])).
fof(f19194,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13810,f18140,f18140,f385])).
fof(f13810,plain,(
  ~sP8('F_0','T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f13796,f372])).
fof(f13796,plain,(
  ~sP7('T_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f9499])).
fof(f9499,plain,(
  ( ! [X28,X29] : (~sP7(X28,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X29) | p(s(bool,X29)) | ~p(s(bool,X28))) ) | $spl4),
  inference(subsumption_resolution,[],[f9378,f247])).
fof(f9378,plain,(
  ( ! [X28,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X28)) | p(s(bool,X29)) | ~sP7(X28,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X29)) ) | $spl4),
  inference(superposition,[],[f379,f7026])).
fof(f21488,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19195,f18140])).
fof(f19195,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10163,f18140,f18140,f385])).
fof(f10163,plain,(
  ~sP8('F_0','T_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f10149,f372])).
fof(f10149,plain,(
  ~sP7('T_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f1094])).
fof(f1094,plain,(
  ( ! [X17,X18] : (~sP7(X17,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X18) | p(s(bool,X18)) | ~p(s(bool,X17))) ) | $spl4),
  inference(subsumption_resolution,[],[f1072,f247])).
fof(f1072,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~sP7(X17,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X18)) ) | $spl4),
  inference(superposition,[],[f379,f491])).
fof(f21487,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19196,f18140])).
fof(f19196,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1441,f18140,f18140,f385])).
fof(f1441,plain,(
  ~sP8('F_0','T_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f1309,f388])).
fof(f21486,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19197,f18140])).
fof(f19197,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1442,f18140,f18140,f385])).
fof(f1442,plain,(
  ~sP8('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1309,f439,f1309,f388])).
fof(f21485,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19199,f18140])).
fof(f19199,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1393,f18140,f18140,f385])).
fof(f21484,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19200,f247])).
fof(f19200,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f581,f247,f18140,f385])).
fof(f21483,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19203,f18140])).
fof(f19203,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1393,f247,f18140,f385])).
fof(f21482,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19208,f247])).
fof(f19208,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f578,f18140,f18140,f385])).
fof(f21481,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19209,f247])).
fof(f19209,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f581,f18140,f18140,f385])).
fof(f21480,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19218,f18140])).
fof(f19218,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1392,f18140,f18140,f385])).
fof(f21479,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19219,f18140])).
fof(f19219,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1393,f18140,f18140,f385])).
fof(f21478,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19224,f18140])).
fof(f19224,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f578,f247,f18140,f385])).
fof(f21477,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19226,f18140])).
fof(f19226,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f578,f18140,f18140,f385])).
fof(f21476,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19228,f247])).
fof(f19228,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f920,f18140,f18140,f385])).
fof(f21475,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19229,f18140])).
fof(f19229,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1437,f18140,f18140,f385])).
fof(f21474,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f19230,f247])).
fof(f19230,plain,(
  ~p(s(bool,'T_0')) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f11345,f18140,f18140,f385])).
fof(f21473,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f19231,f247])).
fof(f19231,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f7544,f18140,f18140,f385])).
fof(f21472,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19232,f247])).
fof(f19232,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13810,f18140,f18140,f385])).
fof(f21471,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19233,f247])).
fof(f19233,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10163,f18140,f18140,f385])).
fof(f21470,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19234,f247])).
fof(f19234,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1441,f18140,f18140,f385])).
fof(f21469,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19235,f18140])).
fof(f19235,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1442,f18140,f18140,f385])).
fof(f21468,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19236,f18140])).
fof(f19236,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1392,f18140,f18140,f385])).
fof(f21467,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19238,f247])).
fof(f19238,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f578,f247,f18140,f385])).
fof(f21466,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19245,f18140])).
fof(f19245,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1392,f247,f18140,f385])).
fof(f21465,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19250,f247])).
fof(f19250,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f578,f18140,f18140,f385])).
fof(f21464,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19251,f247])).
fof(f19251,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f581,f18140,f18140,f385])).
fof(f21463,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19260,f18140])).
fof(f19260,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1392,f18140,f18140,f385])).
fof(f21462,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19261,f18140])).
fof(f19261,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1393,f18140,f18140,f385])).
fof(f21461,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19267,f247])).
fof(f19267,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f920,f247,f18140,f385])).
fof(f21460,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19268,f18140])).
fof(f19268,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1437,f247,f18140,f385])).
fof(f21459,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19270,f18140])).
fof(f19270,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f578,f18140,f18140,f385])).
fof(f21458,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19272,f247])).
fof(f19272,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f920,f18140,f18140,f385])).
fof(f21457,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19273,f18140])).
fof(f19273,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1437,f18140,f18140,f385])).
fof(f21456,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f19274,f247])).
fof(f19274,plain,(
  ~p(s(bool,'T_0')) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f11345,f18140,f18140,f385])).
fof(f21455,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f19275,f247])).
fof(f19275,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f7544,f18140,f18140,f385])).
fof(f21454,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19276,f247])).
fof(f19276,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13810,f18140,f18140,f385])).
fof(f21453,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19277,f247])).
fof(f19277,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10163,f18140,f18140,f385])).
fof(f21452,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19278,f247])).
fof(f19278,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1441,f18140,f18140,f385])).
fof(f21451,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19279,f18140])).
fof(f19279,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1442,f18140,f18140,f385])).
fof(f21450,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19280,f18140])).
fof(f19280,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1392,f18140,f18140,f385])).
fof(f21449,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19283,f247])).
fof(f19283,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f920,f247,f18140,f385])).
fof(f21448,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f19284,f18140])).
fof(f19284,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f11345,f247,f18140,f385])).
fof(f21447,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f19285,f18140])).
fof(f19285,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f7544,f247,f18140,f385])).
fof(f21446,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19286,f18140])).
fof(f19286,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13810,f247,f18140,f385])).
fof(f21445,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19287,f18140])).
fof(f19287,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10163,f247,f18140,f385])).
fof(f21444,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19288,f18140])).
fof(f19288,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1441,f247,f18140,f385])).
fof(f21443,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19295,f18140])).
fof(f19295,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f581,f18140,f18140,f385])).
fof(f21442,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19296,f247])).
fof(f19296,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f920,f18140,f18140,f385])).
fof(f21441,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19297,f247])).
fof(f19297,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1437,f18140,f18140,f385])).
fof(f21440,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f19298,f18140])).
fof(f19298,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f11345,f18140,f18140,f385])).
fof(f21439,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f19299,f18140])).
fof(f19299,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f7544,f18140,f18140,f385])).
fof(f21438,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19300,f18140])).
fof(f19300,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13810,f18140,f18140,f385])).
fof(f21437,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19301,f18140])).
fof(f19301,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10163,f18140,f18140,f385])).
fof(f21436,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19302,f18140])).
fof(f19302,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1441,f18140,f18140,f385])).
fof(f21435,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19303,f18140])).
fof(f19303,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1442,f18140,f18140,f385])).
fof(f21434,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19305,f18140])).
fof(f19305,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1393,f18140,f18140,f385])).
fof(f21433,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19311,f18140])).
fof(f19311,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f703,f439,f18140,f388])).
fof(f703,plain,(
  ( ! [X0] : (sP8('F_0',X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f535,f439,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f378])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f223])).
fof(f21432,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19312,f18140])).
fof(f19312,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f594,f439,f18140,f388])).
fof(f594,plain,(
  ( ! [X0] : (sP8('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f439,f439,f384])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f227])).
fof(f21431,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19313,f18140])).
fof(f19313,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f757,f439,f18140,f388])).
fof(f757,plain,(
  ( ! [X0] : (sP8('F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f439,f536,f425])).
fof(f21430,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19345,f18140])).
fof(f19345,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f586,f439,f18140,f388])).
fof(f586,plain,(
  ( ! [X0] : (sP8('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f439,f439,f383])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f227])).
fof(f21429,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19346,f18140])).
fof(f19346,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f703,f439,f18140,f388])).
fof(f21428,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19348,f18140])).
fof(f19348,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f757,f439,f18140,f388])).
fof(f21427,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19418,f20143])).
fof(f20143,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f458,f18140,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f392])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP9(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f229])).
fof(f229,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1)) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f458,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f439,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f232])).
fof(f232,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f231])).
fof(f231,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f19418,plain,(
  ( ! [X0,X1] : (~sP10('F_0',X0,X1)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f389])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f230])).
fof(f21426,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19458,f20143])).
fof(f19458,plain,(
  ( ! [X0,X1] : (~sP10('F_0',X0,X1)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f390])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f230])).
fof(f21425,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19481,f19541])).
fof(f19541,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f233])).
fof(f19481,plain,(
  ~sP9('F_0','T_0','F_0') | $spl112),
  inference(unit_resulting_resolution,[],[f604,f439,f18140,f393])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f230])).
fof(f604,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f247,f439,f390])).
fof(f21424,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19483,f19541])).
fof(f19483,plain,(
  ~sP9('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0') | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1450,f439,f18140,f393])).
fof(f1450,plain,(
  ( ! [X0] : (~sP10('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f1309,f390])).
fof(f21423,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19509,f19541])).
fof(f19509,plain,(
  ~sP9('F_0','F_0','T_0') | $spl112),
  inference(unit_resulting_resolution,[],[f600,f439,f18140,f394])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f230])).
fof(f600,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f439,f389])).
fof(f21422,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19510,f19541])).
fof(f19510,plain,(
  ~sP9('F_0','F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1445,f439,f18140,f394])).
fof(f1445,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f1309,f389])).
fof(f21421,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19537,f19541])).
fof(f19537,plain,(
  ( ! [X0] : (~sP9(X0,'F_0','F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f439,f18140,f396])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f233])).
fof(f21420,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19543,f247])).
fof(f19543,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f1003,f18140,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f109,f155,f154])).
fof(f109,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) | p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.sat.sat_25_prove_ltb')).
fof(f1003,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f990,f391])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f230])).
fof(f990,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f439,f439,f247,f396])).
fof(f21419,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f19548,f18140])).
fof(f19548,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f11519,f18140,f401])).
fof(f11519,plain,(
  ~sP10(f(s(bool,'T_0')),'F_0','F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f11493,f391])).
fof(f11493,plain,(
  ~sP9(f(s(bool,'T_0')),'F_0','F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f439,f439,f5402])).
fof(f5402,plain,(
  ( ! [X24,X23] : (~sP9(f(s(bool,'T_0')),X23,X24) | p(s(bool,X24)) | p(s(bool,X23))) ) | $spl61),
  inference(subsumption_resolution,[],[f5344,f247])).
fof(f5344,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | p(s(bool,X23)) | p(s(bool,X24)) | ~sP9(f(s(bool,'T_0')),X23,X24)) ) | $spl61),
  inference(superposition,[],[f396,f5088])).
fof(f21418,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f19549,f18140])).
fof(f19549,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f7630,f18140,f401])).
fof(f7630,plain,(
  ~sP10('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),'F_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f7612,f391])).
fof(f7612,plain,(
  ~sP9('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),'F_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f439,f439,f999])).
fof(f999,plain,(
  ( ! [X0,X1] : (~sP9('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f991,f247])).
fof(f991,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP9('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,X1)) ) | $spl6),
  inference(superposition,[],[f396,f490])).
fof(f21417,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19550,f18140])).
fof(f19550,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13974,f18140,f401])).
fof(f13974,plain,(
  ~sP10('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f13956,f391])).
fof(f13956,plain,(
  ~sP9('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f439,f439,f9501])).
fof(f9501,plain,(
  ( ! [X33,X32] : (~sP9('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X32,X33) | p(s(bool,X33)) | p(s(bool,X32))) ) | $spl4),
  inference(subsumption_resolution,[],[f9380,f247])).
fof(f9380,plain,(
  ( ! [X33,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | p(s(bool,X33)) | ~sP9('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X32,X33)) ) | $spl4),
  inference(superposition,[],[f396,f7026])).
fof(f21416,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19551,f18140])).
fof(f19551,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10419,f18140,f401])).
fof(f10419,plain,(
  ~sP10(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f10401,f391])).
fof(f10401,plain,(
  ~sP9(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f439,f439,f1096])).
fof(f1096,plain,(
  ( ! [X21,X22] : (~sP9(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X21,X22) | p(s(bool,X22)) | p(s(bool,X21))) ) | $spl4),
  inference(subsumption_resolution,[],[f1074,f247])).
fof(f1074,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | p(s(bool,X21)) | p(s(bool,X22)) | ~sP9(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X21,X22)) ) | $spl4),
  inference(superposition,[],[f396,f491])).
fof(f21415,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19552,f18140])).
fof(f19552,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1480,f18140,f401])).
fof(f1480,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f439,f439,f1309,f403])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f21414,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19558,f247])).
fof(f19558,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f600,f18140,f401])).
fof(f21413,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19559,f18140])).
fof(f19559,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1445,f18140,f401])).
fof(f21412,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19560,f18140])).
fof(f19560,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f604,f18140,f401])).
fof(f21411,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19561,f18140])).
fof(f19561,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1450,f18140,f401])).
fof(f21410,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19571,f247])).
fof(f19571,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f1003,f18140,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f21409,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f19576,f18140])).
fof(f19576,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f11519,f18140,f402])).
fof(f21408,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f19577,f18140])).
fof(f19577,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f7630,f18140,f402])).
fof(f21407,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19578,f18140])).
fof(f19578,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13974,f18140,f402])).
fof(f21406,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19579,f18140])).
fof(f19579,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10419,f18140,f402])).
fof(f21405,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19580,f18140])).
fof(f19580,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1480,f18140,f402])).
fof(f21404,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19586,f18140])).
fof(f19586,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f600,f18140,f402])).
fof(f21403,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19587,f18140])).
fof(f19587,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1445,f18140,f402])).
fof(f21402,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19588,f247])).
fof(f19588,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f604,f18140,f402])).
fof(f21401,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19589,f18140])).
fof(f19589,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1450,f18140,f402])).
fof(f21400,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19673,f20114])).
fof(f20114,plain,(
  ( ! [X0,X1] : (sP12('F_0',X0,X1)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f461,f18140,f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f409])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP11(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f237])).
fof(f237,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1)) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f461,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f439,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f240])).
fof(f240,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f19673,plain,(
  ( ! [X0,X1] : (~sP12('F_0',X0,X1)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f407])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f238])).
fof(f21399,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19676,f18140])).
fof(f19676,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f461,f645,f18140,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f238])).
fof(f645,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f247,f439,f407])).
fof(f21398,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19677,f18140])).
fof(f19677,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f460,f645,f18140,f410])).
fof(f460,plain,(
  ( ! [X0,X1] : (sP11(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f247,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f241])).
fof(f21397,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19679,f18140])).
fof(f19679,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f538,f645,f18140,f410])).
fof(f538,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f433,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f21396,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19680,f18140])).
fof(f19680,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f461,f1481,f18140,f410])).
fof(f1481,plain,(
  ( ! [X0] : (~sP12('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f1309,f407])).
fof(f21395,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19681,f18140])).
fof(f19681,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1519,f1481,f18140,f410])).
fof(f1519,plain,(
  ( ! [X0,X1] : (sP11(X0,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1309,f415])).
fof(f21394,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19683,f18140])).
fof(f19683,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f538,f1481,f18140,f410])).
fof(f21393,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19691,f19886])).
fof(f19886,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f415])).
fof(f19691,plain,(
  ~sP11('F_0','T_0','T_0') | $spl112),
  inference(unit_resulting_resolution,[],[f645,f247,f18140,f410])).
fof(f21392,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19693,f19886])).
fof(f19693,plain,(
  ~sP11('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'T_0') | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1481,f247,f18140,f410])).
fof(f21391,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19714,f19886])).
fof(f19714,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f644,f18140,f18140,f410])).
fof(f644,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f439,f439,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f238])).
fof(f21390,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19716,f19886])).
fof(f19716,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f645,f18140,f18140,f410])).
fof(f21389,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19718,f19886])).
fof(f19718,plain,(
  ( ! [X0] : (~sP11('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X0)) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1481,f18140,f18140,f410])).
fof(f21388,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19756,f19886])).
fof(f19756,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f644,f18140,f18140,f410])).
fof(f21387,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19757,f19886])).
fof(f19757,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f645,f18140,f18140,f410])).
fof(f21386,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19758,f19886])).
fof(f19758,plain,(
  ( ! [X0] : (~sP11('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X0)) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1481,f18140,f18140,f410])).
fof(f21385,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19798,f19886])).
fof(f19798,plain,(
  ~sP11('F_0','F_0','F_0') | $spl112),
  inference(unit_resulting_resolution,[],[f644,f439,f18140,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f238])).
fof(f21384,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19829,f19886])).
fof(f19829,plain,(
  ~sP11('F_0','F_0','F_0') | $spl112),
  inference(unit_resulting_resolution,[],[f644,f439,f18140,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f238])).
fof(f21383,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19857,f18140])).
fof(f19857,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f538,f439,f18140,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f21382,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19858,f18140])).
fof(f19858,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f461,f439,f18140,f413])).
fof(f21381,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19859,f18140])).
fof(f19859,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f539,f439,f18140,f413])).
fof(f539,plain,(
  ( ! [X0,X1] : (sP11('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f433,f416])).
fof(f21380,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19862,f19886])).
fof(f19862,plain,(
  ( ! [X0,X1] : (~sP11(X0,'F_0',X1)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f439,f18140,f413])).
fof(f21379,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19878,f18140])).
fof(f19878,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f459,f439,f18140,f413])).
fof(f459,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f439,f414])).
fof(f21378,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19879,f18140])).
fof(f19879,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f538,f439,f18140,f413])).
fof(f21377,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19881,f18140])).
fof(f19881,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f539,f439,f18140,f413])).
fof(f21376,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19882,f19886])).
fof(f19882,plain,(
  ( ! [X0] : (~sP11('T_0','F_0',X0)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f247,f439,f18140,f413])).
fof(f21375,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19884,f19886])).
fof(f19884,plain,(
  ( ! [X0,X1] : (~sP11(X0,'F_0',X1)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f439,f18140,f413])).
fof(f21374,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19923,f247])).
fof(f19923,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f1144,f18140,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f139,f158,f157])).
fof(f139,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) => p(s(bool,X20)))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (~p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.sat.sat_27_prove_ltb')).
fof(f1144,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1130,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f238])).
fof(f1130,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f247,f439,f247,f413])).
fof(f21373,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f19924,f247])).
fof(f19924,plain,(
  ~p(s(bool,'T_0')) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f11712,f18140,f419])).
fof(f11712,plain,(
  ~sP12('T_0','F_0',f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f11690,f408])).
fof(f11690,plain,(
  ~sP11('T_0','F_0',f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f439,f247,f5406])).
fof(f5406,plain,(
  ( ! [X31,X32] : (~sP11(X32,X31,f(s(bool,'T_0'))) | ~p(s(bool,X32)) | p(s(bool,X31))) ) | $spl61),
  inference(subsumption_resolution,[],[f5348,f247])).
fof(f5348,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~sP11(X32,X31,f(s(bool,'T_0')))) ) | $spl61),
  inference(superposition,[],[f413,f5088])).
fof(f21372,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f19925,f247])).
fof(f19925,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f7782,f18140,f419])).
fof(f7782,plain,(
  ~sP12('T_0','F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | $spl6),
  inference(unit_resulting_resolution,[],[f7768,f408])).
fof(f7768,plain,(
  ~sP11('T_0','F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | $spl6),
  inference(unit_resulting_resolution,[],[f439,f247,f1139])).
fof(f1139,plain,(
  ( ! [X6,X7] : (~sP11(X7,X6,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | ~p(s(bool,X7)) | p(s(bool,X6))) ) | $spl6),
  inference(subsumption_resolution,[],[f1133,f247])).
fof(f1133,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | ~p(s(bool,X7)) | ~sP11(X7,X6,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) ) | $spl6),
  inference(superposition,[],[f413,f490])).
fof(f21371,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19926,f247])).
fof(f19926,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f14071,f18140,f419])).
fof(f14071,plain,(
  ~sP12('T_0','F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f14057,f408])).
fof(f14057,plain,(
  ~sP11('T_0','F_0','const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f247,f9505])).
fof(f9505,plain,(
  ( ! [X41,X40] : (~sP11(X41,X40,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | ~p(s(bool,X41)) | p(s(bool,X40))) ) | $spl4),
  inference(subsumption_resolution,[],[f9384,f247])).
fof(f9384,plain,(
  ( ! [X41,X40] : (~p(s(bool,'T_0')) | p(s(bool,X40)) | ~p(s(bool,X41)) | ~sP11(X41,X40,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) ) | $spl4),
  inference(superposition,[],[f413,f7026])).
fof(f21370,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19927,f247])).
fof(f19927,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10517,f18140,f419])).
fof(f10517,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f10503,f408])).
fof(f10503,plain,(
  ~sP11('T_0','F_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f247,f1140])).
fof(f1140,plain,(
  ( ! [X8,X9] : (~sP11(X9,X8,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | ~p(s(bool,X9)) | p(s(bool,X8))) ) | $spl4),
  inference(subsumption_resolution,[],[f1134,f247])).
fof(f1134,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X8)) | ~p(s(bool,X9)) | ~sP11(X9,X8,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) ) | $spl4),
  inference(superposition,[],[f413,f491])).
fof(f21369,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19928,f247])).
fof(f19928,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1528,f18140,f419])).
fof(f1528,plain,(
  ~sP12('T_0','F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f1309,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f243])).
fof(f21368,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f19932,f18140])).
fof(f19932,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f11814,f18140,f419])).
fof(f11814,plain,(
  ~sP12(f(s(bool,'T_0')),'F_0','T_0') | $spl61),
  inference(unit_resulting_resolution,[],[f247,f439,f5409])).
fof(f5409,plain,(
  ( ! [X37,X38] : (~sP12(f(s(bool,'T_0')),X38,X37) | p(s(bool,X38)) | ~p(s(bool,X37))) ) | $spl61),
  inference(subsumption_resolution,[],[f5351,f247])).
fof(f5351,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP12(f(s(bool,'T_0')),X38,X37)) ) | $spl61),
  inference(superposition,[],[f420,f5088])).
fof(f21367,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f19933,f18140])).
fof(f19933,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f7822,f18140,f419])).
fof(f7822,plain,(
  ~sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),'F_0','T_0') | $spl6),
  inference(unit_resulting_resolution,[],[f247,f439,f1179])).
fof(f1179,plain,(
  ( ! [X6,X7] : (~sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X7,X6) | p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl6),
  inference(subsumption_resolution,[],[f1173,f247])).
fof(f1173,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X7,X6)) ) | $spl6),
  inference(superposition,[],[f420,f490])).
fof(f21366,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19934,f18140])).
fof(f19934,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f14160,f18140,f419])).
fof(f14160,plain,(
  ~sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f9507])).
fof(f9507,plain,(
  ( ! [X47,X46] : (~sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X47,X46) | p(s(bool,X47)) | ~p(s(bool,X46))) ) | $spl4),
  inference(subsumption_resolution,[],[f9387,f247])).
fof(f9387,plain,(
  ( ! [X47,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X47,X46)) ) | $spl4),
  inference(superposition,[],[f420,f7026])).
fof(f21365,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19935,f18140])).
fof(f19935,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10565,f18140,f419])).
fof(f10565,plain,(
  ~sP12(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f1180])).
fof(f1180,plain,(
  ( ! [X8,X9] : (~sP12(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X9,X8) | p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl4),
  inference(subsumption_resolution,[],[f1174,f247])).
fof(f1174,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~sP12(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X9,X8)) ) | $spl4),
  inference(superposition,[],[f420,f491])).
fof(f21364,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19936,f18140])).
fof(f19936,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1533,f18140,f419])).
fof(f1533,plain,(
  ~sP12('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f247,f439,f1309,f420])).
fof(f21363,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19937,f18140])).
fof(f19937,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1534,f18140,f419])).
fof(f1534,plain,(
  ~sP12('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1309,f439,f1309,f420])).
fof(f21362,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19948,f18140])).
fof(f19948,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f644,f18140,f419])).
fof(f21361,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19949,f247])).
fof(f19949,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f645,f18140,f419])).
fof(f21360,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19950,f18140])).
fof(f19950,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1481,f18140,f419])).
fof(f21359,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f19961,f247])).
fof(f19961,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f648,f439,f18140,f420])).
fof(f648,plain,(
  ( ! [X0] : (sP12('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f439,f247,f418])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f243])).
fof(f21358,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19962,f247])).
fof(f19962,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f770,f439,f18140,f420])).
fof(f770,plain,(
  ( ! [X0] : (sP12('T_0',X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f247,f538,f423])).
fof(f21357,plain,(
  $false | ($spl61 | $spl112)),
  inference(subsumption_resolution,[],[f19965,f18140])).
fof(f19965,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f6556,f439,f18140,f420])).
fof(f6556,plain,(
  ( ! [X0] : (sP12(f(s(bool,'T_0')),X0,'F_0')) ) | $spl61),
  inference(unit_resulting_resolution,[],[f439,f5407])).
fof(f5407,plain,(
  ( ! [X33,X34] : (sP12(f(s(bool,'T_0')),X34,X33) | p(s(bool,X33))) ) | $spl61),
  inference(subsumption_resolution,[],[f5349,f247])).
fof(f5349,plain,(
  ( ! [X33,X34] : (~p(s(bool,'T_0')) | p(s(bool,X33)) | sP12(f(s(bool,'T_0')),X34,X33)) ) | $spl61),
  inference(superposition,[],[f418,f5088])).
fof(f21356,plain,(
  $false | ($spl1 | $spl61 | $spl112)),
  inference(subsumption_resolution,[],[f19966,f18140])).
fof(f19966,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl1 | $spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f7894,f439,f18140,f420])).
fof(f7894,plain,(
  ( ! [X0] : (sP12(f(s(bool,'T_0')),X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) ) | ($spl1 | $spl61)),
  inference(unit_resulting_resolution,[],[f538,f5410])).
fof(f5410,plain,(
  ( ! [X39,X40] : (sP12(f(s(bool,'T_0')),X39,X40) | ~sP11(f(s(bool,'T_0')),X39,X40)) ) | $spl61),
  inference(subsumption_resolution,[],[f5352,f247])).
fof(f5352,plain,(
  ( ! [X39,X40] : (~p(s(bool,'T_0')) | ~sP11(f(s(bool,'T_0')),X39,X40) | sP12(f(s(bool,'T_0')),X39,X40)) ) | $spl61),
  inference(superposition,[],[f423,f5088])).
fof(f21355,plain,(
  $false | ($spl6 | $spl112)),
  inference(subsumption_resolution,[],[f19967,f18140])).
fof(f19967,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f1015,f439,f18140,f420])).
fof(f1015,plain,(
  ( ! [X0] : (sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,'F_0')) ) | $spl6),
  inference(unit_resulting_resolution,[],[f439,f650])).
fof(f650,plain,(
  ( ! [X0,X1] : (sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X1,X0) | p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f649,f247])).
fof(f649,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X1,X0)) ) | $spl6),
  inference(superposition,[],[f418,f490])).
fof(f21354,plain,(
  $false | ($spl1 | $spl6 | $spl112)),
  inference(subsumption_resolution,[],[f19968,f18140])).
fof(f19968,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | ($spl1 | $spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f17640,f439,f18140,f420])).
fof(f17640,plain,(
  ( ! [X0] : (sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) ) | ($spl1 | $spl6)),
  inference(unit_resulting_resolution,[],[f538,f674])).
fof(f674,plain,(
  ( ! [X0,X1] : (sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,X1) | ~sP11('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f673,f247])).
fof(f673,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP11('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,X1) | sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,X1)) ) | $spl6),
  inference(superposition,[],[f423,f490])).
fof(f21353,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f19969,f18140])).
fof(f19969,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1117,f439,f18140,f420])).
fof(f1117,plain,(
  sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f539,f247,f439,f412])).
fof(f21352,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19970,f18140])).
fof(f19970,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1504,f439,f18140,f420])).
fof(f1504,plain,(
  sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)),'F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f439,f539,f1309,f412])).
fof(f21351,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19971,f18140])).
fof(f19971,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10652,f439,f18140,f420])).
fof(f10652,plain,(
  ( ! [X0] : (sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f8033])).
fof(f8033,plain,(
  ( ! [X0,X1] : (sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X1,X0) | p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f8007,f247])).
fof(f8007,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X1,X0)) ) | $spl4),
  inference(superposition,[],[f1793,f491])).
fof(f1793,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (~p(s(bool,happ(s(fun(X133,bool),happ(s(fun(X133,fun(X133,bool)),X134),s(X133,X135))),s(X133,X136)))) | p(s(bool,X137)) | sP12('const.arithmetic.NRC_4'(s(fun(X133,fun(X133,bool)),X134),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X133,X135),s(X133,X136)),X138,X137)) )),
  inference(superposition,[],[f418,f301])).
fof(f21350,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19972,f18140])).
fof(f19972,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f5056,f439,f18140,f420])).
fof(f5056,plain,(
  ( ! [X0] : (sP12(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f1100])).
fof(f1100,plain,(
  ( ! [X30,X29] : (sP12(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X30,X29) | p(s(bool,X29))) ) | $spl4),
  inference(subsumption_resolution,[],[f1078,f247])).
fof(f1078,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP12(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X30,X29)) ) | $spl4),
  inference(superposition,[],[f418,f491])).
fof(f21349,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19973,f18140])).
fof(f19973,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1535,f439,f18140,f420])).
fof(f1535,plain,(
  ( ! [X0] : (sP12('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f459,f1309,f423])).
fof(f21348,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f19974,f18140])).
fof(f19974,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1536,f439,f18140,f420])).
fof(f1536,plain,(
  ( ! [X0] : (sP12('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) ) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f538,f1309,f423])).
fof(f21347,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20018,f18140])).
fof(f20018,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f770,f439,f18140,f420])).
fof(f21346,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20019,f247])).
fof(f20019,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(unit_resulting_resolution,[],[f1116,f439,f18140,f420])).
fof(f1116,plain,(
  sP12('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f461,f247,f439,f412])).
fof(f21345,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20020,f18140])).
fof(f20020,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1520,f439,f18140,f420])).
fof(f1520,plain,(
  sP12('F_0','F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f439,f439,f1309,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f243])).
fof(f21344,plain,(
  $false | ($spl1 | $spl61 | $spl112)),
  inference(subsumption_resolution,[],[f20022,f18140])).
fof(f20022,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f7894,f439,f18140,f420])).
fof(f21343,plain,(
  $false | ($spl1 | $spl6 | $spl112)),
  inference(subsumption_resolution,[],[f20024,f18140])).
fof(f20024,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f17640,f439,f18140,f420])).
fof(f21342,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20025,f247])).
fof(f20025,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1117,f439,f18140,f420])).
fof(f21341,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20026,f18140])).
fof(f20026,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1504,f439,f18140,f420])).
fof(f21340,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20030,f18140])).
fof(f20030,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1536,f439,f18140,f420])).
fof(f21339,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20093,f459])).
fof(f20093,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f644,f18140,f423])).
fof(f21338,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20094,f460])).
fof(f20094,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f645,f18140,f423])).
fof(f21337,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20095,f1519])).
fof(f20095,plain,(
  ( ! [X0] : (~sP11('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X0)) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1481,f18140,f423])).
fof(f21336,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20118,f456])).
fof(f456,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f397])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f233])).
fof(f20118,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl112),
  inference(unit_resulting_resolution,[],[f600,f18140,f424])).
fof(f21335,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20119,f1468])).
fof(f1468,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1309,f397])).
fof(f20119,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1445,f18140,f424])).
fof(f21334,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20120,f457])).
fof(f457,plain,(
  ( ! [X0,X1] : (sP9(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f247,f398])).
fof(f20120,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl112),
  inference(unit_resulting_resolution,[],[f604,f18140,f424])).
fof(f21333,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20121,f1469])).
fof(f1469,plain,(
  ( ! [X0,X1] : (sP9(X0,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1309,f398])).
fof(f20121,plain,(
  ( ! [X0] : (~sP9('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X0)) ) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1450,f18140,f424])).
fof(f21332,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20261,f18140])).
fof(f20261,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f562,f18140,f1618])).
fof(f1618,plain,(
  ( ! [X43,X41,X44,X42] : (sP2(X43,X44,happ(s(fun(X41,bool),X42),s(X41,'const.min.@_1'(s(fun(X41,bool),X42))))) | ~p(s(bool,X43)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X41,bool),X42))))) )),
  inference(superposition,[],[f344,f289])).
fof(f21331,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20268,f18140])).
fof(f20268,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f446,f18140,f1620])).
fof(f1620,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'|dollar|exists_1'(s(fun(X49,bool),X50)))) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP5(X51,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))),X52)) )),
  inference(superposition,[],[f356,f289])).
fof(f21330,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20271,f18140])).
fof(f20271,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f532,f18140,f1620])).
fof(f21329,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20279,f18140])).
fof(f20279,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f449,f18140,f1621])).
fof(f1621,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'|dollar|exists_1'(s(fun(X53,bool),X54)))) | ~p(s(bool,X55)) | p(s(bool,X56)) | ~sP4(X55,happ(s(fun(X53,bool),X54),s(X53,'const.min.@_1'(s(fun(X53,bool),X54)))),X56)) )),
  inference(superposition,[],[f360,f289])).
fof(f21328,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20282,f18140])).
fof(f20282,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f534,f18140,f1621])).
fof(f21327,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20290,f18140])).
fof(f20290,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f455,f18140,f1622])).
fof(f1622,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'|dollar|exists_1'(s(fun(X57,bool),X58)))) | ~p(s(bool,X59)) | p(s(bool,X60)) | ~sP7(X59,happ(s(fun(X57,bool),X58),s(X57,'const.min.@_1'(s(fun(X57,bool),X58)))),X60)) )),
  inference(superposition,[],[f379,f289])).
fof(f21326,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20293,f18140])).
fof(f20293,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f536,f18140,f1622])).
fof(f21325,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20297,f18140])).
fof(f20297,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f594,f18140,f1623])).
fof(f1623,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'|dollar|exists_1'(s(fun(X61,bool),X62)))) | ~p(s(bool,X63)) | p(s(bool,X64)) | ~sP8(X64,X63,happ(s(fun(X61,bool),X62),s(X61,'const.min.@_1'(s(fun(X61,bool),X62)))))) )),
  inference(superposition,[],[f388,f289])).
fof(f21324,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20300,f18140])).
fof(f20300,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f757,f18140,f1623])).
fof(f21323,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20310,f18140])).
fof(f20310,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f461,f18140,f1628])).
fof(f1628,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),X82)))) | p(s(bool,X83)) | ~p(s(bool,X84)) | ~sP11(X84,X83,happ(s(fun(X81,bool),X82),s(X81,'const.min.@_1'(s(fun(X81,bool),X82)))))) )),
  inference(superposition,[],[f413,f289])).
fof(f21322,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20313,f18140])).
fof(f20313,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f539,f18140,f1628])).
fof(f21321,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20322,f18140])).
fof(f20322,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f562,f18140,f1751])).
fof(f1751,plain,(
  ( ! [X94,X99,X97,X95,X98,X96] : (~p(s(bool,X98)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X94,fun(X94,bool)),X95),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X94,X96),s(X94,X97)))) | sP2(X98,X99,happ(s(fun(X94,bool),happ(s(fun(X94,fun(X94,bool)),X95),s(X94,X96))),s(X94,X97)))) )),
  inference(superposition,[],[f344,f301])).
fof(f21320,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20325,f18140])).
fof(f20325,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f446,f18140,f1753])).
fof(f1753,plain,(
  ( ! [X111,X109,X107,X110,X108,X106] : (~p(s(bool,X110)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X106,fun(X106,bool)),X107),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X106,X108),s(X106,X109)))) | p(s(bool,X111)) | ~sP5(X110,happ(s(fun(X106,bool),happ(s(fun(X106,fun(X106,bool)),X107),s(X106,X108))),s(X106,X109)),X111)) )),
  inference(superposition,[],[f356,f301])).
fof(f21319,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20326,f18140])).
fof(f20326,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f532,f18140,f1753])).
fof(f21318,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20336,f18140])).
fof(f20336,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f449,f18140,f1754])).
fof(f1754,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (~p(s(bool,X116)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X112,fun(X112,bool)),X113),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X112,X114),s(X112,X115)))) | p(s(bool,X117)) | ~sP4(X116,happ(s(fun(X112,bool),happ(s(fun(X112,fun(X112,bool)),X113),s(X112,X114))),s(X112,X115)),X117)) )),
  inference(superposition,[],[f360,f301])).
fof(f21317,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20337,f18140])).
fof(f20337,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f534,f18140,f1754])).
fof(f21316,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20347,f18140])).
fof(f20347,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f455,f18140,f1755])).
fof(f1755,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (~p(s(bool,X122)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X118,fun(X118,bool)),X119),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X118,X120),s(X118,X121)))) | p(s(bool,X123)) | ~sP7(X122,happ(s(fun(X118,bool),happ(s(fun(X118,fun(X118,bool)),X119),s(X118,X120))),s(X118,X121)),X123)) )),
  inference(superposition,[],[f379,f301])).
fof(f21315,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20348,f18140])).
fof(f20348,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f536,f18140,f1755])).
fof(f21314,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20358,f18140])).
fof(f20358,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f594,f18140,f1756])).
fof(f1756,plain,(
  ( ! [X127,X125,X128,X126,X124,X129] : (~p(s(bool,X128)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X124,fun(X124,bool)),X125),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X124,X126),s(X124,X127)))) | p(s(bool,X129)) | ~sP8(X129,X128,happ(s(fun(X124,bool),happ(s(fun(X124,fun(X124,bool)),X125),s(X124,X126))),s(X124,X127)))) )),
  inference(superposition,[],[f388,f301])).
fof(f21313,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20359,f18140])).
fof(f20359,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f757,f18140,f1756])).
fof(f21312,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20365,f18140])).
fof(f20365,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10419,f18140,f1758])).
fof(f1758,plain,(
  ( ! [X140,X138,X136,X141,X139,X137] : (~p(s(bool,X140)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X136,fun(X136,bool)),X137),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X136,X138),s(X136,X139)))) | sP10(happ(s(fun(X136,bool),happ(s(fun(X136,fun(X136,bool)),X137),s(X136,X138))),s(X136,X139)),X141,X140)) )),
  inference(superposition,[],[f401,f301])).
fof(f21311,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20368,f18140])).
fof(f20368,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10419,f18140,f1759])).
fof(f1759,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (~p(s(bool,X146)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X142,fun(X142,bool)),X143),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X142,X144),s(X142,X145)))) | sP10(happ(s(fun(X142,bool),happ(s(fun(X142,fun(X142,bool)),X143),s(X142,X144))),s(X142,X145)),X146,X147)) )),
  inference(superposition,[],[f402,f301])).
fof(f21310,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20371,f18140])).
fof(f20371,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f461,f18140,f1761])).
fof(f1761,plain,(
  ( ! [X158,X156,X154,X159,X157,X155] : (~p(s(bool,X159)) | p(s(bool,X158)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X154,fun(X154,bool)),X155),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X154,X156),s(X154,X157)))) | ~sP11(X159,X158,happ(s(fun(X154,bool),happ(s(fun(X154,fun(X154,bool)),X155),s(X154,X156))),s(X154,X157)))) )),
  inference(superposition,[],[f413,f301])).
fof(f21309,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20372,f18140])).
fof(f20372,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f539,f18140,f1761])).
fof(f21308,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20382,f18140])).
fof(f20382,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10565,f18140,f1763])).
fof(f1763,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (~p(s(bool,X170)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X166,fun(X166,bool)),X167),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X166,X168),s(X166,X169)))) | sP12(happ(s(fun(X166,bool),happ(s(fun(X166,fun(X166,bool)),X167),s(X166,X168))),s(X166,X169)),X170,X171)) )),
  inference(superposition,[],[f419,f301])).
fof(f21307,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20385,f18140])).
fof(f20385,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f5056,f18140,f1764])).
fof(f1764,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (~p(s(bool,X176)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X172,fun(X172,bool)),X173),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X172,X174),s(X172,X175)))) | p(s(bool,X177)) | ~sP12(happ(s(fun(X172,bool),happ(s(fun(X172,fun(X172,bool)),X173),s(X172,X174))),s(X172,X175)),X177,X176)) )),
  inference(superposition,[],[f420,f301])).
fof(f21306,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20393,f18140])).
fof(f20393,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f562,f18140,f1782])).
fof(f1782,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,X71)) | ~p(s(bool,happ(s(fun(X67,bool),happ(s(fun(X67,fun(X67,bool)),X68),s(X67,X69))),s(X67,X70)))) | sP2(X71,X72,'const.arithmetic.NRC_4'(s(fun(X67,fun(X67,bool)),X68),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X67,X69),s(X67,X70)))) )),
  inference(superposition,[],[f344,f301])).
fof(f21305,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20400,f18140])).
fof(f20400,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f446,f18140,f1784])).
fof(f1784,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (~p(s(bool,X83)) | ~p(s(bool,happ(s(fun(X79,bool),happ(s(fun(X79,fun(X79,bool)),X80),s(X79,X81))),s(X79,X82)))) | p(s(bool,X84)) | ~sP5(X83,'const.arithmetic.NRC_4'(s(fun(X79,fun(X79,bool)),X80),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X79,X81),s(X79,X82)),X84)) )),
  inference(superposition,[],[f356,f301])).
fof(f21304,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20403,f18140])).
fof(f20403,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f532,f18140,f1784])).
fof(f21303,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20411,f18140])).
fof(f20411,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f449,f18140,f1785])).
fof(f1785,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,X89)) | ~p(s(bool,happ(s(fun(X85,bool),happ(s(fun(X85,fun(X85,bool)),X86),s(X85,X87))),s(X85,X88)))) | p(s(bool,X90)) | ~sP4(X89,'const.arithmetic.NRC_4'(s(fun(X85,fun(X85,bool)),X86),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X85,X87),s(X85,X88)),X90)) )),
  inference(superposition,[],[f360,f301])).
fof(f21302,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20414,f18140])).
fof(f20414,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f534,f18140,f1785])).
fof(f21301,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20422,f18140])).
fof(f20422,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f455,f18140,f1786])).
fof(f1786,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (~p(s(bool,X95)) | ~p(s(bool,happ(s(fun(X91,bool),happ(s(fun(X91,fun(X91,bool)),X92),s(X91,X93))),s(X91,X94)))) | p(s(bool,X96)) | ~sP7(X95,'const.arithmetic.NRC_4'(s(fun(X91,fun(X91,bool)),X92),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X91,X93),s(X91,X94)),X96)) )),
  inference(superposition,[],[f379,f301])).
fof(f21300,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20425,f18140])).
fof(f20425,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f536,f18140,f1786])).
fof(f21299,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20429,f18140])).
fof(f20429,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f594,f18140,f1787])).
fof(f1787,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (~p(s(bool,X101)) | ~p(s(bool,happ(s(fun(X97,bool),happ(s(fun(X97,fun(X97,bool)),X98),s(X97,X99))),s(X97,X100)))) | p(s(bool,X102)) | ~sP8(X102,X101,'const.arithmetic.NRC_4'(s(fun(X97,fun(X97,bool)),X98),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X97,X99),s(X97,X100)))) )),
  inference(superposition,[],[f388,f301])).
fof(f21298,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20432,f18140])).
fof(f20432,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f757,f18140,f1787])).
fof(f21297,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20436,f18140])).
fof(f20436,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13974,f18140,f1789])).
fof(f1789,plain,(
  ( ! [X111,X109,X114,X112,X110,X113] : (~p(s(bool,X113)) | ~p(s(bool,happ(s(fun(X109,bool),happ(s(fun(X109,fun(X109,bool)),X110),s(X109,X111))),s(X109,X112)))) | sP10('const.arithmetic.NRC_4'(s(fun(X109,fun(X109,bool)),X110),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X109,X111),s(X109,X112)),X114,X113)) )),
  inference(superposition,[],[f401,f301])).
fof(f21296,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20439,f18140])).
fof(f20439,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13974,f18140,f1790])).
fof(f1790,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (~p(s(bool,X119)) | ~p(s(bool,happ(s(fun(X115,bool),happ(s(fun(X115,fun(X115,bool)),X116),s(X115,X117))),s(X115,X118)))) | sP10('const.arithmetic.NRC_4'(s(fun(X115,fun(X115,bool)),X116),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X115,X117),s(X115,X118)),X119,X120)) )),
  inference(superposition,[],[f402,f301])).
fof(f21295,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20446,f18140])).
fof(f20446,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f439,f461,f18140,f1792])).
fof(f1792,plain,(
  ( ! [X127,X132,X130,X128,X131,X129] : (~p(s(bool,X132)) | p(s(bool,X131)) | ~p(s(bool,happ(s(fun(X127,bool),happ(s(fun(X127,fun(X127,bool)),X128),s(X127,X129))),s(X127,X130)))) | ~sP11(X132,X131,'const.arithmetic.NRC_4'(s(fun(X127,fun(X127,bool)),X128),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X127,X129),s(X127,X130)))) )),
  inference(superposition,[],[f413,f301])).
fof(f21294,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20449,f18140])).
fof(f20449,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f539,f18140,f1792])).
fof(f21293,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20453,f18140])).
fof(f20453,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f14160,f18140,f1794])).
fof(f1794,plain,(
  ( ! [X144,X142,X140,X143,X141,X139] : (~p(s(bool,X143)) | ~p(s(bool,happ(s(fun(X139,bool),happ(s(fun(X139,fun(X139,bool)),X140),s(X139,X141))),s(X139,X142)))) | sP12('const.arithmetic.NRC_4'(s(fun(X139,fun(X139,bool)),X140),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X139,X141),s(X139,X142)),X143,X144)) )),
  inference(superposition,[],[f419,f301])).
fof(f21292,plain,(
  $false | ($spl4 | $spl112)),
  inference(subsumption_resolution,[],[f20456,f18140])).
fof(f20456,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f10652,f18140,f1795])).
fof(f1795,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (~p(s(bool,X149)) | ~p(s(bool,happ(s(fun(X145,bool),happ(s(fun(X145,fun(X145,bool)),X146),s(X145,X147))),s(X145,X148)))) | p(s(bool,X150)) | ~sP12('const.arithmetic.NRC_4'(s(fun(X145,fun(X145,bool)),X146),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X145,X147),s(X145,X148)),X150,X149)) )),
  inference(superposition,[],[f420,f301])).
fof(f21291,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20705,f20707])).
fof(f20705,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f1740])).
fof(f1740,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X39,fun(X39,bool)),X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X39,X41),s(X39,X42)))) | s(bool,'F_0') != s(bool,happ(s(fun(X39,bool),happ(s(fun(X39,fun(X39,bool)),X40),s(X39,X41))),s(X39,X42)))) )),
  inference(superposition,[],[f266,f301])).
fof(f21290,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20784,f18140])).
fof(f20784,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f532,f18140,f1753])).
fof(f21289,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20796,f18140])).
fof(f20796,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f534,f18140,f1754])).
fof(f21288,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20808,f18140])).
fof(f20808,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f536,f18140,f1755])).
fof(f21287,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20816,f18140])).
fof(f20816,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f757,f18140,f1756])).
fof(f21286,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f20848,f18140])).
fof(f20848,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f539,f18140,f1761])).
fof(f21285,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f20955,f20957])).
fof(f20955,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3)))) ) | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f1771])).
fof(f1771,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(X12,fun(X12,bool)),X13),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X12,X14),s(X12,X15))) | ~p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X12,fun(X12,bool)),X13),s(X12,X14))),s(X12,X15))))) )),
  inference(superposition,[],[f266,f301])).
fof(f21284,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f21013,f18140])).
fof(f21013,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f532,f18140,f1784])).
fof(f21283,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f21025,f18140])).
fof(f21025,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f534,f18140,f1785])).
fof(f21282,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f21037,f18140])).
fof(f21037,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f536,f18140,f1786])).
fof(f21281,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f21045,f18140])).
fof(f21045,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f757,f18140,f1787])).
fof(f21280,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f21077,f18140])).
fof(f21077,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f539,f18140,f1792])).
fof(f21279,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f21162,f18140])).
fof(f21162,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f532,f18140,f1620])).
fof(f21278,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f21174,f18140])).
fof(f21174,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f534,f18140,f1621])).
fof(f21277,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f21186,f18140])).
fof(f21186,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f536,f18140,f1622])).
fof(f21276,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f21194,f18140])).
fof(f21194,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f757,f18140,f1623])).
fof(f21275,plain,(
  $false | ($spl1 | $spl112)),
  inference(subsumption_resolution,[],[f21220,f18140])).
fof(f21220,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f539,f18140,f1628])).
fof(f21222,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f539,f18140,f1628])).
fof(f21219,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f461,f18140,f1628])).
fof(f21196,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f757,f18140,f1623])).
fof(f21193,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f439,f594,f18140,f1623])).
fof(f21188,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f536,f18140,f1622])).
fof(f21185,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f455,f18140,f1622])).
fof(f21176,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f534,f18140,f1621])).
fof(f21173,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f449,f18140,f1621])).
fof(f21164,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f532,f18140,f1620])).
fof(f21161,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f446,f18140,f1620])).
fof(f21118,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f562,f18140,f1618])).
fof(f21090,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f10652,f18140,f1795])).
fof(f21085,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f18140,f14160,f18140,f1794])).
fof(f21079,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f539,f18140,f1792])).
fof(f21076,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f461,f18140,f1792])).
fof(f21064,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f18140,f13974,f18140,f1790])).
fof(f21060,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f18140,f13974,f18140,f1789])).
fof(f21047,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f757,f18140,f1787])).
fof(f21044,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f439,f594,f18140,f1787])).
fof(f21039,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f536,f18140,f1786])).
fof(f21036,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f455,f18140,f1786])).
fof(f21027,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f534,f18140,f1785])).
fof(f21024,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f449,f18140,f1785])).
fof(f21015,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f532,f18140,f1784])).
fof(f21012,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f446,f18140,f1784])).
fof(f20967,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f562,f18140,f1782])).
fof(f20861,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f5056,f18140,f1764])).
fof(f20856,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f18140,f10565,f18140,f1763])).
fof(f20850,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f539,f18140,f1761])).
fof(f20847,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f461,f18140,f1761])).
fof(f20835,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f18140,f10419,f18140,f1759])).
fof(f20831,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f18140,f10419,f18140,f1758])).
fof(f20818,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f757,f18140,f1756])).
fof(f20815,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f439,f594,f18140,f1756])).
fof(f20810,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f536,f18140,f1755])).
fof(f20807,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f455,f18140,f1755])).
fof(f20798,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f534,f18140,f1754])).
fof(f20795,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f449,f18140,f1754])).
fof(f20786,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f532,f18140,f1753])).
fof(f20783,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f446,f18140,f1753])).
fof(f20717,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f562,f18140,f1751])).
fof(f20696,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f10652,f18140,f9507])).
fof(f20693,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f14160,f18140,f9506])).
fof(f9506,plain,(
  ( ! [X45,X44] : (sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X44,X45) | ~p(s(bool,X44))) ) | $spl4),
  inference(subsumption_resolution,[],[f9386,f247])).
fof(f9386,plain,(
  ( ! [X45,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X44,X45)) ) | $spl4),
  inference(superposition,[],[f419,f7026])).
fof(f20691,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f539,f18140,f9505])).
fof(f20689,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f461,f18140,f9505])).
fof(f20682,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13974,f18140,f9503])).
fof(f9503,plain,(
  ( ! [X37,X36] : (sP10('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X36,X37) | ~p(s(bool,X36))) ) | $spl4),
  inference(subsumption_resolution,[],[f9382,f247])).
fof(f9382,plain,(
  ( ! [X37,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | sP10('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X36,X37)) ) | $spl4),
  inference(superposition,[],[f402,f7026])).
fof(f20680,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13974,f18140,f9502])).
fof(f9502,plain,(
  ( ! [X35,X34] : (sP10('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X35,X34) | ~p(s(bool,X34))) ) | $spl4),
  inference(subsumption_resolution,[],[f9381,f247])).
fof(f9381,plain,(
  ( ! [X35,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X34)) | sP10('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X35,X34)) ) | $spl4),
  inference(superposition,[],[f401,f7026])).
fof(f20678,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f757,f18140,f9500])).
fof(f9500,plain,(
  ( ! [X30,X31] : (~sP8(X31,X30,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | p(s(bool,X31)) | ~p(s(bool,X30))) ) | $spl4),
  inference(subsumption_resolution,[],[f9379,f247])).
fof(f9379,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP8(X31,X30,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) ) | $spl4),
  inference(superposition,[],[f388,f7026])).
fof(f20676,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f594,f18140,f9500])).
fof(f20673,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f536,f18140,f9499])).
fof(f20671,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f455,f18140,f9499])).
fof(f20664,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f534,f18140,f9498])).
fof(f9498,plain,(
  ( ! [X26,X27] : (~sP4(X26,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X27) | p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl4),
  inference(subsumption_resolution,[],[f9377,f247])).
fof(f9377,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | p(s(bool,X27)) | ~sP4(X26,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)),X27)) ) | $spl4),
  inference(superposition,[],[f360,f7026])).
fof(f20662,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f449,f18140,f9498])).
fof(f20655,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f532,f18140,f9497])).
fof(f20653,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f446,f18140,f9497])).
fof(f20646,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f562,f18140,f9495])).
fof(f9495,plain,(
  ( ! [X21,X20] : (sP2(X20,X21,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14))) | ~p(s(bool,X20))) ) | $spl4),
  inference(subsumption_resolution,[],[f9374,f247])).
fof(f9374,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | sP2(X20,X21,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK13,sK18),s(sK13,sK14)))) ) | $spl4),
  inference(superposition,[],[f344,f7026])).
fof(f20639,plain,(
  $false | ($spl1 | $spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f7894,f18140,f5409])).
fof(f20637,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f6556,f18140,f5409])).
fof(f20634,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f11814,f18140,f5408])).
fof(f5408,plain,(
  ( ! [X35,X36] : (sP12(f(s(bool,'T_0')),X35,X36) | ~p(s(bool,X35))) ) | $spl61),
  inference(subsumption_resolution,[],[f5350,f247])).
fof(f5350,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | sP12(f(s(bool,'T_0')),X35,X36)) ) | $spl61),
  inference(superposition,[],[f419,f5088])).
fof(f20632,plain,(
  $false | ($spl1 | $spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f539,f18140,f5406])).
fof(f20630,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f461,f18140,f5406])).
fof(f20623,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f11519,f18140,f5404])).
fof(f5404,plain,(
  ( ! [X28,X27] : (sP10(f(s(bool,'T_0')),X27,X28) | ~p(s(bool,X27))) ) | $spl61),
  inference(subsumption_resolution,[],[f5346,f247])).
fof(f5346,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | sP10(f(s(bool,'T_0')),X27,X28)) ) | $spl61),
  inference(superposition,[],[f402,f5088])).
fof(f20621,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f11519,f18140,f5403])).
fof(f5403,plain,(
  ( ! [X26,X25] : (sP10(f(s(bool,'T_0')),X26,X25) | ~p(s(bool,X25))) ) | $spl61),
  inference(subsumption_resolution,[],[f5345,f247])).
fof(f5345,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X25)) | sP10(f(s(bool,'T_0')),X26,X25)) ) | $spl61),
  inference(superposition,[],[f401,f5088])).
fof(f20619,plain,(
  $false | ($spl1 | $spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f757,f18140,f5401])).
fof(f5401,plain,(
  ( ! [X21,X22] : (~sP8(X22,X21,f(s(bool,'T_0'))) | p(s(bool,X22)) | ~p(s(bool,X21))) ) | $spl61),
  inference(subsumption_resolution,[],[f5343,f247])).
fof(f5343,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | p(s(bool,X22)) | ~sP8(X22,X21,f(s(bool,'T_0')))) ) | $spl61),
  inference(superposition,[],[f388,f5088])).
fof(f20617,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f594,f18140,f5401])).
fof(f20614,plain,(
  $false | ($spl1 | $spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f536,f18140,f5400])).
fof(f20612,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f455,f18140,f5400])).
fof(f20605,plain,(
  $false | ($spl1 | $spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f534,f18140,f5399])).
fof(f5399,plain,(
  ( ! [X17,X18] : (~sP4(X17,f(s(bool,'T_0')),X18) | p(s(bool,X18)) | ~p(s(bool,X17))) ) | $spl61),
  inference(subsumption_resolution,[],[f5341,f247])).
fof(f5341,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~sP4(X17,f(s(bool,'T_0')),X18)) ) | $spl61),
  inference(superposition,[],[f360,f5088])).
fof(f20603,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f449,f18140,f5399])).
fof(f20596,plain,(
  $false | ($spl1 | $spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f532,f18140,f5398])).
fof(f5398,plain,(
  ( ! [X15,X16] : (~sP5(X15,f(s(bool,'T_0')),X16) | p(s(bool,X16)) | ~p(s(bool,X15))) ) | $spl61),
  inference(subsumption_resolution,[],[f5340,f247])).
fof(f5340,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | p(s(bool,X16)) | ~sP5(X15,f(s(bool,'T_0')),X16)) ) | $spl61),
  inference(superposition,[],[f356,f5088])).
fof(f20594,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f446,f18140,f5398])).
fof(f20587,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f562,f18140,f5396])).
fof(f5396,plain,(
  ( ! [X12,X11] : (sP2(X11,X12,f(s(bool,'T_0'))) | ~p(s(bool,X11))) ) | $spl61),
  inference(subsumption_resolution,[],[f5338,f247])).
fof(f5338,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | sP2(X11,X12,f(s(bool,'T_0')))) ) | $spl61),
  inference(superposition,[],[f344,f5088])).
fof(f20580,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f1536,f18140,f4524])).
fof(f4524,plain,(
  ( ! [X37,X38] : (~sP12('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X38,X37) | p(s(bool,X38)) | ~p(s(bool,X37))) ) | $spl4),
  inference(subsumption_resolution,[],[f4476,f247])).
fof(f4476,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP12('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X38,X37)) ) | $spl4),
  inference(superposition,[],[f420,f1314])).
fof(f20578,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f1535,f18140,f4524])).
fof(f20575,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1533,f18140,f4523])).
fof(f4523,plain,(
  ( ! [X35,X36] : (sP12('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X35,X36) | ~p(s(bool,X35))) ) | $spl4),
  inference(subsumption_resolution,[],[f4475,f247])).
fof(f4475,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | sP12('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X35,X36)) ) | $spl4),
  inference(superposition,[],[f419,f1314])).
fof(f20573,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f539,f18140,f4521])).
fof(f4521,plain,(
  ( ! [X31,X32] : (~sP11(X32,X31,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ~p(s(bool,X32)) | p(s(bool,X31))) ) | $spl4),
  inference(subsumption_resolution,[],[f4473,f247])).
fof(f4473,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~sP11(X32,X31,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | $spl4),
  inference(superposition,[],[f413,f1314])).
fof(f20571,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f461,f18140,f4521])).
fof(f20564,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1480,f18140,f4519])).
fof(f4519,plain,(
  ( ! [X28,X27] : (sP10('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X27,X28) | ~p(s(bool,X27))) ) | $spl4),
  inference(subsumption_resolution,[],[f4471,f247])).
fof(f4471,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | sP10('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X27,X28)) ) | $spl4),
  inference(superposition,[],[f402,f1314])).
fof(f20562,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1480,f18140,f4518])).
fof(f4518,plain,(
  ( ! [X26,X25] : (sP10('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X26,X25) | ~p(s(bool,X25))) ) | $spl4),
  inference(subsumption_resolution,[],[f4470,f247])).
fof(f4470,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X25)) | sP10('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X26,X25)) ) | $spl4),
  inference(superposition,[],[f401,f1314])).
fof(f20557,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f757,f18140,f4516])).
fof(f4516,plain,(
  ( ! [X21,X22] : (~sP8(X22,X21,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | p(s(bool,X22)) | ~p(s(bool,X21))) ) | $spl4),
  inference(subsumption_resolution,[],[f4468,f247])).
fof(f4468,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | p(s(bool,X22)) | ~sP8(X22,X21,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | $spl4),
  inference(superposition,[],[f388,f1314])).
fof(f20555,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f594,f18140,f4516])).
fof(f20550,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f536,f18140,f4515])).
fof(f4515,plain,(
  ( ! [X19,X20] : (~sP7(X19,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X20) | p(s(bool,X20)) | ~p(s(bool,X19))) ) | $spl4),
  inference(subsumption_resolution,[],[f4467,f247])).
fof(f4467,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | p(s(bool,X20)) | ~sP7(X19,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X20)) ) | $spl4),
  inference(superposition,[],[f379,f1314])).
fof(f20548,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f455,f18140,f4515])).
fof(f20541,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f534,f18140,f4514])).
fof(f4514,plain,(
  ( ! [X17,X18] : (~sP4(X17,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X18) | p(s(bool,X18)) | ~p(s(bool,X17))) ) | $spl4),
  inference(subsumption_resolution,[],[f4466,f247])).
fof(f4466,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~sP4(X17,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X18)) ) | $spl4),
  inference(superposition,[],[f360,f1314])).
fof(f20539,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f449,f18140,f4514])).
fof(f20532,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f532,f18140,f4513])).
fof(f4513,plain,(
  ( ! [X15,X16] : (~sP5(X15,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X16) | p(s(bool,X16)) | ~p(s(bool,X15))) ) | $spl4),
  inference(subsumption_resolution,[],[f4465,f247])).
fof(f4465,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | p(s(bool,X16)) | ~sP5(X15,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),X16)) ) | $spl4),
  inference(superposition,[],[f356,f1314])).
fof(f20530,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f446,f18140,f4513])).
fof(f20523,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f562,f18140,f4510])).
fof(f4510,plain,(
  ( ! [X12,X11] : (sP2(X11,X12,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ~p(s(bool,X11))) ) | $spl4),
  inference(subsumption_resolution,[],[f4463,f247])).
fof(f4463,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | sP2(X11,X12,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))))) ) | $spl4),
  inference(superposition,[],[f344,f1314])).
fof(f20521,plain,(
  $false | ($spl40 | $spl112)),
  inference(unit_resulting_resolution,[],[f648,f439,f18140,f2691])).
fof(f20513,plain,(
  $false | ($spl38 | $spl112)),
  inference(unit_resulting_resolution,[],[f1144,f18140,f2688])).
fof(f20511,plain,(
  $false | ($spl34 | $spl112)),
  inference(unit_resulting_resolution,[],[f461,f439,f18140,f2682])).
fof(f20503,plain,(
  $false | ($spl30 | $spl112)),
  inference(unit_resulting_resolution,[],[f1003,f18140,f2676])).
fof(f20501,plain,(
  $false | ($spl28 | $spl112)),
  inference(unit_resulting_resolution,[],[f1003,f18140,f2673])).
fof(f20499,plain,(
  $false | ($spl24 | $spl112)),
  inference(unit_resulting_resolution,[],[f594,f439,f18140,f2667])).
fof(f20487,plain,(
  $false | ($spl22 | $spl112)),
  inference(unit_resulting_resolution,[],[f455,f439,f18140,f2664])).
fof(f20479,plain,(
  $false | ($spl20 | $spl112)),
  inference(unit_resulting_resolution,[],[f449,f439,f18140,f2661])).
fof(f20471,plain,(
  $false | ($spl18 | $spl112)),
  inference(unit_resulting_resolution,[],[f446,f439,f18140,f2658])).
fof(f20461,plain,(
  $false | ($spl14 | $spl112)),
  inference(unit_resulting_resolution,[],[f562,f18140,f2651])).
fof(f20458,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f10652,f18140,f1795])).
fof(f20454,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f18140,f14160,f18140,f1794])).
fof(f20451,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f539,f18140,f1792])).
fof(f20448,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f461,f18140,f1792])).
fof(f20440,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f18140,f13974,f18140,f1790])).
fof(f20437,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f18140,f13974,f18140,f1789])).
fof(f20434,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f757,f18140,f1787])).
fof(f20431,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f594,f18140,f1787])).
fof(f20427,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f536,f18140,f1786])).
fof(f20424,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f455,f18140,f1786])).
fof(f20416,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f534,f18140,f1785])).
fof(f20413,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f449,f18140,f1785])).
fof(f20405,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f532,f18140,f1784])).
fof(f20402,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f446,f18140,f1784])).
fof(f20394,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f562,f18140,f1782])).
fof(f20388,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f5056,f18140,f18140,f1764])).
fof(f20384,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10565,f18140,f18140,f1763])).
fof(f20381,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f539,f18140,f18140,f1761])).
fof(f20379,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f461,f18140,f18140,f1761])).
fof(f20370,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10419,f18140,f18140,f1759])).
fof(f20367,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10419,f18140,f18140,f1758])).
fof(f20364,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f757,f18140,f18140,f1756])).
fof(f20362,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f594,f18140,f18140,f1756])).
fof(f20357,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f536,f18140,f18140,f1755])).
fof(f20355,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f455,f18140,f18140,f1755])).
fof(f20346,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f534,f18140,f18140,f1754])).
fof(f20344,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f449,f18140,f18140,f1754])).
fof(f20335,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f532,f18140,f18140,f1753])).
fof(f20333,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f446,f18140,f18140,f1753])).
fof(f20324,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f562,f18140,f18140,f1751])).
fof(f20315,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f539,f18140,f1628])).
fof(f20312,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f461,f18140,f1628])).
fof(f20302,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f757,f18140,f1623])).
fof(f20299,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f594,f18140,f1623])).
fof(f20295,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f536,f18140,f1622])).
fof(f20292,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f455,f18140,f1622])).
fof(f20284,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f534,f18140,f1621])).
fof(f20281,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f449,f18140,f1621])).
fof(f20273,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f18140,f532,f18140,f1620])).
fof(f20270,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140,f446,f18140,f1620])).
fof(f20262,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f562,f18140,f1618])).
fof(f20257,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f5056,f18140,f1180])).
fof(f20250,plain,(
  $false | ($spl1 | $spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f17640,f18140,f1179])).
fof(f20248,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f1015,f18140,f1179])).
fof(f20245,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f539,f18140,f1140])).
fof(f20243,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f461,f18140,f1140])).
fof(f20236,plain,(
  $false | ($spl1 | $spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f539,f18140,f1139])).
fof(f20234,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f461,f18140,f1139])).
fof(f20227,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10565,f18140,f1101])).
fof(f1101,plain,(
  ( ! [X31,X32] : (sP12(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X31,X32) | ~p(s(bool,X31))) ) | $spl4),
  inference(subsumption_resolution,[],[f1079,f247])).
fof(f1079,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | sP12(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X31,X32)) ) | $spl4),
  inference(superposition,[],[f419,f491])).
fof(f20225,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10419,f18140,f1098])).
fof(f1098,plain,(
  ( ! [X26,X25] : (sP10(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X25,X26) | ~p(s(bool,X25))) ) | $spl4),
  inference(subsumption_resolution,[],[f1076,f247])).
fof(f1076,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X25)) | sP10(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X25,X26)) ) | $spl4),
  inference(superposition,[],[f402,f491])).
fof(f20223,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f10419,f18140,f1097])).
fof(f1097,plain,(
  ( ! [X24,X23] : (sP10(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X24,X23) | ~p(s(bool,X23))) ) | $spl4),
  inference(subsumption_resolution,[],[f1075,f247])).
fof(f1075,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | sP10(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X24,X23)) ) | $spl4),
  inference(superposition,[],[f401,f491])).
fof(f20221,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f757,f18140,f1095])).
fof(f1095,plain,(
  ( ! [X19,X20] : (~sP8(X20,X19,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | p(s(bool,X20)) | ~p(s(bool,X19))) ) | $spl4),
  inference(subsumption_resolution,[],[f1073,f247])).
fof(f1073,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | p(s(bool,X20)) | ~sP8(X20,X19,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) ) | $spl4),
  inference(superposition,[],[f388,f491])).
fof(f20219,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f594,f18140,f1095])).
fof(f20216,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f536,f18140,f1094])).
fof(f20214,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f455,f18140,f1094])).
fof(f20207,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f534,f18140,f1093])).
fof(f1093,plain,(
  ( ! [X15,X16] : (~sP4(X15,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X16) | p(s(bool,X16)) | ~p(s(bool,X15))) ) | $spl4),
  inference(subsumption_resolution,[],[f1071,f247])).
fof(f1071,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | p(s(bool,X16)) | ~sP4(X15,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X16)) ) | $spl4),
  inference(superposition,[],[f360,f491])).
fof(f20205,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f449,f18140,f1093])).
fof(f20198,plain,(
  $false | ($spl1 | $spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f532,f18140,f1092])).
fof(f20196,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f446,f18140,f1092])).
fof(f20189,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f562,f18140,f1091])).
fof(f1091,plain,(
  ( ! [X12,X11] : (sP2(X11,X12,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14))) | ~p(s(bool,X11))) ) | $spl4),
  inference(subsumption_resolution,[],[f1069,f247])).
fof(f1069,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | sP2(X11,X12,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) ) | $spl4),
  inference(superposition,[],[f344,f491])).
fof(f20187,plain,(
  $false | ($spl1 | $spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f757,f18140,f947])).
fof(f947,plain,(
  ( ! [X0,X1] : (~sP8(X1,X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f940,f247])).
fof(f940,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP8(X1,X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) ) | $spl6),
  inference(superposition,[],[f388,f490])).
fof(f20185,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f594,f18140,f947])).
fof(f20182,plain,(
  $false | ($spl1 | $spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f536,f18140,f916])).
fof(f20180,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f455,f18140,f916])).
fof(f20173,plain,(
  $false | ($spl1 | $spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f534,f18140,f865])).
fof(f865,plain,(
  ( ! [X0,X1] : (~sP4(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f858,f247])).
fof(f858,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP4(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X1)) ) | $spl6),
  inference(superposition,[],[f360,f490])).
fof(f20171,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f449,f18140,f865])).
fof(f20164,plain,(
  $false | ($spl1 | $spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f532,f18140,f843])).
fof(f843,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f836,f247])).
fof(f836,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X1)) ) | $spl6),
  inference(superposition,[],[f356,f490])).
fof(f20162,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f439,f446,f18140,f843])).
fof(f20155,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f7822,f18140,f658])).
fof(f658,plain,(
  ( ! [X0,X1] : (sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,X1) | ~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f657,f247])).
fof(f657,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP12('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,X1)) ) | $spl6),
  inference(superposition,[],[f419,f490])).
fof(f20153,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f7630,f18140,f640])).
fof(f640,plain,(
  ( ! [X0,X1] : (sP10('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,X1) | ~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f639,f247])).
fof(f639,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP10('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X0,X1)) ) | $spl6),
  inference(superposition,[],[f402,f490])).
fof(f20151,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f7630,f18140,f635])).
fof(f635,plain,(
  ( ! [X10,X11] : (sP10('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X11,X10) | ~p(s(bool,X10))) ) | $spl6),
  inference(subsumption_resolution,[],[f626,f247])).
fof(f626,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | sP10('const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X11,X10)) ) | $spl6),
  inference(superposition,[],[f401,f490])).
fof(f20149,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f562,f18140,f634])).
fof(f634,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18))) | ~p(s(bool,X8))) ) | $spl6),
  inference(subsumption_resolution,[],[f625,f247])).
fof(f625,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | sP2(X8,X9,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) ) | $spl6),
  inference(superposition,[],[f344,f490])).
fof(f20147,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f557,f18140,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f341])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f206])).
fof(f20144,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f600,f458,f18140,f424])).
fof(f20142,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1450,f1469,f18140,f424])).
fof(f20140,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f600,f484,f18140,f424])).
fof(f484,plain,(
  ( ! [X0,X1] : (sP9(X0,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f436,f398])).
fof(f20138,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f600,f472,f18140,f424])).
fof(f472,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f438,f398])).
fof(f20136,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f604,f457,f18140,f424])).
fof(f20134,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1445,f1468,f18140,f424])).
fof(f20132,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f604,f483,f18140,f424])).
fof(f483,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f436,f397])).
fof(f20130,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f604,f471,f18140,f424])).
fof(f471,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f438,f397])).
fof(f20128,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f600,f456,f18140,f424])).
fof(f20115,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f644,f461,f18140,f423])).
fof(f20113,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1481,f1519,f18140,f423])).
fof(f20111,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f644,f485,f18140,f423])).
fof(f485,plain,(
  ( ! [X0,X1] : (sP11(X0,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f436,f415])).
fof(f20109,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f644,f473,f18140,f423])).
fof(f473,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f438,f415])).
fof(f20107,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f645,f460,f18140,f423])).
fof(f20105,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f645,f538,f18140,f423])).
fof(f20103,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f644,f459,f18140,f423])).
fof(f20086,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f648,f439,f18140,f18140,f420])).
fof(f20052,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1116,f439,f247,f18140,f420])).
fof(f20016,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f648,f439,f18140,f18140,f420])).
fof(f19982,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f648,f439,f247,f18140,f420])).
fof(f19960,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1144,f18140,f18140,f419])).
fof(f19958,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f645,f247,f18140,f419])).
fof(f19941,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1144,f18140,f18140,f419])).
fof(f19939,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1144,f247,f18140,f419])).
fof(f19922,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f644,f439,f18140,f418])).
fof(f19905,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f644,f439,f439,f18140,f417])).
fof(f19887,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1130,f18140,f415])).
fof(f19885,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f459,f18140,f439,f18140,f413])).
fof(f19883,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f459,f247,f439,f18140,f413])).
fof(f19863,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f459,f18140,f439,f18140,f413])).
fof(f19861,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f461,f247,f439,f18140,f413])).
fof(f19840,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f644,f461,f439,f18140,f412])).
fof(f19837,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f644,f459,f439,f18140,f412])).
fof(f19809,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f644,f461,f439,f18140,f411])).
fof(f19806,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f644,f459,f439,f18140,f411])).
fof(f19778,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f644,f461,f18140,f18140,f410])).
fof(f19776,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1481,f1519,f18140,f18140,f410])).
fof(f19774,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f644,f485,f18140,f18140,f410])).
fof(f19772,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f644,f473,f18140,f18140,f410])).
fof(f19770,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f645,f460,f18140,f18140,f410])).
fof(f19768,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f645,f538,f18140,f18140,f410])).
fof(f19766,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f644,f459,f18140,f18140,f410])).
fof(f19719,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f461,f1481,f18140,f18140,f410])).
fof(f19717,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f461,f645,f18140,f18140,f410])).
fof(f19715,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f461,f644,f18140,f18140,f410])).
fof(f19694,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f461,f1481,f247,f18140,f410])).
fof(f19692,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f461,f645,f247,f18140,f410])).
fof(f19674,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1116,f439,f18140,f407])).
fof(f19640,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f972,f439,f439,f18140,f403])).
fof(f972,plain,(
  sP10('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f458,f439,f439,f395])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f230])).
fof(f19598,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1003,f18140,f18140,f402])).
fof(f19596,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f604,f247,f18140,f402])).
fof(f19584,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1003,f18140,f18140,f402])).
fof(f19582,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1003,f247,f18140,f402])).
fof(f19570,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1003,f18140,f18140,f401])).
fof(f19568,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f600,f247,f18140,f401])).
fof(f19556,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1003,f18140,f18140,f401])).
fof(f19554,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1003,f247,f18140,f401])).
fof(f19542,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f990,f18140,f398])).
fof(f19540,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f990,f18140,f397])).
fof(f19538,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f458,f439,f439,f18140,f396])).
fof(f19523,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f600,f458,f439,f18140,f394])).
fof(f19521,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1445,f1468,f439,f18140,f394])).
fof(f19517,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f600,f456,f439,f18140,f394])).
fof(f19484,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f458,f1450,f439,f18140,f393])).
fof(f19482,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f458,f604,f439,f18140,f393])).
fof(f19459,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f972,f439,f18140,f390])).
fof(f19419,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f972,f439,f18140,f389])).
fof(f19379,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f586,f439,f18140,f18140,f388])).
fof(f19360,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f586,f439,f247,f18140,f388])).
fof(f19344,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f586,f439,f18140,f18140,f388])).
fof(f19325,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f594,f439,f247,f18140,f388])).
fof(f19309,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f578,f18140,f18140,f18140,f385])).
fof(f19307,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f581,f247,f18140,f18140,f385])).
fof(f19293,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f578,f18140,f247,f18140,f385])).
fof(f19291,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f920,f247,f247,f18140,f385])).
fof(f19265,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f578,f18140,f18140,f18140,f385])).
fof(f19263,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f578,f247,f18140,f18140,f385])).
fof(f19249,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f578,f18140,f247,f18140,f385])).
fof(f19247,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f578,f247,f247,f18140,f385])).
fof(f19223,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f578,f18140,f18140,f18140,f385])).
fof(f19221,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f578,f247,f18140,f18140,f385])).
fof(f19207,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f581,f18140,f247,f18140,f385])).
fof(f19205,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f581,f247,f247,f18140,f385])).
fof(f19185,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f908,f18140,f380])).
fof(f19183,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f454,f439,f18140,f18140,f379])).
fof(f19173,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f454,f439,f247,f18140,f379])).
fof(f19161,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f454,f439,f18140,f18140,f379])).
fof(f19151,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f455,f439,f247,f18140,f379])).
fof(f19139,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f455,f1393,f18140,f18140,f375])).
fof(f19137,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f454,f1392,f18140,f18140,f375])).
fof(f19127,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f455,f581,f18140,f18140,f375])).
fof(f19125,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f454,f578,f18140,f18140,f375])).
fof(f19115,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f454,f1392,f247,f18140,f375])).
fof(f19107,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f454,f578,f247,f18140,f375])).
fof(f19089,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f455,f1393,f18140,f18140,f375])).
fof(f19087,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f454,f1392,f18140,f18140,f375])).
fof(f19077,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f455,f581,f18140,f18140,f375])).
fof(f19075,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f454,f578,f18140,f18140,f375])).
fof(f19065,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f455,f1393,f247,f18140,f375])).
fof(f19061,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f455,f581,f247,f18140,f375])).
fof(f19043,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f594,f439,f18140,f374])).
fof(f19024,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f586,f439,f18140,f373])).
fof(f19005,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1195,f523,f18140,f370])).
fof(f19003,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f785,f487,f18140,f370])).
fof(f19001,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1290,f523,f18140,f370])).
fof(f18999,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f786,f487,f18140,f370])).
fof(f18931,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1599,f289,f18140,f369])).
fof(f18929,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1710,f301,f18140,f369])).
fof(f18927,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1971,f524,f18140,f369])).
fof(f18925,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1706,f1314,f18140,f369])).
fof(f18923,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1052,f491,f18140,f369])).
fof(f18921,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f9338,f7026,f18140,f369])).
fof(f18919,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f815,f490,f18140,f369])).
fof(f18917,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f5313,f5088,f18140,f369])).
fof(f18914,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1597,f289,f18140,f369])).
fof(f18912,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1708,f301,f18140,f369])).
fof(f18910,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f1969,f524,f18140,f369])).
fof(f18908,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1701,f1314,f18140,f369])).
fof(f18906,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1050,f491,f18140,f369])).
fof(f18904,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f9336,f7026,f18140,f369])).
fof(f18902,plain,(
  $false | ($spl6 | $spl112)),
  inference(unit_resulting_resolution,[],[f813,f490,f18140,f369])).
fof(f18900,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f5311,f5088,f18140,f369])).
fof(f21274,plain,(
  $false | ($spl4 | $spl112)),
  inference(trivial_inequality_removal,[],[f18894])).
fof(f18894,plain,(
  s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) != s(bool,'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f13037,f18140,f369])).
fof(f13037,plain,(
  ~sP6('F_0','|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1376,f350])).
fof(f1376,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'|dollar|exists_1'(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18)))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1309,f439,f1309,f360])).
fof(f21273,plain,(
  $false | $spl112),
  inference(trivial_inequality_removal,[],[f18882])).
fof(f18882,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl112),
  inference(unit_resulting_resolution,[],[f890,f18140,f369])).
fof(f890,plain,(
  ~sP6('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f877,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f211])).
fof(f877,plain,(
  ~sP3('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f439,f439,f439,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f219])).
fof(f21272,plain,(
  $false | $spl112),
  inference(trivial_inequality_removal,[],[f18876])).
fof(f18876,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl112),
  inference(unit_resulting_resolution,[],[f869,f18140,f369])).
fof(f869,plain,(
  ~sP6('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f857,f350])).
fof(f857,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f439,f247,f247,f360])).
fof(f18864,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f877,f18140,f367])).
fof(f18862,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f877,f18140,f366])).
fof(f18860,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f877,f18140,f365])).
fof(f18858,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f857,f18140,f361])).
fof(f18856,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f448,f439,f18140,f18140,f360])).
fof(f18846,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f448,f439,f247,f18140,f360])).
fof(f18834,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f448,f439,f18140,f18140,f360])).
fof(f18824,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f449,f439,f247,f18140,f360])).
fof(f18812,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f835,f18140,f357])).
fof(f835,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f439,f247,f247,f356])).
fof(f18810,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f445,f439,f18140,f18140,f356])).
fof(f18798,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f445,f439,f247,f18140,f356])).
fof(f18784,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f445,f439,f18140,f18140,f356])).
fof(f18772,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f446,f439,f247,f18140,f356])).
fof(f18758,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1384,f448,f445,f1343,f18140,f353])).
fof(f18752,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f448,f445,f1390,f18140,f353])).
fof(f18727,plain,(
  $false | ($spl4 | $spl112)),
  inference(unit_resulting_resolution,[],[f1384,f448,f445,f1342,f18140,f353])).
fof(f18722,plain,(
  $false | ($spl61 | $spl112)),
  inference(unit_resulting_resolution,[],[f450,f448,f445,f5084,f18140,f353])).
fof(f18712,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f451,f533,f531,f1951,f18140,f353])).
fof(f18708,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f451,f448,f445,f806,f18140,f353])).
fof(f18459,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1227,f18140,f439,f18140,f352])).
fof(f18457,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1227,f247,f439,f18140,f352])).
fof(f18360,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f1227,f439,f18140,f18140,f352])).
fof(f18293,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f786,f439,f247,f18140,f352])).
fof(f18266,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f557,f18140,f18140,f344])).
fof(f18264,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f557,f247,f18140,f344])).
fof(f18260,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f557,f18140,f18140,f344])).
fof(f18258,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f562,f247,f18140,f344])).
fof(f18254,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f557,f18140,f18140,f343])).
fof(f18252,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f557,f247,f18140,f343])).
fof(f18248,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f557,f18140,f18140,f343])).
fof(f18243,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f557,f18140,f18140,f342])).
fof(f18238,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f557,f18140,f18140,f342])).
fof(f18236,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f562,f247,f18140,f342])).
fof(f18232,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f498,f439,f18140,f333])).
fof(f498,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f439,f439,f328])).
fof(f328,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f201])).
fof(f18222,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f498,f439,f18140,f330])).
fof(f18212,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f502,f18140,f18140,f329])).
fof(f18210,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f502,f247,f18140,f329])).
fof(f18194,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f502,f18140,f18140,f329])).
fof(f18192,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f507,f247,f18140,f329])).
fof(f18176,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f516,f18140,f315])).
fof(f315,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f516,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f439,f439,f441])).
fof(f441,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f313,f439])).
fof(f313,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f194])).
fof(f18174,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f516,f18140,f314])).
fof(f314,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f194])).
fof(f18172,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f18140,f523,f18140,f308])).
fof(f18170,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f487,f18140,f308])).
fof(f18168,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f247,f523,f18140,f308])).
fof(f18166,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f247,f487,f18140,f308])).
fof(f18162,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f18140,f523,f18140,f308])).
fof(f18160,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f18140,f487,f18140,f308])).
fof(f18158,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f247,f523,f18140,f308])).
fof(f18156,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f247,f487,f18140,f308])).
fof(f18150,plain,(
  $false | ($spl1 | $spl112)),
  inference(unit_resulting_resolution,[],[f524,f18140,f266])).
fof(f18148,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f487,f18140,f263])).
fof(f18146,plain,(
  $false | $spl112),
  inference(unit_resulting_resolution,[],[f439,f18140])).
fof(f18142,plain,(
  ( ! [X39,X38] : (sP1(X38,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X39)) ) | $spl114),
  inference(cnf_transformation,[],[f18142_D])).
fof(f18142_D,plain,(
  ( ! [X39,X38] : (sP1(X38,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X39)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f18141,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | sP1(X38,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | $spl1),
  inference(subsumption_resolution,[],[f18117,f262])).
fof(f18117,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | sP1(X38,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X37)) ) | $spl1),
  inference(superposition,[],[f2702,f2047])).
fof(f2702,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f2331,f247])).
fof(f2331,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f330,f711])).
fof(f18139,plain,(
  ( ! [X31,X32] : (sP1(X31,'F_0') | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) | $spl110),
  inference(cnf_transformation,[],[f18139_D])).
fof(f18139_D,plain,(
  ( ! [X31,X32] : (sP1(X31,'F_0') | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f18138,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | sP1(X31,'F_0') | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f18114,f262])).
fof(f18114,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | sP1(X31,'F_0') | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f2702,f255])).
fof(f13017,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,sK20('T_0',happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3)),X1))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))))) ) | $spl108),
  inference(cnf_transformation,[],[f13017_D])).
fof(f13017_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,sK20('T_0',happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3)),X1))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f2645,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f2645_D])).
fof(f2645_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f13019,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,sK20('T_0',happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3)),X1))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))))) )),
  inference(condensation,[],[f13018])).
fof(f13018,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X41,fun(X41,bool)),X42),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X41,X43),s(X41,sK20('T_0',happ(s(fun(X41,fun(X41,bool)),X42),s(X41,X43)),X41))))) | p(s(bool,happ(s(fun(X41,bool),happ(s(fun(X41,fun(X41,bool)),X42),s(X41,X43))),s(X41,X44)))) | s(bool,X40) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f12994,f769])).
fof(f12994,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X41,fun(X41,bool)),X42),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X41,X43),s(X41,sK20('T_0',happ(s(fun(X41,fun(X41,bool)),X42),s(X41,X43)),X41))))) | p(s(bool,happ(s(fun(X41,bool),happ(s(fun(X41,fun(X41,bool)),X42),s(X41,X43))),s(X41,X44)))) | s(bool,X40) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f1734,f711])).
fof(f1734,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,X15)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X12,fun(X12,bool)),X13),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X12,X14),s(X12,sK20(X15,happ(s(fun(X12,fun(X12,bool)),X13),s(X12,X14)),X12))))) | p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X12,fun(X12,bool)),X13),s(X12,X14))),s(X12,X16))))) )),
  inference(superposition,[],[f295,f301])).
fof(f295,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0))))) | ~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) | ~p(s(bool,X2)) | ! [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0))))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20])],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2)) | ! [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2)) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (((? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1,X2] : ((! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X13,X17] : ((! [X8] : p(s(bool,happ(s(fun(X6,bool),X13),s(X6,X8)))) & p(s(bool,X17))) <=> ! [X8] : (p(s(bool,happ(s(fun(X6,bool),X13),s(X6,X8)))) & p(s(bool,X17))))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.bool.LEFT_AND_FORALL_THM')).
fof(f13016,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,sK20('T_0',happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3)),X1))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))))) )),
  inference(condensation,[],[f13015])).
fof(f13015,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,sK20('T_0',happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37)),X35))))) | p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | s(bool,X34) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f12993,f769])).
fof(f12993,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,sK20('T_0',happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37)),X35))))) | p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | s(bool,X34) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f1734,f711])).
fof(f12967,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,sK19('T_0',happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3)),X1))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))))) ) | $spl106),
  inference(cnf_transformation,[],[f12967_D])).
fof(f12967_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,sK19('T_0',happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3)),X1))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f12969,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,sK19('T_0',happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3)),X1))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))))) )),
  inference(condensation,[],[f12968])).
fof(f12968,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X41,fun(X41,bool)),X42),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X41,X43),s(X41,sK19('T_0',happ(s(fun(X41,fun(X41,bool)),X42),s(X41,X43)),X41))))) | p(s(bool,happ(s(fun(X41,bool),happ(s(fun(X41,fun(X41,bool)),X42),s(X41,X43))),s(X41,X44)))) | s(bool,X40) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f12944,f769])).
fof(f12944,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X41,fun(X41,bool)),X42),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X41,X43),s(X41,sK19('T_0',happ(s(fun(X41,fun(X41,bool)),X42),s(X41,X43)),X41))))) | p(s(bool,happ(s(fun(X41,bool),happ(s(fun(X41,fun(X41,bool)),X42),s(X41,X43))),s(X41,X44)))) | s(bool,X40) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f1733,f711])).
fof(f1733,plain,(
  ( ! [X10,X8,X7,X11,X9] : (~p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X7,fun(X7,bool)),X8),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X7,X9),s(X7,sK19(X10,happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X9)),X7))))) | p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X9))),s(X7,X11))))) )),
  inference(superposition,[],[f293,f301])).
fof(f293,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) | ~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f12966,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,sK19('T_0',happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3)),X1))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))))) )),
  inference(condensation,[],[f12965])).
fof(f12965,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,sK19('T_0',happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37)),X35))))) | p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | s(bool,X34) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f12943,f769])).
fof(f12943,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,sK19('T_0',happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37)),X35))))) | p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | s(bool,X34) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f1733,f711])).
fof(f10910,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | ~sP12('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5,'T_0') | p(s(bool,X5))) ) | $spl104),
  inference(cnf_transformation,[],[f10910_D])).
fof(f10910_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | ~sP12('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5,'T_0') | p(s(bool,X5))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f10912,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | p(s(bool,X5)) | ~sP12('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5,'T_0')) )),
  inference(condensation,[],[f10911])).
fof(f10911,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | p(s(bool,X46)) | ~sP12('const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),X46,'T_0') | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10890,f769])).
fof(f10890,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | p(s(bool,X46)) | ~sP12('const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),X46,'T_0') | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1795,f711])).
fof(f10909,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | p(s(bool,X5)) | ~sP12('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5,'T_0')) )),
  inference(condensation,[],[f10908])).
fof(f10908,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | p(s(bool,X39)) | ~sP12('const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),X39,'T_0') | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f10889,f769])).
fof(f10889,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | p(s(bool,X39)) | ~sP12('const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),X39,'T_0') | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1795,f711])).
fof(f10867,plain,(
  ( ! [X4,X2,X5,X3,X1] : (p(s(bool,X1)) | ~sP11('T_0',X1,'const.arithmetic.NRC_4'(s(fun(X2,fun(X2,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X2,X4),s(X2,X5))) | ~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X3),s(X2,X4))),s(X2,X5))))) ) | $spl102),
  inference(cnf_transformation,[],[f10867_D])).
fof(f10867_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (p(s(bool,X1)) | ~sP11('T_0',X1,'const.arithmetic.NRC_4'(s(fun(X2,fun(X2,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X2,X4),s(X2,X5))) | ~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X3),s(X2,X4))),s(X2,X5))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f10869,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X3),s(X2,X4))),s(X2,X5)))) | ~sP11('T_0',X1,'const.arithmetic.NRC_4'(s(fun(X2,fun(X2,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X2,X4),s(X2,X5)))) )),
  inference(condensation,[],[f10868])).
fof(f10868,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X45))),s(X43,X46)))) | ~sP11('T_0',X42,'const.arithmetic.NRC_4'(s(fun(X43,fun(X43,bool)),X44),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X43,X45),s(X43,X46))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10847,f769])).
fof(f10847,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X45))),s(X43,X46)))) | ~sP11('T_0',X42,'const.arithmetic.NRC_4'(s(fun(X43,fun(X43,bool)),X44),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X43,X45),s(X43,X46))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1792,f711])).
fof(f10866,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X3),s(X2,X4))),s(X2,X5)))) | ~sP11('T_0',X1,'const.arithmetic.NRC_4'(s(fun(X2,fun(X2,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X2,X4),s(X2,X5)))) )),
  inference(condensation,[],[f10865])).
fof(f10865,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (p(s(bool,X35)) | ~p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X38))),s(X36,X39)))) | ~sP11('T_0',X35,'const.arithmetic.NRC_4'(s(fun(X36,fun(X36,bool)),X37),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X36,X38),s(X36,X39))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f10846,f769])).
fof(f10846,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | p(s(bool,X35)) | ~p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X38))),s(X36,X39)))) | ~sP11('T_0',X35,'const.arithmetic.NRC_4'(s(fun(X36,fun(X36,bool)),X37),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X36,X38),s(X36,X39))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1792,f711])).
fof(f10684,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | ~sP8(X5,'T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4))) | p(s(bool,X5))) ) | $spl100),
  inference(cnf_transformation,[],[f10684_D])).
fof(f10684_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | ~sP8(X5,'T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4))) | p(s(bool,X5))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f10686,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | p(s(bool,X5)) | ~sP8(X5,'T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) )),
  inference(condensation,[],[f10685])).
fof(f10685,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | p(s(bool,X46)) | ~sP8(X46,'T_0','const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10664,f769])).
fof(f10664,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | p(s(bool,X46)) | ~sP8(X46,'T_0','const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1787,f711])).
fof(f10683,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | p(s(bool,X5)) | ~sP8(X5,'T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) )),
  inference(condensation,[],[f10682])).
fof(f10682,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | p(s(bool,X39)) | ~sP8(X39,'T_0','const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f10663,f769])).
fof(f10663,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | p(s(bool,X39)) | ~sP8(X39,'T_0','const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1787,f711])).
fof(f10641,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | ~sP7('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5) | p(s(bool,X5))) ) | $spl98),
  inference(cnf_transformation,[],[f10641_D])).
fof(f10641_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | ~sP7('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5) | p(s(bool,X5))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f10643,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | p(s(bool,X5)) | ~sP7('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5)) )),
  inference(condensation,[],[f10642])).
fof(f10642,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | p(s(bool,X46)) | ~sP7('T_0','const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10621,f769])).
fof(f10621,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | p(s(bool,X46)) | ~sP7('T_0','const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1786,f711])).
fof(f10640,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | p(s(bool,X5)) | ~sP7('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5)) )),
  inference(condensation,[],[f10639])).
fof(f10639,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | p(s(bool,X39)) | ~sP7('T_0','const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f10620,f769])).
fof(f10620,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | p(s(bool,X39)) | ~sP7('T_0','const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1786,f711])).
fof(f10594,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | ~sP4('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5) | p(s(bool,X5))) ) | $spl96),
  inference(cnf_transformation,[],[f10594_D])).
fof(f10594_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | ~sP4('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5) | p(s(bool,X5))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f10596,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | p(s(bool,X5)) | ~sP4('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5)) )),
  inference(condensation,[],[f10595])).
fof(f10595,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | p(s(bool,X46)) | ~sP4('T_0','const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10574,f769])).
fof(f10574,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | p(s(bool,X46)) | ~sP4('T_0','const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1785,f711])).
fof(f10593,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | p(s(bool,X5)) | ~sP4('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5)) )),
  inference(condensation,[],[f10592])).
fof(f10592,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | p(s(bool,X39)) | ~sP4('T_0','const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f10573,f769])).
fof(f10573,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | p(s(bool,X39)) | ~sP4('T_0','const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1785,f711])).
fof(f10545,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | ~sP5('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5) | p(s(bool,X5))) ) | $spl94),
  inference(cnf_transformation,[],[f10545_D])).
fof(f10545_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | ~sP5('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5) | p(s(bool,X5))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f10547,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | p(s(bool,X5)) | ~sP5('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5)) )),
  inference(condensation,[],[f10546])).
fof(f10546,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | p(s(bool,X46)) | ~sP5('T_0','const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10525,f769])).
fof(f10525,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | p(s(bool,X46)) | ~sP5('T_0','const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1784,f711])).
fof(f10544,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | p(s(bool,X5)) | ~sP5('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5)) )),
  inference(condensation,[],[f10543])).
fof(f10543,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | p(s(bool,X39)) | ~sP5('T_0','const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f10524,f769])).
fof(f10524,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | p(s(bool,X39)) | ~sP5('T_0','const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1784,f711])).
fof(f10492,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | ~sP12(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5,'T_0') | p(s(bool,X5))) ) | $spl92),
  inference(cnf_transformation,[],[f10492_D])).
fof(f10492_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | ~sP12(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5,'T_0') | p(s(bool,X5))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f10494,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP12(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5,'T_0')) )),
  inference(condensation,[],[f10493])).
fof(f10493,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP12(happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),X46,'T_0') | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10472,f769])).
fof(f10472,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP12(happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),X46,'T_0') | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1764,f711])).
fof(f10491,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP12(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5,'T_0')) )),
  inference(condensation,[],[f10490])).
fof(f10490,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP12(happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),X39,'T_0') | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f10471,f769])).
fof(f10471,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP12(happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),X39,'T_0') | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1764,f711])).
fof(f10447,plain,(
  ( ! [X4,X2,X5,X3,X1] : (p(s(bool,X1)) | ~sP11('T_0',X1,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X3),s(X2,X4))),s(X2,X5))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X2,fun(X2,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X2,X4),s(X2,X5))))) ) | $spl90),
  inference(cnf_transformation,[],[f10447_D])).
fof(f10447_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (p(s(bool,X1)) | ~sP11('T_0',X1,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X3),s(X2,X4))),s(X2,X5))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X2,fun(X2,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X2,X4),s(X2,X5))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f10449,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X2,fun(X2,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X2,X4),s(X2,X5)))) | ~sP11('T_0',X1,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X3),s(X2,X4))),s(X2,X5)))) )),
  inference(condensation,[],[f10448])).
fof(f10448,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X43,fun(X43,bool)),X44),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X43,X45),s(X43,X46)))) | ~sP11('T_0',X42,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X45))),s(X43,X46))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10427,f769])).
fof(f10427,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X43,fun(X43,bool)),X44),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X43,X45),s(X43,X46)))) | ~sP11('T_0',X42,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X45))),s(X43,X46))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1761,f711])).
fof(f10446,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X2,fun(X2,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X2,X4),s(X2,X5)))) | ~sP11('T_0',X1,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X3),s(X2,X4))),s(X2,X5)))) )),
  inference(condensation,[],[f10445])).
fof(f10445,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (p(s(bool,X35)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X36,fun(X36,bool)),X37),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X36,X38),s(X36,X39)))) | ~sP11('T_0',X35,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X38))),s(X36,X39))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f10426,f769])).
fof(f10426,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | p(s(bool,X35)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X36,fun(X36,bool)),X37),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X36,X38),s(X36,X39)))) | ~sP11('T_0',X35,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X38))),s(X36,X39))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1761,f711])).
fof(f10138,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | ~sP8(X5,'T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))) | p(s(bool,X5))) ) | $spl88),
  inference(cnf_transformation,[],[f10138_D])).
fof(f10138_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | ~sP8(X5,'T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))) | p(s(bool,X5))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f10140,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP8(X5,'T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) )),
  inference(condensation,[],[f10139])).
fof(f10139,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP8(X46,'T_0',happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10118,f769])).
fof(f10118,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP8(X46,'T_0',happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1756,f711])).
fof(f10137,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP8(X5,'T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) )),
  inference(condensation,[],[f10136])).
fof(f10136,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP8(X39,'T_0',happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f10117,f769])).
fof(f10117,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP8(X39,'T_0',happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1756,f711])).
fof(f10087,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | ~sP7('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) | $spl86),
  inference(cnf_transformation,[],[f10087_D])).
fof(f10087_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | ~sP7('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f10089,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP7('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f10088])).
fof(f10088,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP7('T_0',happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10067,f769])).
fof(f10067,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP7('T_0',happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1755,f711])).
fof(f10086,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP7('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f10085])).
fof(f10085,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP7('T_0',happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f10066,f769])).
fof(f10066,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP7('T_0',happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1755,f711])).
fof(f10034,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | ~sP4('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) | $spl84),
  inference(cnf_transformation,[],[f10034_D])).
fof(f10034_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | ~sP4('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f10036,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP4('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f10035])).
fof(f10035,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP4('T_0',happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10014,f769])).
fof(f10014,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP4('T_0',happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1754,f711])).
fof(f10033,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP4('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f10032])).
fof(f10032,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP4('T_0',happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f10013,f769])).
fof(f10013,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP4('T_0',happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1754,f711])).
fof(f9994,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | ~sP5('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) | $spl82),
  inference(cnf_transformation,[],[f9994_D])).
fof(f9994_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | ~sP5('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f9996,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP5('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f9995])).
fof(f9995,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP5('T_0',happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f9974,f769])).
fof(f9974,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP5('T_0',happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1753,f711])).
fof(f9993,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP5('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f9992])).
fof(f9992,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP5('T_0',happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f9973,f769])).
fof(f9973,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP5('T_0',happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1753,f711])).
fof(f8084,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP12('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),'T_0',X5)) ) | $spl80),
  inference(cnf_transformation,[],[f8084_D])).
fof(f8084_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP12('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),'T_0',X5)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f8086,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP12('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),'T_0',X5)) )),
  inference(condensation,[],[f8085])).
fof(f8085,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | sP12('const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),'T_0',X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f8069,f769])).
fof(f8069,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | sP12('const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),'T_0',X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1794,f711])).
fof(f8083,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP12('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),'T_0',X5)) )),
  inference(condensation,[],[f8082])).
fof(f8082,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | sP12('const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),'T_0',X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f8068,f769])).
fof(f8068,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | sP12('const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),'T_0',X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1794,f711])).
fof(f7939,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP10('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),'T_0',X5)) ) | $spl78),
  inference(cnf_transformation,[],[f7939_D])).
fof(f7939_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP10('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),'T_0',X5)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f7941,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP10('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),'T_0',X5)) )),
  inference(condensation,[],[f7940])).
fof(f7940,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | sP10('const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),'T_0',X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f7924,f769])).
fof(f7924,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | sP10('const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),'T_0',X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1790,f711])).
fof(f7938,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP10('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),'T_0',X5)) )),
  inference(condensation,[],[f7937])).
fof(f7937,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | sP10('const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),'T_0',X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7923,f769])).
fof(f7923,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | sP10('const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),'T_0',X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1790,f711])).
fof(f7885,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP10('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5,'T_0')) ) | $spl76),
  inference(cnf_transformation,[],[f7885_D])).
fof(f7885_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP10('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5,'T_0')) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f7887,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP10('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5,'T_0')) )),
  inference(condensation,[],[f7886])).
fof(f7886,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | sP10('const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),X46,'T_0') | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f7870,f769])).
fof(f7870,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | sP10('const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),X46,'T_0') | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1789,f711])).
fof(f7884,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP10('const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5,'T_0')) )),
  inference(condensation,[],[f7883])).
fof(f7883,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | sP10('const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),X39,'T_0') | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7869,f769])).
fof(f7869,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | sP10('const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),X39,'T_0') | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1789,f711])).
fof(f7846,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP2('T_0',X5,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) ) | $spl74),
  inference(cnf_transformation,[],[f7846_D])).
fof(f7846_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP2('T_0',X5,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f7848,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP2('T_0',X5,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) )),
  inference(condensation,[],[f7847])).
fof(f7847,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | sP2('T_0',X46,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f7831,f769])).
fof(f7831,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | sP2('T_0',X46,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1782,f711])).
fof(f7845,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP2('T_0',X5,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) )),
  inference(condensation,[],[f7844])).
fof(f7844,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | sP2('T_0',X39,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7830,f769])).
fof(f7830,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | sP2('T_0',X39,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1782,f711])).
fof(f7805,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP2('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5)) ) | $spl72),
  inference(cnf_transformation,[],[f7805_D])).
fof(f7805_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP2('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f7807,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP2('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5)) )),
  inference(condensation,[],[f7806])).
fof(f7806,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | sP2('T_0','const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f7790,f769])).
fof(f7790,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)))) | sP2('T_0','const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1781,f711])).
fof(f1781,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,X65)) | ~p(s(bool,happ(s(fun(X61,bool),happ(s(fun(X61,fun(X61,bool)),X62),s(X61,X63))),s(X61,X64)))) | sP2(X65,'const.arithmetic.NRC_4'(s(fun(X61,fun(X61,bool)),X62),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X61,X63),s(X61,X64)),X66)) )),
  inference(superposition,[],[f343,f301])).
fof(f7804,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) | sP2('T_0','const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)),X5)) )),
  inference(condensation,[],[f7803])).
fof(f7803,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | sP2('T_0','const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7789,f769])).
fof(f7789,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | sP2('T_0','const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1781,f711])).
fof(f7760,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP12(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),'T_0',X5)) ) | $spl70),
  inference(cnf_transformation,[],[f7760_D])).
fof(f7760_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP12(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),'T_0',X5)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f7762,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP12(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),'T_0',X5)) )),
  inference(condensation,[],[f7761])).
fof(f7761,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | sP12(happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),'T_0',X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f7745,f769])).
fof(f7745,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | sP12(happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),'T_0',X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1763,f711])).
fof(f7759,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP12(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),'T_0',X5)) )),
  inference(condensation,[],[f7758])).
fof(f7758,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | sP12(happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),'T_0',X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7744,f769])).
fof(f7744,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | sP12(happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),'T_0',X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1763,f711])).
fof(f7604,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP10(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),'T_0',X5)) ) | $spl68),
  inference(cnf_transformation,[],[f7604_D])).
fof(f7604_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP10(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),'T_0',X5)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f7606,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP10(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),'T_0',X5)) )),
  inference(condensation,[],[f7605])).
fof(f7605,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | sP10(happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),'T_0',X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f7589,f769])).
fof(f7589,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | sP10(happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),'T_0',X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1759,f711])).
fof(f7603,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP10(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),'T_0',X5)) )),
  inference(condensation,[],[f7602])).
fof(f7602,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | sP10(happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),'T_0',X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7588,f769])).
fof(f7588,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | sP10(happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),'T_0',X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1759,f711])).
fof(f7567,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP10(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5,'T_0')) ) | $spl66),
  inference(cnf_transformation,[],[f7567_D])).
fof(f7567_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP10(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5,'T_0')) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f7569,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP10(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5,'T_0')) )),
  inference(condensation,[],[f7568])).
fof(f7568,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | sP10(happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),X46,'T_0') | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f7552,f769])).
fof(f7552,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | sP10(happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),X46,'T_0') | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1758,f711])).
fof(f7566,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP10(happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5,'T_0')) )),
  inference(condensation,[],[f7565])).
fof(f7565,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | sP10(happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),X39,'T_0') | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7551,f769])).
fof(f7551,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | sP10(happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),X39,'T_0') | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1758,f711])).
fof(f7522,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP2('T_0',X5,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) ) | $spl64),
  inference(cnf_transformation,[],[f7522_D])).
fof(f7522_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP2('T_0',X5,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f7524,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP2('T_0',X5,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) )),
  inference(condensation,[],[f7523])).
fof(f7523,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | sP2('T_0',X46,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f7507,f769])).
fof(f7507,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | sP2('T_0',X46,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1751,f711])).
fof(f7521,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP2('T_0',X5,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)))) )),
  inference(condensation,[],[f7520])).
fof(f7520,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | sP2('T_0',X39,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7506,f769])).
fof(f7506,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | sP2('T_0',X39,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1751,f711])).
fof(f7479,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP2('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5)) ) | $spl62),
  inference(cnf_transformation,[],[f7479_D])).
fof(f7479_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP2('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f7481,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP2('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f7480])).
fof(f7480,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | sP2('T_0',happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f7464,f769])).
fof(f7464,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X42,fun(X42,bool)),X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X42,X44),s(X42,X45)))) | sP2('T_0',happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1750,f711])).
fof(f1750,plain,(
  ( ! [X92,X90,X88,X93,X91,X89] : (~p(s(bool,X92)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X88,fun(X88,bool)),X89),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X88,X90),s(X88,X91)))) | sP2(X92,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X90))),s(X88,X91)),X93)) )),
  inference(superposition,[],[f343,f301])).
fof(f7478,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X1,X3),s(X1,X4)))) | sP2('T_0',happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f7477])).
fof(f7477,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | sP2('T_0',happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7463,f769])).
fof(f7463,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | sP2('T_0',happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1750,f711])).
fof(f5073,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK24(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK24(X0,X1,bool,bool))))) ) | $spl58),
  inference(cnf_transformation,[],[f5073_D])).
fof(f5073_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK24(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK24(X0,X1,bool,bool))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f5075,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl61),
  inference(cnf_transformation,[],[f5075_D])).
fof(f5075_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f5063,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK24(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK24(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1697,f711])).
fof(f1697,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK24(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK24(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f338,f255])).
fof(f5062,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK24(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK24(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1697,f711])).
fof(f4980,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl52),
  inference(cnf_transformation,[],[f4980_D])).
fof(f4980_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4976,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK24(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl48),
  inference(cnf_transformation,[],[f4976_D])).
fof(f4976_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK24(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f4983,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl57),
  inference(cnf_transformation,[],[f4983_D])).
fof(f4983_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f4974,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK24(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1696,f255])).
fof(f1696,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK24(X0,X1,X2,bool))) )),
  inference(superposition,[],[f338,f255])).
fof(f4981,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl54),
  inference(cnf_transformation,[],[f4981_D])).
fof(f4981_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4973,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK24(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1696,f711])).
fof(f4972,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK24(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1696,f711])).
fof(f4975,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl46),
  inference(cnf_transformation,[],[f4975_D])).
fof(f4975_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4978,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl51),
  inference(cnf_transformation,[],[f4978_D])).
fof(f4978_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f4971,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK24(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1696,f711])).
fof(f4970,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK24(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1696,f711])).
fof(f3222,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197,f3199,f3201,f3203,f3205,f3207,f3209,f3211,f3213,f3215,f3217,f3219,f3221])).
fof(f3221,plain,(
  ( ! [X103,X101,X102,X104] : (~sP9(happ(s(fun(X101,bool),X102),s(X101,'const.min.@_1'(s(fun(X101,bool),X102)))),X103,X104) | sP10(happ(s(fun(X101,bool),X102),s(X101,'const.min.@_1'(s(fun(X101,bool),X102)))),X103,X104)) ) | $spl10),
  inference(subsumption_resolution,[],[f2965,f247])).
fof(f2965,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,'T_0')) | ~sP9(happ(s(fun(X101,bool),X102),s(X101,'const.min.@_1'(s(fun(X101,bool),X102)))),X103,X104) | sP10(happ(s(fun(X101,bool),X102),s(X101,'const.min.@_1'(s(fun(X101,bool),X102)))),X103,X104)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1633])).
fof(f3219,plain,(
  ( ! [X99,X97,X100,X98] : (~sP11(happ(s(fun(X97,bool),X98),s(X97,'const.min.@_1'(s(fun(X97,bool),X98)))),X99,X100) | sP12(happ(s(fun(X97,bool),X98),s(X97,'const.min.@_1'(s(fun(X97,bool),X98)))),X99,X100)) ) | $spl10),
  inference(subsumption_resolution,[],[f2964,f247])).
fof(f2964,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | ~sP11(happ(s(fun(X97,bool),X98),s(X97,'const.min.@_1'(s(fun(X97,bool),X98)))),X99,X100) | sP12(happ(s(fun(X97,bool),X98),s(X97,'const.min.@_1'(s(fun(X97,bool),X98)))),X99,X100)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1632])).
fof(f3217,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,X95)) | p(s(bool,X96)) | ~sP12(happ(s(fun(X93,bool),X94),s(X93,'const.min.@_1'(s(fun(X93,bool),X94)))),X96,X95)) ) | $spl10),
  inference(subsumption_resolution,[],[f2963,f247])).
fof(f2963,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X95)) | p(s(bool,X96)) | ~sP12(happ(s(fun(X93,bool),X94),s(X93,'const.min.@_1'(s(fun(X93,bool),X94)))),X96,X95)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1631])).
fof(f3215,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X91)) | sP12(happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90)))),X91,X92)) ) | $spl10),
  inference(subsumption_resolution,[],[f2962,f247])).
fof(f2962,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,X91)) | sP12(happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90)))),X91,X92)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1630])).
fof(f3213,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X87)) | sP12(happ(s(fun(X85,bool),X86),s(X85,'const.min.@_1'(s(fun(X85,bool),X86)))),X88,X87)) ) | $spl10),
  inference(subsumption_resolution,[],[f2961,f247])).
fof(f2961,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'T_0')) | p(s(bool,X87)) | sP12(happ(s(fun(X85,bool),X86),s(X85,'const.min.@_1'(s(fun(X85,bool),X86)))),X88,X87)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1629])).
fof(f3211,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,X83)) | ~p(s(bool,X84)) | ~sP11(X84,X83,happ(s(fun(X81,bool),X82),s(X81,'const.min.@_1'(s(fun(X81,bool),X82)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2960,f247])).
fof(f2960,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,'T_0')) | p(s(bool,X83)) | ~p(s(bool,X84)) | ~sP11(X84,X83,happ(s(fun(X81,bool),X82),s(X81,'const.min.@_1'(s(fun(X81,bool),X82)))))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1628])).
fof(f3209,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,X79)) | p(s(bool,X80)) | ~sP10(happ(s(fun(X77,bool),X78),s(X77,'const.min.@_1'(s(fun(X77,bool),X78)))),X80,X79)) ) | $spl10),
  inference(subsumption_resolution,[],[f2959,f247])).
fof(f2959,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,'T_0')) | p(s(bool,X79)) | p(s(bool,X80)) | ~sP10(happ(s(fun(X77,bool),X78),s(X77,'const.min.@_1'(s(fun(X77,bool),X78)))),X80,X79)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1627])).
fof(f3207,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,X75)) | sP10(happ(s(fun(X73,bool),X74),s(X73,'const.min.@_1'(s(fun(X73,bool),X74)))),X75,X76)) ) | $spl10),
  inference(subsumption_resolution,[],[f2958,f247])).
fof(f2958,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X75)) | sP10(happ(s(fun(X73,bool),X74),s(X73,'const.min.@_1'(s(fun(X73,bool),X74)))),X75,X76)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1626])).
fof(f3205,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,X71)) | sP10(happ(s(fun(X69,bool),X70),s(X69,'const.min.@_1'(s(fun(X69,bool),X70)))),X72,X71)) ) | $spl10),
  inference(subsumption_resolution,[],[f2957,f247])).
fof(f2957,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'T_0')) | ~p(s(bool,X71)) | sP10(happ(s(fun(X69,bool),X70),s(X69,'const.min.@_1'(s(fun(X69,bool),X70)))),X72,X71)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1625])).
fof(f3203,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,X67)) | p(s(bool,X68)) | ~sP9(happ(s(fun(X65,bool),X66),s(X65,'const.min.@_1'(s(fun(X65,bool),X66)))),X67,X68)) ) | $spl10),
  inference(subsumption_resolution,[],[f2956,f247])).
fof(f2956,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X67)) | p(s(bool,X68)) | ~sP9(happ(s(fun(X65,bool),X66),s(X65,'const.min.@_1'(s(fun(X65,bool),X66)))),X67,X68)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1624])).
fof(f3201,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X63)) | p(s(bool,X64)) | ~sP8(X64,X63,happ(s(fun(X61,bool),X62),s(X61,'const.min.@_1'(s(fun(X61,bool),X62)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2955,f247])).
fof(f2955,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X63)) | p(s(bool,X64)) | ~sP8(X64,X63,happ(s(fun(X61,bool),X62),s(X61,'const.min.@_1'(s(fun(X61,bool),X62)))))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1623])).
fof(f3199,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X59)) | p(s(bool,X60)) | ~sP7(X59,happ(s(fun(X57,bool),X58),s(X57,'const.min.@_1'(s(fun(X57,bool),X58)))),X60)) ) | $spl10),
  inference(subsumption_resolution,[],[f2954,f247])).
fof(f2954,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X59)) | p(s(bool,X60)) | ~sP7(X59,happ(s(fun(X57,bool),X58),s(X57,'const.min.@_1'(s(fun(X57,bool),X58)))),X60)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1622])).
fof(f3197,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X55)) | p(s(bool,X56)) | ~sP4(X55,happ(s(fun(X53,bool),X54),s(X53,'const.min.@_1'(s(fun(X53,bool),X54)))),X56)) ) | $spl10),
  inference(subsumption_resolution,[],[f2953,f247])).
fof(f2953,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X55)) | p(s(bool,X56)) | ~sP4(X55,happ(s(fun(X53,bool),X54),s(X53,'const.min.@_1'(s(fun(X53,bool),X54)))),X56)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1621])).
fof(f3195,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP5(X51,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))),X52)) ) | $spl10),
  inference(subsumption_resolution,[],[f2952,f247])).
fof(f2952,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP5(X51,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))),X52)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1620])).
fof(f3193,plain,(
  ( ! [X47,X45,X48,X46] : (sP6(X47,X48,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46))))) | ~sP5(X47,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X48) | ~sP4(X48,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X47) | ~sP3(happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X48,X47)) ) | $spl10),
  inference(subsumption_resolution,[],[f2951,f247])).
fof(f2951,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'T_0')) | sP6(X47,X48,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46))))) | ~sP5(X47,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X48) | ~sP4(X48,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X47) | ~sP3(happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X48,X47)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1619])).
fof(f3191,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X43)) | sP2(X43,X44,happ(s(fun(X41,bool),X42),s(X41,'const.min.@_1'(s(fun(X41,bool),X42)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2950,f247])).
fof(f2950,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | sP2(X43,X44,happ(s(fun(X41,bool),X42),s(X41,'const.min.@_1'(s(fun(X41,bool),X42)))))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1618])).
fof(f3189,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X39)) | sP2(X39,happ(s(fun(X37,bool),X38),s(X37,'const.min.@_1'(s(fun(X37,bool),X38)))),X40)) ) | $spl10),
  inference(subsumption_resolution,[],[f2949,f247])).
fof(f2949,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | sP2(X39,happ(s(fun(X37,bool),X38),s(X37,'const.min.@_1'(s(fun(X37,bool),X38)))),X40)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1617])).
fof(f3187,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X32)) | sP1(X32,happ(s(fun(X30,bool),X31),s(X30,'const.min.@_1'(s(fun(X30,bool),X31)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2948,f247])).
fof(f2948,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | sP1(X32,happ(s(fun(X30,bool),X31),s(X30,'const.min.@_1'(s(fun(X30,bool),X31)))))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1615])).
fof(f3185,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X29)) | sP1(happ(s(fun(X27,bool),X28),s(X27,'const.min.@_1'(s(fun(X27,bool),X28)))),X29)) ) | $spl10),
  inference(subsumption_resolution,[],[f2947,f247])).
fof(f2947,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(happ(s(fun(X27,bool),X28),s(X27,'const.min.@_1'(s(fun(X27,bool),X28)))),X29)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1614])).
fof(f3183,plain,(
  ( ! [X26,X24,X25] : (sP0(happ(s(fun(X24,bool),X25),s(X24,'const.min.@_1'(s(fun(X24,bool),X25)))),X26)) ) | $spl10),
  inference(subsumption_resolution,[],[f2946,f247])).
fof(f2946,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X24,bool),X25),s(X24,'const.min.@_1'(s(fun(X24,bool),X25)))),X26)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1613])).
fof(f3181,plain,(
  ( ! [X23,X21,X22] : (sP0(X23,happ(s(fun(X21,bool),X22),s(X21,'const.min.@_1'(s(fun(X21,bool),X22)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2945,f247])).
fof(f2945,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'T_0')) | sP0(X23,happ(s(fun(X21,bool),X22),s(X21,'const.min.@_1'(s(fun(X21,bool),X22)))))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1612])).
fof(f3179,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (~sP9('const.arithmetic.NRC_4'(s(fun(X157,fun(X157,bool)),X158),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X157,X159),s(X157,X160)),X161,X162) | sP10('const.arithmetic.NRC_4'(s(fun(X157,fun(X157,bool)),X158),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X157,X159),s(X157,X160)),X161,X162)) ) | $spl10),
  inference(subsumption_resolution,[],[f2934,f247])).
fof(f2934,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (~p(s(bool,'T_0')) | ~sP9('const.arithmetic.NRC_4'(s(fun(X157,fun(X157,bool)),X158),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X157,X159),s(X157,X160)),X161,X162) | sP10('const.arithmetic.NRC_4'(s(fun(X157,fun(X157,bool)),X158),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X157,X159),s(X157,X160)),X161,X162)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1797])).
fof(f3177,plain,(
  ( ! [X156,X154,X152,X151,X155,X153] : (~sP11('const.arithmetic.NRC_4'(s(fun(X151,fun(X151,bool)),X152),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X151,X153),s(X151,X154)),X155,X156) | sP12('const.arithmetic.NRC_4'(s(fun(X151,fun(X151,bool)),X152),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X151,X153),s(X151,X154)),X155,X156)) ) | $spl10),
  inference(subsumption_resolution,[],[f2933,f247])).
fof(f2933,plain,(
  ( ! [X156,X154,X152,X151,X155,X153] : (~p(s(bool,'T_0')) | ~sP11('const.arithmetic.NRC_4'(s(fun(X151,fun(X151,bool)),X152),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X151,X153),s(X151,X154)),X155,X156) | sP12('const.arithmetic.NRC_4'(s(fun(X151,fun(X151,bool)),X152),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X151,X153),s(X151,X154)),X155,X156)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1796])).
fof(f3175,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (~p(s(bool,X149)) | p(s(bool,X150)) | ~sP12('const.arithmetic.NRC_4'(s(fun(X145,fun(X145,bool)),X146),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X145,X147),s(X145,X148)),X150,X149)) ) | $spl10),
  inference(subsumption_resolution,[],[f2932,f247])).
fof(f2932,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (~p(s(bool,'T_0')) | ~p(s(bool,X149)) | p(s(bool,X150)) | ~sP12('const.arithmetic.NRC_4'(s(fun(X145,fun(X145,bool)),X146),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X145,X147),s(X145,X148)),X150,X149)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1795])).
fof(f3173,plain,(
  ( ! [X144,X142,X140,X143,X141,X139] : (~p(s(bool,X143)) | sP12('const.arithmetic.NRC_4'(s(fun(X139,fun(X139,bool)),X140),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X139,X141),s(X139,X142)),X143,X144)) ) | $spl10),
  inference(subsumption_resolution,[],[f2931,f247])).
fof(f2931,plain,(
  ( ! [X144,X142,X140,X143,X141,X139] : (~p(s(bool,'T_0')) | ~p(s(bool,X143)) | sP12('const.arithmetic.NRC_4'(s(fun(X139,fun(X139,bool)),X140),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X139,X141),s(X139,X142)),X143,X144)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1794])).
fof(f3171,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (p(s(bool,X137)) | sP12('const.arithmetic.NRC_4'(s(fun(X133,fun(X133,bool)),X134),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X133,X135),s(X133,X136)),X138,X137)) ) | $spl10),
  inference(subsumption_resolution,[],[f2930,f247])).
fof(f2930,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (~p(s(bool,'T_0')) | p(s(bool,X137)) | sP12('const.arithmetic.NRC_4'(s(fun(X133,fun(X133,bool)),X134),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X133,X135),s(X133,X136)),X138,X137)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1793])).
fof(f3169,plain,(
  ( ! [X127,X132,X130,X128,X131,X129] : (p(s(bool,X131)) | ~p(s(bool,X132)) | ~sP11(X132,X131,'const.arithmetic.NRC_4'(s(fun(X127,fun(X127,bool)),X128),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X127,X129),s(X127,X130)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2929,f247])).
fof(f2929,plain,(
  ( ! [X127,X132,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,X131)) | ~p(s(bool,X132)) | ~sP11(X132,X131,'const.arithmetic.NRC_4'(s(fun(X127,fun(X127,bool)),X128),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X127,X129),s(X127,X130)))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1792])).
fof(f3167,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (p(s(bool,X125)) | p(s(bool,X126)) | ~sP10('const.arithmetic.NRC_4'(s(fun(X121,fun(X121,bool)),X122),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X121,X123),s(X121,X124)),X126,X125)) ) | $spl10),
  inference(subsumption_resolution,[],[f2928,f247])).
fof(f2928,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (~p(s(bool,'T_0')) | p(s(bool,X125)) | p(s(bool,X126)) | ~sP10('const.arithmetic.NRC_4'(s(fun(X121,fun(X121,bool)),X122),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X121,X123),s(X121,X124)),X126,X125)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1791])).
fof(f3165,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (~p(s(bool,X119)) | sP10('const.arithmetic.NRC_4'(s(fun(X115,fun(X115,bool)),X116),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X115,X117),s(X115,X118)),X119,X120)) ) | $spl10),
  inference(subsumption_resolution,[],[f2927,f247])).
fof(f2927,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (~p(s(bool,'T_0')) | ~p(s(bool,X119)) | sP10('const.arithmetic.NRC_4'(s(fun(X115,fun(X115,bool)),X116),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X115,X117),s(X115,X118)),X119,X120)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1790])).
fof(f3163,plain,(
  ( ! [X111,X109,X114,X112,X110,X113] : (~p(s(bool,X113)) | sP10('const.arithmetic.NRC_4'(s(fun(X109,fun(X109,bool)),X110),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X109,X111),s(X109,X112)),X114,X113)) ) | $spl10),
  inference(subsumption_resolution,[],[f2926,f247])).
fof(f2926,plain,(
  ( ! [X111,X109,X114,X112,X110,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,X113)) | sP10('const.arithmetic.NRC_4'(s(fun(X109,fun(X109,bool)),X110),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X109,X111),s(X109,X112)),X114,X113)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1789])).
fof(f3161,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (p(s(bool,X107)) | p(s(bool,X108)) | ~sP9('const.arithmetic.NRC_4'(s(fun(X103,fun(X103,bool)),X104),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X103,X105),s(X103,X106)),X107,X108)) ) | $spl10),
  inference(subsumption_resolution,[],[f2925,f247])).
fof(f2925,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X107)) | p(s(bool,X108)) | ~sP9('const.arithmetic.NRC_4'(s(fun(X103,fun(X103,bool)),X104),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X103,X105),s(X103,X106)),X107,X108)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1788])).
fof(f3159,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (~p(s(bool,X101)) | p(s(bool,X102)) | ~sP8(X102,X101,'const.arithmetic.NRC_4'(s(fun(X97,fun(X97,bool)),X98),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X97,X99),s(X97,X100)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2924,f247])).
fof(f2924,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (~p(s(bool,'T_0')) | ~p(s(bool,X101)) | p(s(bool,X102)) | ~sP8(X102,X101,'const.arithmetic.NRC_4'(s(fun(X97,fun(X97,bool)),X98),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X97,X99),s(X97,X100)))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1787])).
fof(f3157,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (~p(s(bool,X95)) | p(s(bool,X96)) | ~sP7(X95,'const.arithmetic.NRC_4'(s(fun(X91,fun(X91,bool)),X92),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X91,X93),s(X91,X94)),X96)) ) | $spl10),
  inference(subsumption_resolution,[],[f2923,f247])).
fof(f2923,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X95)) | p(s(bool,X96)) | ~sP7(X95,'const.arithmetic.NRC_4'(s(fun(X91,fun(X91,bool)),X92),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X91,X93),s(X91,X94)),X96)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1786])).
fof(f3155,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,X89)) | p(s(bool,X90)) | ~sP4(X89,'const.arithmetic.NRC_4'(s(fun(X85,fun(X85,bool)),X86),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X85,X87),s(X85,X88)),X90)) ) | $spl10),
  inference(subsumption_resolution,[],[f2922,f247])).
fof(f2922,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,'T_0')) | ~p(s(bool,X89)) | p(s(bool,X90)) | ~sP4(X89,'const.arithmetic.NRC_4'(s(fun(X85,fun(X85,bool)),X86),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X85,X87),s(X85,X88)),X90)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1785])).
fof(f3153,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (~p(s(bool,X83)) | p(s(bool,X84)) | ~sP5(X83,'const.arithmetic.NRC_4'(s(fun(X79,fun(X79,bool)),X80),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X79,X81),s(X79,X82)),X84)) ) | $spl10),
  inference(subsumption_resolution,[],[f2921,f247])).
fof(f2921,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X83)) | p(s(bool,X84)) | ~sP5(X83,'const.arithmetic.NRC_4'(s(fun(X79,fun(X79,bool)),X80),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X79,X81),s(X79,X82)),X84)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1784])).
fof(f3151,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (sP6(X77,X78,'const.arithmetic.NRC_4'(s(fun(X73,fun(X73,bool)),X74),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X73,X75),s(X73,X76))) | ~sP5(X77,'const.arithmetic.NRC_4'(s(fun(X73,fun(X73,bool)),X74),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X73,X75),s(X73,X76)),X78) | ~sP4(X78,'const.arithmetic.NRC_4'(s(fun(X73,fun(X73,bool)),X74),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X73,X75),s(X73,X76)),X77) | ~sP3('const.arithmetic.NRC_4'(s(fun(X73,fun(X73,bool)),X74),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X73,X75),s(X73,X76)),X78,X77)) ) | $spl10),
  inference(subsumption_resolution,[],[f2920,f247])).
fof(f2920,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (~p(s(bool,'T_0')) | sP6(X77,X78,'const.arithmetic.NRC_4'(s(fun(X73,fun(X73,bool)),X74),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X73,X75),s(X73,X76))) | ~sP5(X77,'const.arithmetic.NRC_4'(s(fun(X73,fun(X73,bool)),X74),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X73,X75),s(X73,X76)),X78) | ~sP4(X78,'const.arithmetic.NRC_4'(s(fun(X73,fun(X73,bool)),X74),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X73,X75),s(X73,X76)),X77) | ~sP3('const.arithmetic.NRC_4'(s(fun(X73,fun(X73,bool)),X74),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X73,X75),s(X73,X76)),X78,X77)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1783])).
fof(f3149,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,X71)) | sP2(X71,X72,'const.arithmetic.NRC_4'(s(fun(X67,fun(X67,bool)),X68),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X67,X69),s(X67,X70)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2919,f247])).
fof(f2919,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X71)) | sP2(X71,X72,'const.arithmetic.NRC_4'(s(fun(X67,fun(X67,bool)),X68),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X67,X69),s(X67,X70)))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1782])).
fof(f3147,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,X65)) | sP2(X65,'const.arithmetic.NRC_4'(s(fun(X61,fun(X61,bool)),X62),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X61,X63),s(X61,X64)),X66)) ) | $spl10),
  inference(subsumption_resolution,[],[f2918,f247])).
fof(f2918,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X65)) | sP2(X65,'const.arithmetic.NRC_4'(s(fun(X61,fun(X61,bool)),X62),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X61,X63),s(X61,X64)),X66)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1781])).
fof(f3145,plain,(
  ( ! [X54,X52,X50,X53,X51] : (p(s(bool,X54)) | sP1(X54,'const.arithmetic.NRC_4'(s(fun(X50,fun(X50,bool)),X51),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X50,X52),s(X50,X53)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2917,f247])).
fof(f2917,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X54)) | sP1(X54,'const.arithmetic.NRC_4'(s(fun(X50,fun(X50,bool)),X51),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X50,X52),s(X50,X53)))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1779])).
fof(f3143,plain,(
  ( ! [X47,X45,X48,X46,X49] : (p(s(bool,X49)) | sP1('const.arithmetic.NRC_4'(s(fun(X45,fun(X45,bool)),X46),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X45,X47),s(X45,X48)),X49)) ) | $spl10),
  inference(subsumption_resolution,[],[f2916,f247])).
fof(f2916,plain,(
  ( ! [X47,X45,X48,X46,X49] : (~p(s(bool,'T_0')) | p(s(bool,X49)) | sP1('const.arithmetic.NRC_4'(s(fun(X45,fun(X45,bool)),X46),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X45,X47),s(X45,X48)),X49)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1778])).
fof(f3141,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP0('const.arithmetic.NRC_4'(s(fun(X40,fun(X40,bool)),X41),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X40,X42),s(X40,X43)),X44)) ) | $spl10),
  inference(subsumption_resolution,[],[f2915,f247])).
fof(f2915,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.NRC_4'(s(fun(X40,fun(X40,bool)),X41),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X40,X42),s(X40,X43)),X44)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1777])).
fof(f3139,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP0(X39,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2914,f247])).
fof(f2914,plain,(
  ( ! [X39,X37,X35,X38,X36] : (~p(s(bool,'T_0')) | sP0(X39,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1776])).
fof(f3065,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (~sP9(happ(s(fun(X184,bool),happ(s(fun(X184,fun(X184,bool)),X185),s(X184,X186))),s(X184,X187)),X188,X189) | sP10(happ(s(fun(X184,bool),happ(s(fun(X184,fun(X184,bool)),X185),s(X184,X186))),s(X184,X187)),X188,X189)) ) | $spl10),
  inference(subsumption_resolution,[],[f2843,f247])).
fof(f2843,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (~p(s(bool,'T_0')) | ~sP9(happ(s(fun(X184,bool),happ(s(fun(X184,fun(X184,bool)),X185),s(X184,X186))),s(X184,X187)),X188,X189) | sP10(happ(s(fun(X184,bool),happ(s(fun(X184,fun(X184,bool)),X185),s(X184,X186))),s(X184,X187)),X188,X189)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1766])).
fof(f3063,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (~sP11(happ(s(fun(X178,bool),happ(s(fun(X178,fun(X178,bool)),X179),s(X178,X180))),s(X178,X181)),X182,X183) | sP12(happ(s(fun(X178,bool),happ(s(fun(X178,fun(X178,bool)),X179),s(X178,X180))),s(X178,X181)),X182,X183)) ) | $spl10),
  inference(subsumption_resolution,[],[f2842,f247])).
fof(f2842,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (~p(s(bool,'T_0')) | ~sP11(happ(s(fun(X178,bool),happ(s(fun(X178,fun(X178,bool)),X179),s(X178,X180))),s(X178,X181)),X182,X183) | sP12(happ(s(fun(X178,bool),happ(s(fun(X178,fun(X178,bool)),X179),s(X178,X180))),s(X178,X181)),X182,X183)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1765])).
fof(f3061,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (~p(s(bool,X176)) | p(s(bool,X177)) | ~sP12(happ(s(fun(X172,bool),happ(s(fun(X172,fun(X172,bool)),X173),s(X172,X174))),s(X172,X175)),X177,X176)) ) | $spl10),
  inference(subsumption_resolution,[],[f2841,f247])).
fof(f2841,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (~p(s(bool,'T_0')) | ~p(s(bool,X176)) | p(s(bool,X177)) | ~sP12(happ(s(fun(X172,bool),happ(s(fun(X172,fun(X172,bool)),X173),s(X172,X174))),s(X172,X175)),X177,X176)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1764])).
fof(f3059,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (~p(s(bool,X170)) | sP12(happ(s(fun(X166,bool),happ(s(fun(X166,fun(X166,bool)),X167),s(X166,X168))),s(X166,X169)),X170,X171)) ) | $spl10),
  inference(subsumption_resolution,[],[f2840,f247])).
fof(f2840,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (~p(s(bool,'T_0')) | ~p(s(bool,X170)) | sP12(happ(s(fun(X166,bool),happ(s(fun(X166,fun(X166,bool)),X167),s(X166,X168))),s(X166,X169)),X170,X171)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1763])).
fof(f3057,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (p(s(bool,X164)) | sP12(happ(s(fun(X160,bool),happ(s(fun(X160,fun(X160,bool)),X161),s(X160,X162))),s(X160,X163)),X165,X164)) ) | $spl10),
  inference(subsumption_resolution,[],[f2839,f247])).
fof(f2839,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (~p(s(bool,'T_0')) | p(s(bool,X164)) | sP12(happ(s(fun(X160,bool),happ(s(fun(X160,fun(X160,bool)),X161),s(X160,X162))),s(X160,X163)),X165,X164)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1762])).
fof(f3055,plain,(
  ( ! [X158,X156,X154,X159,X157,X155] : (p(s(bool,X158)) | ~p(s(bool,X159)) | ~sP11(X159,X158,happ(s(fun(X154,bool),happ(s(fun(X154,fun(X154,bool)),X155),s(X154,X156))),s(X154,X157)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2838,f247])).
fof(f2838,plain,(
  ( ! [X158,X156,X154,X159,X157,X155] : (~p(s(bool,'T_0')) | p(s(bool,X158)) | ~p(s(bool,X159)) | ~sP11(X159,X158,happ(s(fun(X154,bool),happ(s(fun(X154,fun(X154,bool)),X155),s(X154,X156))),s(X154,X157)))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1761])).
fof(f3053,plain,(
  ( ! [X152,X151,X149,X153,X150,X148] : (p(s(bool,X152)) | p(s(bool,X153)) | ~sP10(happ(s(fun(X148,bool),happ(s(fun(X148,fun(X148,bool)),X149),s(X148,X150))),s(X148,X151)),X153,X152)) ) | $spl10),
  inference(subsumption_resolution,[],[f2837,f247])).
fof(f2837,plain,(
  ( ! [X152,X151,X149,X153,X150,X148] : (~p(s(bool,'T_0')) | p(s(bool,X152)) | p(s(bool,X153)) | ~sP10(happ(s(fun(X148,bool),happ(s(fun(X148,fun(X148,bool)),X149),s(X148,X150))),s(X148,X151)),X153,X152)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1760])).
fof(f3051,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (~p(s(bool,X146)) | sP10(happ(s(fun(X142,bool),happ(s(fun(X142,fun(X142,bool)),X143),s(X142,X144))),s(X142,X145)),X146,X147)) ) | $spl10),
  inference(subsumption_resolution,[],[f2836,f247])).
fof(f2836,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (~p(s(bool,'T_0')) | ~p(s(bool,X146)) | sP10(happ(s(fun(X142,bool),happ(s(fun(X142,fun(X142,bool)),X143),s(X142,X144))),s(X142,X145)),X146,X147)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1759])).
fof(f3049,plain,(
  ( ! [X140,X138,X136,X141,X139,X137] : (~p(s(bool,X140)) | sP10(happ(s(fun(X136,bool),happ(s(fun(X136,fun(X136,bool)),X137),s(X136,X138))),s(X136,X139)),X141,X140)) ) | $spl10),
  inference(subsumption_resolution,[],[f2835,f247])).
fof(f2835,plain,(
  ( ! [X140,X138,X136,X141,X139,X137] : (~p(s(bool,'T_0')) | ~p(s(bool,X140)) | sP10(happ(s(fun(X136,bool),happ(s(fun(X136,fun(X136,bool)),X137),s(X136,X138))),s(X136,X139)),X141,X140)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1758])).
fof(f3047,plain,(
  ( ! [X134,X132,X130,X135,X133,X131] : (p(s(bool,X134)) | p(s(bool,X135)) | ~sP9(happ(s(fun(X130,bool),happ(s(fun(X130,fun(X130,bool)),X131),s(X130,X132))),s(X130,X133)),X134,X135)) ) | $spl10),
  inference(subsumption_resolution,[],[f2834,f247])).
fof(f2834,plain,(
  ( ! [X134,X132,X130,X135,X133,X131] : (~p(s(bool,'T_0')) | p(s(bool,X134)) | p(s(bool,X135)) | ~sP9(happ(s(fun(X130,bool),happ(s(fun(X130,fun(X130,bool)),X131),s(X130,X132))),s(X130,X133)),X134,X135)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1757])).
fof(f3045,plain,(
  ( ! [X127,X125,X128,X126,X124,X129] : (~p(s(bool,X128)) | p(s(bool,X129)) | ~sP8(X129,X128,happ(s(fun(X124,bool),happ(s(fun(X124,fun(X124,bool)),X125),s(X124,X126))),s(X124,X127)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2833,f247])).
fof(f2833,plain,(
  ( ! [X127,X125,X128,X126,X124,X129] : (~p(s(bool,'T_0')) | ~p(s(bool,X128)) | p(s(bool,X129)) | ~sP8(X129,X128,happ(s(fun(X124,bool),happ(s(fun(X124,fun(X124,bool)),X125),s(X124,X126))),s(X124,X127)))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1756])).
fof(f3043,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (~p(s(bool,X122)) | p(s(bool,X123)) | ~sP7(X122,happ(s(fun(X118,bool),happ(s(fun(X118,fun(X118,bool)),X119),s(X118,X120))),s(X118,X121)),X123)) ) | $spl10),
  inference(subsumption_resolution,[],[f2832,f247])).
fof(f2832,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (~p(s(bool,'T_0')) | ~p(s(bool,X122)) | p(s(bool,X123)) | ~sP7(X122,happ(s(fun(X118,bool),happ(s(fun(X118,fun(X118,bool)),X119),s(X118,X120))),s(X118,X121)),X123)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1755])).
fof(f3041,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (~p(s(bool,X116)) | p(s(bool,X117)) | ~sP4(X116,happ(s(fun(X112,bool),happ(s(fun(X112,fun(X112,bool)),X113),s(X112,X114))),s(X112,X115)),X117)) ) | $spl10),
  inference(subsumption_resolution,[],[f2831,f247])).
fof(f2831,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,X116)) | p(s(bool,X117)) | ~sP4(X116,happ(s(fun(X112,bool),happ(s(fun(X112,fun(X112,bool)),X113),s(X112,X114))),s(X112,X115)),X117)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1754])).
fof(f3039,plain,(
  ( ! [X111,X109,X107,X110,X108,X106] : (~p(s(bool,X110)) | p(s(bool,X111)) | ~sP5(X110,happ(s(fun(X106,bool),happ(s(fun(X106,fun(X106,bool)),X107),s(X106,X108))),s(X106,X109)),X111)) ) | $spl10),
  inference(subsumption_resolution,[],[f2830,f247])).
fof(f2830,plain,(
  ( ! [X111,X109,X107,X110,X108,X106] : (~p(s(bool,'T_0')) | ~p(s(bool,X110)) | p(s(bool,X111)) | ~sP5(X110,happ(s(fun(X106,bool),happ(s(fun(X106,fun(X106,bool)),X107),s(X106,X108))),s(X106,X109)),X111)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1753])).
fof(f3037,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (sP6(X104,X105,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103))) | ~sP5(X104,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103)),X105) | ~sP4(X105,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103)),X104) | ~sP3(happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103)),X105,X104)) ) | $spl10),
  inference(subsumption_resolution,[],[f2829,f247])).
fof(f2829,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (~p(s(bool,'T_0')) | sP6(X104,X105,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103))) | ~sP5(X104,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103)),X105) | ~sP4(X105,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103)),X104) | ~sP3(happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103)),X105,X104)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1752])).
fof(f3035,plain,(
  ( ! [X94,X99,X97,X95,X98,X96] : (~p(s(bool,X98)) | sP2(X98,X99,happ(s(fun(X94,bool),happ(s(fun(X94,fun(X94,bool)),X95),s(X94,X96))),s(X94,X97)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2828,f247])).
fof(f2828,plain,(
  ( ! [X94,X99,X97,X95,X98,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X98)) | sP2(X98,X99,happ(s(fun(X94,bool),happ(s(fun(X94,fun(X94,bool)),X95),s(X94,X96))),s(X94,X97)))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1751])).
fof(f3033,plain,(
  ( ! [X92,X90,X88,X93,X91,X89] : (~p(s(bool,X92)) | sP2(X92,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X90))),s(X88,X91)),X93)) ) | $spl10),
  inference(subsumption_resolution,[],[f2827,f247])).
fof(f2827,plain,(
  ( ! [X92,X90,X88,X93,X91,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,X92)) | sP2(X92,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X90))),s(X88,X91)),X93)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1750])).
fof(f3031,plain,(
  ( ! [X80,X78,X81,X79,X77] : (p(s(bool,X81)) | sP1(X81,happ(s(fun(X77,bool),happ(s(fun(X77,fun(X77,bool)),X78),s(X77,X79))),s(X77,X80)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2826,f247])).
fof(f2826,plain,(
  ( ! [X80,X78,X81,X79,X77] : (~p(s(bool,'T_0')) | p(s(bool,X81)) | sP1(X81,happ(s(fun(X77,bool),happ(s(fun(X77,fun(X77,bool)),X78),s(X77,X79))),s(X77,X80)))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1748])).
fof(f3029,plain,(
  ( ! [X76,X74,X72,X75,X73] : (p(s(bool,X76)) | sP1(happ(s(fun(X72,bool),happ(s(fun(X72,fun(X72,bool)),X73),s(X72,X74))),s(X72,X75)),X76)) ) | $spl10),
  inference(subsumption_resolution,[],[f2825,f247])).
fof(f2825,plain,(
  ( ! [X76,X74,X72,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X76)) | sP1(happ(s(fun(X72,bool),happ(s(fun(X72,fun(X72,bool)),X73),s(X72,X74))),s(X72,X75)),X76)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1747])).
fof(f3027,plain,(
  ( ! [X70,X68,X71,X69,X67] : (sP0(happ(s(fun(X67,bool),happ(s(fun(X67,fun(X67,bool)),X68),s(X67,X69))),s(X67,X70)),X71)) ) | $spl10),
  inference(subsumption_resolution,[],[f2824,f247])).
fof(f2824,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X67,bool),happ(s(fun(X67,fun(X67,bool)),X68),s(X67,X69))),s(X67,X70)),X71)) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1746])).
fof(f3025,plain,(
  ( ! [X66,X64,X62,X65,X63] : (sP0(X66,happ(s(fun(X62,bool),happ(s(fun(X62,fun(X62,bool)),X63),s(X62,X64))),s(X62,X65)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2823,f247])).
fof(f2823,plain,(
  ( ! [X66,X64,X62,X65,X63] : (~p(s(bool,'T_0')) | sP0(X66,happ(s(fun(X62,bool),happ(s(fun(X62,fun(X62,bool)),X63),s(X62,X64))),s(X62,X65)))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1745])).
fof(f2993,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl10),
  inference(subsumption_resolution,[],[f2770,f2645])).
fof(f2770,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f266])).
fof(f2385,plain,(
  ( ! [X251,X249,X250,X248] : (s(X250,f(s(bool,'T_0'))) != s(X250,g(s(bool,'T_0'))) | s(fun(bool,X250),X248) = s(fun(bool,X250),X249) | s(bool,X251) = s(bool,sK24(X248,X249,X250,bool)) | s(bool,'T_0') = s(bool,X251)) )),
  inference(superposition,[],[f338,f711])).
fof(f2384,plain,(
  ( ! [X246,X244,X247,X245] : (~p(s(bool,happ(s(fun(bool,bool),X244),s(bool,'T_0')))) | p(s(bool,X245)) | p(s(bool,happ(s(fun(bool,bool),X244),s(bool,X246)))) | s(bool,X247) = s(bool,sK22(X244,X245,bool)) | s(bool,'T_0') = s(bool,X247)) )),
  inference(superposition,[],[f300,f711])).
fof(f2383,plain,(
  ( ! [X243,X241,X242,X240] : (~p(s(bool,happ(s(fun(bool,bool),X240),s(bool,'T_0')))) | p(s(bool,X241)) | p(s(bool,happ(s(fun(bool,bool),X240),s(bool,X242)))) | s(bool,X243) = s(bool,sK21(X240,X241,bool)) | s(bool,'T_0') = s(bool,X243)) )),
  inference(superposition,[],[f298,f711])).
fof(f2382,plain,(
  ( ! [X239,X237,X238,X236] : (~p(s(bool,happ(s(fun(bool,bool),X237),s(bool,'T_0')))) | ~p(s(bool,X236)) | p(s(bool,happ(s(fun(bool,bool),X237),s(bool,X238)))) | s(bool,X239) = s(bool,sK20(X236,X237,bool)) | s(bool,'T_0') = s(bool,X239)) )),
  inference(superposition,[],[f295,f711])).
fof(f2381,plain,(
  ( ! [X235,X233,X234,X232] : (~p(s(bool,happ(s(fun(bool,bool),X233),s(bool,'T_0')))) | ~p(s(bool,X232)) | p(s(bool,happ(s(fun(bool,bool),X233),s(bool,X234)))) | s(bool,X235) = s(bool,sK19(X232,X233,bool)) | s(bool,'T_0') = s(bool,X235)) )),
  inference(superposition,[],[f293,f711])).
fof(f2380,plain,(
  ( ! [X231,X230] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X230))) = s(bool,happ(s(fun(bool,bool),X230),s(bool,'T_0'))) | s(bool,X231) = s(bool,'const.min.@_1'(s(fun(bool,bool),X230))) | s(bool,'T_0') = s(bool,X231)) )),
  inference(superposition,[],[f289,f711])).
fof(f2379,plain,(
  ( ! [X229,X227,X228,X226] : (s(bool,'T_0') != s(bool,f(s(X226,sK24(X227,X228,bool,X226)))) | s(fun(X226,bool),X227) = s(fun(X226,bool),X228) | s(bool,X229) = s(bool,g(s(X226,sK24(X227,X228,bool,X226)))) | s(bool,'T_0') = s(bool,X229)) )),
  inference(superposition,[],[f338,f711])).
fof(f2377,plain,(
  ( ! [X218,X216,X215,X213,X219,X217,X214] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X214),s('type.num.num',X215),s(bool,X216),s(bool,'T_0')))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X214),s('type.num.num',X217),s(bool,X213),s(bool,X218)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X214),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X215),s('type.num.num',X217))),s(bool,X216),s(bool,X218)))) | s(bool,X213) = s(bool,X219) | s(bool,'T_0') = s(bool,X219)) )),
  inference(superposition,[],[f304,f711])).
fof(f2376,plain,(
  ( ! [X212,X210,X208,X206,X211,X209,X207] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X207),s('type.num.num',X208),s(bool,'T_0'),s(bool,X209)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X207),s('type.num.num',X210),s(bool,X209),s(bool,X211)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X207),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X208),s('type.num.num',X210))),s(bool,X206),s(bool,X211)))) | s(bool,X206) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(superposition,[],[f304,f711])).
fof(f2375,plain,(
  ( ! [X204,X202,X200,X205,X203,X201] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X201),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X202),s('type.num.num',X203))),s(bool,X204),s(bool,'T_0')))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X201),s('type.num.num',X203),s(bool,sK23(X201,X202,X203,X204,X200,bool)),s(bool,X200)))) | s(bool,X200) = s(bool,X205) | s(bool,'T_0') = s(bool,X205)) )),
  inference(superposition,[],[f303,f711])).
fof(f2374,plain,(
  ( ! [X198,X196,X194,X199,X197,X195] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X195),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X196),s('type.num.num',X197))),s(bool,'T_0'),s(bool,X198)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X195),s('type.num.num',X197),s(bool,sK23(X195,X196,X197,X194,X198,bool)),s(bool,X198)))) | s(bool,X194) = s(bool,X199) | s(bool,'T_0') = s(bool,X199)) )),
  inference(superposition,[],[f303,f711])).
fof(f2373,plain,(
  ( ! [X191,X189,X192,X190,X188,X193] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X189),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X190),s('type.num.num',X191))),s(bool,X192),s(bool,'T_0')))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X189),s('type.num.num',X190),s(bool,X192),s(bool,sK23(X189,X190,X191,X192,X188,bool))))) | s(bool,X188) = s(bool,X193) | s(bool,'T_0') = s(bool,X193)) )),
  inference(superposition,[],[f302,f711])).
fof(f2372,plain,(
  ( ! [X187,X185,X182,X186,X184,X183] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X183),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X184),s('type.num.num',X185))),s(bool,'T_0'),s(bool,X186)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X183),s('type.num.num',X184),s(bool,X182),s(bool,sK23(X183,X184,X185,X182,X186,bool))))) | s(bool,X182) = s(bool,X187) | s(bool,'T_0') = s(bool,X187)) )),
  inference(superposition,[],[f302,f711])).
fof(f2728,plain,(
  ( ! [X167,X165,X166,X164,X168] : (p(s(bool,X166)) | p(s(bool,happ(s(fun(X164,bool),X165),s(X164,X167)))) | s(bool,X168) = s(bool,happ(s(fun(X164,bool),X165),s(X164,sK22(X165,X166,X164)))) | s(bool,'T_0') = s(bool,X168)) )),
  inference(subsumption_resolution,[],[f2368,f247])).
fof(f2368,plain,(
  ( ! [X167,X165,X166,X164,X168] : (~p(s(bool,'T_0')) | p(s(bool,X166)) | p(s(bool,happ(s(fun(X164,bool),X165),s(X164,X167)))) | s(bool,X168) = s(bool,happ(s(fun(X164,bool),X165),s(X164,sK22(X165,X166,X164)))) | s(bool,'T_0') = s(bool,X168)) )),
  inference(superposition,[],[f300,f711])).
fof(f2727,plain,(
  ( ! [X163,X161,X159,X162,X160] : (p(s(bool,X161)) | p(s(bool,happ(s(fun(X159,bool),X160),s(X159,X162)))) | s(bool,X163) = s(bool,happ(s(fun(X159,bool),X160),s(X159,sK21(X160,X161,X159)))) | s(bool,'T_0') = s(bool,X163)) )),
  inference(subsumption_resolution,[],[f2367,f247])).
fof(f2367,plain,(
  ( ! [X163,X161,X159,X162,X160] : (~p(s(bool,'T_0')) | p(s(bool,X161)) | p(s(bool,happ(s(fun(X159,bool),X160),s(X159,X162)))) | s(bool,X163) = s(bool,happ(s(fun(X159,bool),X160),s(X159,sK21(X160,X161,X159)))) | s(bool,'T_0') = s(bool,X163)) )),
  inference(superposition,[],[f298,f711])).
fof(f2726,plain,(
  ( ! [X158,X156,X154,X157,X155] : (~p(s(bool,X156)) | p(s(bool,happ(s(fun(X154,bool),X155),s(X154,X157)))) | s(bool,X158) = s(bool,happ(s(fun(X154,bool),X155),s(X154,sK20(X156,X155,X154)))) | s(bool,'T_0') = s(bool,X158)) )),
  inference(subsumption_resolution,[],[f2366,f247])).
fof(f2366,plain,(
  ( ! [X158,X156,X154,X157,X155] : (~p(s(bool,'T_0')) | ~p(s(bool,X156)) | p(s(bool,happ(s(fun(X154,bool),X155),s(X154,X157)))) | s(bool,X158) = s(bool,happ(s(fun(X154,bool),X155),s(X154,sK20(X156,X155,X154)))) | s(bool,'T_0') = s(bool,X158)) )),
  inference(superposition,[],[f295,f711])).
fof(f2725,plain,(
  ( ! [X152,X151,X149,X153,X150] : (~p(s(bool,X151)) | p(s(bool,happ(s(fun(X149,bool),X150),s(X149,X152)))) | s(bool,X153) = s(bool,happ(s(fun(X149,bool),X150),s(X149,sK19(X151,X150,X149)))) | s(bool,'T_0') = s(bool,X153)) )),
  inference(subsumption_resolution,[],[f2365,f247])).
fof(f2365,plain,(
  ( ! [X152,X151,X149,X153,X150] : (~p(s(bool,'T_0')) | ~p(s(bool,X151)) | p(s(bool,happ(s(fun(X149,bool),X150),s(X149,X152)))) | s(bool,X153) = s(bool,happ(s(fun(X149,bool),X150),s(X149,sK19(X151,X150,X149)))) | s(bool,'T_0') = s(bool,X153)) )),
  inference(superposition,[],[f293,f711])).
fof(f2724,plain,(
  ( ! [X144,X142,X145,X143] : (p(s(bool,'|dollar|exists_1'(s(fun(X142,bool),X143)))) | s(bool,X145) = s(bool,happ(s(fun(X142,bool),X143),s(X142,X144))) | s(bool,'T_0') = s(bool,X145)) )),
  inference(subsumption_resolution,[],[f2363,f247])).
fof(f2363,plain,(
  ( ! [X144,X142,X145,X143] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X142,bool),X143)))) | s(bool,X145) = s(bool,happ(s(fun(X142,bool),X143),s(X142,X144))) | s(bool,'T_0') = s(bool,X145)) )),
  inference(superposition,[],[f440,f711])).
fof(f2723,plain,(
  ( ! [X134,X140,X138,X136,X135,X141,X139,X137] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X134,fun(X134,bool)),X135),s('type.num.num',X139),s(X134,X138),s(X134,X140)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X134,fun(X134,bool)),X135),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X136),s('type.num.num',X139))),s(X134,X137),s(X134,X140)))) | s(bool,X141) = s(bool,'const.arithmetic.NRC_4'(s(fun(X134,fun(X134,bool)),X135),s('type.num.num',X136),s(X134,X137),s(X134,X138))) | s(bool,'T_0') = s(bool,X141)) )),
  inference(subsumption_resolution,[],[f2362,f247])).
fof(f2362,plain,(
  ( ! [X134,X140,X138,X136,X135,X141,X139,X137] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X134,fun(X134,bool)),X135),s('type.num.num',X139),s(X134,X138),s(X134,X140)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X134,fun(X134,bool)),X135),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X136),s('type.num.num',X139))),s(X134,X137),s(X134,X140)))) | s(bool,X141) = s(bool,'const.arithmetic.NRC_4'(s(fun(X134,fun(X134,bool)),X135),s('type.num.num',X136),s(X134,X137),s(X134,X138))) | s(bool,'T_0') = s(bool,X141)) )),
  inference(superposition,[],[f304,f711])).
fof(f2722,plain,(
  ( ! [X127,X132,X130,X128,X133,X131,X129] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X127,fun(X127,bool)),X128),s('type.num.num',X129),s(X127,X131),s(X127,sK23(X128,X129,X130,X131,X132,X127))))) | s(bool,X133) = s(bool,'const.arithmetic.NRC_4'(s(fun(X127,fun(X127,bool)),X128),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X129),s('type.num.num',X130))),s(X127,X131),s(X127,X132))) | s(bool,'T_0') = s(bool,X133)) )),
  inference(subsumption_resolution,[],[f2361,f247])).
fof(f2361,plain,(
  ( ! [X127,X132,X130,X128,X133,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X127,fun(X127,bool)),X128),s('type.num.num',X129),s(X127,X131),s(X127,sK23(X128,X129,X130,X131,X132,X127))))) | s(bool,X133) = s(bool,'const.arithmetic.NRC_4'(s(fun(X127,fun(X127,bool)),X128),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X129),s('type.num.num',X130))),s(X127,X131),s(X127,X132))) | s(bool,'T_0') = s(bool,X133)) )),
  inference(superposition,[],[f302,f711])).
fof(f2721,plain,(
  ( ! [X125,X123,X121,X126,X124,X122,X120] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X120,fun(X120,bool)),X121),s('type.num.num',X123),s(X120,sK23(X121,X122,X123,X124,X125,X120)),s(X120,X125)))) | s(bool,X126) = s(bool,'const.arithmetic.NRC_4'(s(fun(X120,fun(X120,bool)),X121),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X122),s('type.num.num',X123))),s(X120,X124),s(X120,X125))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(subsumption_resolution,[],[f2360,f247])).
fof(f2360,plain,(
  ( ! [X125,X123,X121,X126,X124,X122,X120] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X120,fun(X120,bool)),X121),s('type.num.num',X123),s(X120,sK23(X121,X122,X123,X124,X125,X120)),s(X120,X125)))) | s(bool,X126) = s(bool,'const.arithmetic.NRC_4'(s(fun(X120,fun(X120,bool)),X121),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X122),s('type.num.num',X123))),s(X120,X124),s(X120,X125))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f303,f711])).
fof(f2720,plain,(
  ( ! [X101,X99,X102,X100] : (~sP9(X99,X100,X101) | sP10(X99,X100,X101) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f2350,f247])).
fof(f2350,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,'T_0')) | ~sP9(X99,X100,X101) | sP10(X99,X100,X101) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f424,f711])).
fof(f2719,plain,(
  ( ! [X97,X95,X98,X96] : (~sP11(X95,X96,X97) | sP12(X95,X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f2349,f247])).
fof(f2349,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,'T_0')) | ~sP11(X95,X96,X97) | sP12(X95,X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f423,f711])).
fof(f2718,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X92)) | p(s(bool,X93)) | ~sP12(X91,X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f2348,f247])).
fof(f2348,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | ~p(s(bool,X92)) | p(s(bool,X93)) | ~sP12(X91,X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f420,f711])).
fof(f2717,plain,(
  ( ! [X90,X88,X87,X89] : (sP12(X87,X88,X89) | ~p(s(bool,X88)) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f2347,f247])).
fof(f2347,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,X88)) | sP12(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f419,f711])).
fof(f2716,plain,(
  ( ! [X85,X83,X86,X84] : (sP12(X83,X85,X84) | p(s(bool,X84)) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f2346,f247])).
fof(f2346,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X84)) | sP12(X83,X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f418,f711])).
fof(f2715,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | ~p(s(bool,X81)) | ~sP11(X81,X80,X79) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f2345,f247])).
fof(f2345,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | p(s(bool,X80)) | ~p(s(bool,X81)) | ~sP11(X81,X80,X79) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f413,f711])).
fof(f2714,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | p(s(bool,X77)) | ~sP10(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f2344,f247])).
fof(f2344,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X76)) | p(s(bool,X77)) | ~sP10(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f403,f711])).
fof(f2713,plain,(
  ( ! [X74,X72,X71,X73] : (sP10(X71,X72,X73) | ~p(s(bool,X72)) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f2343,f247])).
fof(f2343,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X72)) | sP10(X71,X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f402,f711])).
fof(f2712,plain,(
  ( ! [X70,X68,X69,X67] : (sP10(X67,X69,X68) | ~p(s(bool,X68)) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f2342,f247])).
fof(f2342,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X68)) | sP10(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f401,f711])).
fof(f2711,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f2341,f247])).
fof(f2341,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f396,f711])).
fof(f2710,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | p(s(bool,X61)) | ~sP8(X61,X60,X59) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f2340,f247])).
fof(f2340,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | p(s(bool,X61)) | ~sP8(X61,X60,X59) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f388,f711])).
fof(f2709,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP7(X56,X55,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f2339,f247])).
fof(f2339,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP7(X56,X55,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f379,f711])).
fof(f2708,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP4(X52,X51,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f2338,f247])).
fof(f2338,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP4(X52,X51,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f360,f711])).
fof(f2707,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP5(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f2337,f247])).
fof(f2337,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP5(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f356,f711])).
fof(f2706,plain,(
  ( ! [X45,X43,X46,X44] : (sP6(X44,X45,X43) | ~sP5(X44,X43,X45) | ~sP4(X45,X43,X44) | ~sP3(X43,X45,X44) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f2336,f247])).
fof(f2336,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | sP6(X44,X45,X43) | ~sP5(X44,X43,X45) | ~sP4(X45,X43,X44) | ~sP3(X43,X45,X44) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f353,f711])).
fof(f2705,plain,(
  ( ! [X39,X41,X42,X40] : (sP2(X40,X41,X39) | ~p(s(bool,X40)) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f2335,f247])).
fof(f2335,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | sP2(X40,X41,X39) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f344,f711])).
fof(f2704,plain,(
  ( ! [X37,X35,X38,X36] : (sP2(X36,X35,X37) | ~p(s(bool,X36)) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f2334,f247])).
fof(f2334,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | sP2(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f343,f711])).
fof(f2701,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f2330,f247])).
fof(f2330,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f315,f711])).
fof(f2700,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f2329,f247])).
fof(f2329,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f314,f711])).
fof(f2699,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f2328,f247])).
fof(f2328,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f308,f711])).
fof(f2698,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2324,f247])).
fof(f2324,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f266,f711])).
fof(f2319,plain,(
  ( ! [X156,X154,X152,X151,X155,X153,X150] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X151),s('type.num.num',X152),s(bool,X153),s(bool,X150)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X151),s('type.num.num',X154),s(bool,'T_0'),s(bool,X155)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X151),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X152),s('type.num.num',X154))),s(bool,X153),s(bool,X155)))) | s(bool,X150) = s(bool,X156) | s(bool,'T_0') = s(bool,X156)) )),
  inference(superposition,[],[f304,f711])).
fof(f2318,plain,(
  ( ! [X146,X144,X149,X147,X145,X143,X148] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X144),s('type.num.num',X145),s(bool,X143),s(bool,X146)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X144),s('type.num.num',X147),s(bool,X146),s(bool,X148)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X144),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X145),s('type.num.num',X147))),s(bool,'T_0'),s(bool,X148)))) | s(bool,X143) = s(bool,X149) | s(bool,'T_0') = s(bool,X149)) )),
  inference(superposition,[],[f304,f711])).
fof(f2317,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X138),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X139),s('type.num.num',X140))),s(bool,X141),s(bool,X137)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X138),s('type.num.num',X140),s(bool,sK23(X138,X139,X140,X141,'T_0',bool)),s(bool,'T_0')))) | s(bool,X137) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f303,f711])).
fof(f2316,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X132),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X133),s('type.num.num',X134))),s(bool,X131),s(bool,X135)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X132),s('type.num.num',X134),s(bool,sK23(X132,X133,X134,'T_0',X135,bool)),s(bool,X135)))) | s(bool,X131) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f303,f711])).
fof(f2315,plain,(
  ( ! [X127,X125,X130,X128,X126,X129] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X126),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X127),s('type.num.num',X128))),s(bool,X129),s(bool,X125)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X126),s('type.num.num',X127),s(bool,X129),s(bool,sK23(X126,X127,X128,X129,'T_0',bool))))) | s(bool,X125) = s(bool,X130) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f302,f711])).
fof(f2314,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X120),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X121),s('type.num.num',X122))),s(bool,X119),s(bool,X123)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X120),s('type.num.num',X121),s(bool,'T_0'),s(bool,sK23(X120,X121,X122,'T_0',X123,bool))))) | s(bool,X119) = s(bool,X124) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f302,f711])).
fof(f2280,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f262,f711])).
fof(f2212,plain,(
  ( ! [X255,X253,X254,X256] : (s(X255,f(s(bool,X256))) != s(X255,g(s(bool,X256))) | s(fun(bool,X255),X253) = s(fun(bool,X255),X254) | s(bool,'T_0') = s(bool,X256) | s(bool,'T_0') = s(bool,sK24(X253,X254,X255,bool))) )),
  inference(superposition,[],[f338,f711])).
fof(f2211,plain,(
  ( ! [X251,X249,X252,X250] : (~p(s(bool,happ(s(fun(bool,bool),X249),s(bool,X251)))) | p(s(bool,X250)) | p(s(bool,happ(s(fun(bool,bool),X249),s(bool,X252)))) | s(bool,'T_0') = s(bool,X251) | s(bool,'T_0') = s(bool,sK22(X249,X250,bool))) )),
  inference(superposition,[],[f300,f711])).
fof(f2210,plain,(
  ( ! [X246,X248,X247,X245] : (~p(s(bool,happ(s(fun(bool,bool),X245),s(bool,X247)))) | p(s(bool,X246)) | p(s(bool,happ(s(fun(bool,bool),X245),s(bool,X248)))) | s(bool,'T_0') = s(bool,X247) | s(bool,'T_0') = s(bool,sK21(X245,X246,bool))) )),
  inference(superposition,[],[f298,f711])).
fof(f2209,plain,(
  ( ! [X243,X241,X244,X242] : (~p(s(bool,happ(s(fun(bool,bool),X242),s(bool,X243)))) | ~p(s(bool,X241)) | p(s(bool,happ(s(fun(bool,bool),X242),s(bool,X244)))) | s(bool,'T_0') = s(bool,X243) | s(bool,'T_0') = s(bool,sK20(X241,X242,bool))) )),
  inference(superposition,[],[f295,f711])).
fof(f2208,plain,(
  ( ! [X239,X237,X240,X238] : (~p(s(bool,happ(s(fun(bool,bool),X238),s(bool,X239)))) | ~p(s(bool,X237)) | p(s(bool,happ(s(fun(bool,bool),X238),s(bool,X240)))) | s(bool,'T_0') = s(bool,X239) | s(bool,'T_0') = s(bool,sK19(X237,X238,bool))) )),
  inference(superposition,[],[f293,f711])).
fof(f2207,plain,(
  ( ! [X235,X236] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X235))) = s(bool,happ(s(fun(bool,bool),X235),s(bool,X236))) | s(bool,'T_0') = s(bool,X236) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X235)))) )),
  inference(superposition,[],[f289,f711])).
fof(f2206,plain,(
  ( ! [X231,X233,X234,X232] : (s(bool,X234) != s(bool,f(s(X231,sK24(X232,X233,bool,X231)))) | s(fun(X231,bool),X232) = s(fun(X231,bool),X233) | s(bool,'T_0') = s(bool,X234) | s(bool,'T_0') = s(bool,g(s(X231,sK24(X232,X233,bool,X231))))) )),
  inference(superposition,[],[f338,f711])).
fof(f2204,plain,(
  ( ! [X218,X222,X220,X219,X217,X223,X221] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X219),s('type.num.num',X220),s(bool,X221),s(bool,X218)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X219),s('type.num.num',X222),s(bool,X217),s(bool,X223)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X219),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X220),s('type.num.num',X222))),s(bool,X221),s(bool,X223)))) | s(bool,'T_0') = s(bool,X218) | s(bool,'T_0') = s(bool,X217)) )),
  inference(superposition,[],[f304,f711])).
fof(f2203,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X214] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X212),s('type.num.num',X213),s(bool,X211),s(bool,X214)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X212),s('type.num.num',X215),s(bool,X214),s(bool,X216)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X212),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X213),s('type.num.num',X215))),s(bool,X210),s(bool,X216)))) | s(bool,'T_0') = s(bool,X211) | s(bool,'T_0') = s(bool,X210)) )),
  inference(superposition,[],[f304,f711])).
fof(f2202,plain,(
  ( ! [X208,X206,X204,X209,X207,X205] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X206),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X207),s('type.num.num',X208))),s(bool,X209),s(bool,X205)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X206),s('type.num.num',X208),s(bool,sK23(X206,X207,X208,X209,X204,bool)),s(bool,X204)))) | s(bool,'T_0') = s(bool,X205) | s(bool,'T_0') = s(bool,X204)) )),
  inference(superposition,[],[f303,f711])).
fof(f2201,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X200),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X201),s('type.num.num',X202))),s(bool,X199),s(bool,X203)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X200),s('type.num.num',X202),s(bool,sK23(X200,X201,X202,X198,X203,bool)),s(bool,X203)))) | s(bool,'T_0') = s(bool,X199) | s(bool,'T_0') = s(bool,X198)) )),
  inference(superposition,[],[f303,f711])).
fof(f2200,plain,(
  ( ! [X196,X194,X192,X197,X195,X193] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X194),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X195),s('type.num.num',X196))),s(bool,X197),s(bool,X193)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X194),s('type.num.num',X195),s(bool,X197),s(bool,sK23(X194,X195,X196,X197,X192,bool))))) | s(bool,'T_0') = s(bool,X193) | s(bool,'T_0') = s(bool,X192)) )),
  inference(superposition,[],[f302,f711])).
fof(f2199,plain,(
  ( ! [X191,X189,X187,X190,X188,X186] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X188),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X189),s('type.num.num',X190))),s(bool,X187),s(bool,X191)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X188),s('type.num.num',X189),s(bool,X186),s(bool,sK23(X188,X189,X190,X186,X191,bool))))) | s(bool,'T_0') = s(bool,X187) | s(bool,'T_0') = s(bool,X186)) )),
  inference(superposition,[],[f302,f711])).
fof(f2130,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21))) | s(bool,X22) = s(bool,happ(s(fun(X20,bool),X21),s(X20,'const.min.@_1'(s(fun(X20,bool),X21))))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f711,f289])).
fof(f2118,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f711,f711])).
fof(f2103,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,X35) = s(bool,'const.arithmetic.NRC_4'(s(fun(X31,fun(X31,bool)),X32),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X31,X33),s(X31,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,happ(s(fun(X31,bool),happ(s(fun(X31,fun(X31,bool)),X32),s(X31,X33))),s(X31,X34)))) )),
  inference(superposition,[],[f711,f301])).
fof(f2102,plain,(
  ( ! [X30,X28,X29] : (s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(X28,bool),X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,happ(s(fun(X28,bool),X29),s(X28,'const.min.@_1'(s(fun(X28,bool),X29)))))) )),
  inference(superposition,[],[f711,f289])).
fof(f2100,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,X26) = s(bool,happ(s(fun(X22,bool),happ(s(fun(X22,fun(X22,bool)),X23),s(X22,X24))),s(X22,X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(X22,fun(X22,bool)),X23),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X22,X24),s(X22,X25)))) )),
  inference(superposition,[],[f711,f301])).
fof(f2045,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X15,fun(X15,bool)),X16),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X15,X17),s(X15,X18)))) | s(bool,X19) = s(bool,happ(s(fun(X15,bool),happ(s(fun(X15,fun(X15,bool)),X16),s(X15,X17))),s(X15,X18))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f769,f301])).
fof(f2042,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) | s(bool,X10) = s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X6,X8),s(X6,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f769,f301])).
fof(f1920,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X3),s(X2,X4))),s(X2,X5)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X2,fun(X2,bool)),X3),s('type.num.num',X6),s(X2,X5),s(X2,X7)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X2,fun(X2,bool)),X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6))),s(X2,X4),s(X2,X7))))) )),
  inference(superposition,[],[f304,f301])).
fof(f1917,plain,(
  ( ! [X39,X45,X43,X41,X38,X46,X44,X42,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X42),s('type.num.num',X43),s(bool,X44),s(bool,'const.arithmetic.NRC_4'(s(fun(X38,fun(X38,bool)),X39),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X38,X40),s(X38,X41)))))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X42),s('type.num.num',X45),s(bool,happ(s(fun(X38,bool),happ(s(fun(X38,fun(X38,bool)),X39),s(X38,X40))),s(X38,X41))),s(bool,X46)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X42),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X45))),s(bool,X44),s(bool,X46))))) )),
  inference(superposition,[],[f304,f301])).
fof(f1914,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X21),s('type.num.num',X24),s(bool,'const.arithmetic.NRC_4'(s(fun(X17,fun(X17,bool)),X18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X17,X19),s(X17,X20))),s(bool,X25)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X21),s('type.num.num',X22),s(bool,X23),s(bool,happ(s(fun(X17,bool),happ(s(fun(X17,fun(X17,bool)),X18),s(X17,X19))),s(X17,X20)))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X21),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num',X24))),s(bool,X23),s(bool,X25))))) )),
  inference(superposition,[],[f304,f301])).
fof(f1912,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X7),s('type.num.num',X10),s(bool,'F_0'),s(bool,X11)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X7),s('type.num.num',X8),s(bool,X9),s(bool,X6)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X10))),s(bool,X9),s(bool,X11)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f304,f255])).
fof(f1911,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num',X2),s(bool,X3),s(bool,'F_0')))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num',X4),s(bool,X0),s(bool,X5)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X4))),s(bool,X3),s(bool,X5)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f304,f255])).
fof(f1909,plain,(
  ( ! [X39,X45,X43,X41,X38,X46,X44,X42,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X42),s('type.num.num',X45),s(bool,X44),s(bool,X46)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X42),s('type.num.num',X43),s(bool,'const.arithmetic.NRC_4'(s(fun(X38,fun(X38,bool)),X39),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X38,X40),s(X38,X41))),s(bool,X44)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X42),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X45))),s(bool,happ(s(fun(X38,bool),happ(s(fun(X38,fun(X38,bool)),X39),s(X38,X40))),s(X38,X41))),s(bool,X46))))) )),
  inference(superposition,[],[f304,f301])).
fof(f1906,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X21),s('type.num.num',X24),s(bool,X23),s(bool,X25)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X21),s('type.num.num',X22),s(bool,happ(s(fun(X17,bool),happ(s(fun(X17,fun(X17,bool)),X18),s(X17,X19))),s(X17,X20))),s(bool,X23)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X21),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num',X24))),s(bool,'const.arithmetic.NRC_4'(s(fun(X17,fun(X17,bool)),X18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X17,X19),s(X17,X20))),s(bool,X25))))) )),
  inference(superposition,[],[f304,f301])).
fof(f1904,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X7),s('type.num.num',X10),s(bool,X9),s(bool,X11)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X7),s('type.num.num',X8),s(bool,X6),s(bool,X9)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X10))),s(bool,'F_0'),s(bool,X11)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f304,f255])).
fof(f1903,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num',X2),s(bool,'F_0'),s(bool,X3)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num',X4),s(bool,X3),s(bool,X5)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X4))),s(bool,X0),s(bool,X5)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f304,f255])).
fof(f1891,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun('type.num.num',fun('type.num.num',bool)),X41),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num',X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X40)))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun('type.num.num',fun('type.num.num',bool)),X41),s('type.num.num',X43),s('type.num.num',sK23(X41,X42,X43,X44,'const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),'type.num.num')),s('type.num.num','const.num.SUC_1'(s('type.num.num',X40))))))) )),
  inference(forward_demodulation,[],[f1885,f254])).
fof(f1885,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun('type.num.num',fun('type.num.num',bool)),X41),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num',X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X40)))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun('type.num.num',fun('type.num.num',bool)),X41),s('type.num.num',X43),s('type.num.num',sK23(X41,X42,X43,X44,'const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),'type.num.num')),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f303,f254])).
fof(f1884,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X36),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38))),s(bool,X39),s(bool,'const.arithmetic.NRC_4'(s(fun(X32,fun(X32,bool)),X33),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X32,X34),s(X32,X35)))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X36),s('type.num.num',X38),s(bool,sK23(X36,X37,X38,X39,happ(s(fun(X32,bool),happ(s(fun(X32,fun(X32,bool)),X33),s(X32,X34))),s(X32,X35)),bool)),s(bool,happ(s(fun(X32,bool),happ(s(fun(X32,fun(X32,bool)),X33),s(X32,X34))),s(X32,X35))))))) )),
  inference(superposition,[],[f303,f301])).
fof(f1890,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X28),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X30))),s(bool,X31),s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27)))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X28),s('type.num.num',X30),s(bool,sK23(X28,X29,X30,X31,happ(s(fun(X26,bool),X27),s(X26,'const.min.@_1'(s(fun(X26,bool),X27)))),bool)),s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27))))))) )),
  inference(forward_demodulation,[],[f1883,f289])).
fof(f1883,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X28),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X30))),s(bool,X31),s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27)))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X28),s('type.num.num',X30),s(bool,sK23(X28,X29,X30,X31,happ(s(fun(X26,bool),X27),s(X26,'const.min.@_1'(s(fun(X26,bool),X27)))),bool)),s(bool,happ(s(fun(X26,bool),X27),s(X26,'const.min.@_1'(s(fun(X26,bool),X27))))))))) )),
  inference(superposition,[],[f303,f289])).
fof(f1881,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X18),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))),s(bool,X21),s(bool,happ(s(fun(X14,bool),happ(s(fun(X14,fun(X14,bool)),X15),s(X14,X16))),s(X14,X17)))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X18),s('type.num.num',X20),s(bool,sK23(X18,X19,X20,X21,'const.arithmetic.NRC_4'(s(fun(X14,fun(X14,bool)),X15),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X14,X16),s(X14,X17)),bool)),s(bool,'const.arithmetic.NRC_4'(s(fun(X14,fun(X14,bool)),X15),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X14,X16),s(X14,X17))))))) )),
  inference(superposition,[],[f303,f301])).
fof(f1879,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X8))),s(bool,X9),s(bool,X5)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X6),s('type.num.num',X8),s(bool,sK23(X6,X7,X8,X9,'F_0',bool)),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f303,f255])).
fof(f1878,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s(bool,X4),s(bool,'F_0')))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num',X3),s(bool,sK23(X1,X2,X3,X4,X0,bool)),s(bool,X0)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f303,f255])).
fof(f1877,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun('type.num.num',fun('type.num.num',bool)),X41),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X40))),s('type.num.num',X44)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun('type.num.num',fun('type.num.num',bool)),X41),s('type.num.num',X43),s('type.num.num',sK23(X41,X42,X43,'const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X44,'type.num.num')),s('type.num.num',X44))))) )),
  inference(superposition,[],[f303,f254])).
fof(f1876,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X36),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38))),s(bool,'const.arithmetic.NRC_4'(s(fun(X32,fun(X32,bool)),X33),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X32,X34),s(X32,X35))),s(bool,X39)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X36),s('type.num.num',X38),s(bool,sK23(X36,X37,X38,happ(s(fun(X32,bool),happ(s(fun(X32,fun(X32,bool)),X33),s(X32,X34))),s(X32,X35)),X39,bool)),s(bool,X39))))) )),
  inference(superposition,[],[f303,f301])).
fof(f1875,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X28),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X30))),s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27))),s(bool,X31)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X28),s('type.num.num',X30),s(bool,sK23(X28,X29,X30,happ(s(fun(X26,bool),X27),s(X26,'const.min.@_1'(s(fun(X26,bool),X27)))),X31,bool)),s(bool,X31))))) )),
  inference(superposition,[],[f303,f289])).
fof(f1873,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X18),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))),s(bool,happ(s(fun(X14,bool),happ(s(fun(X14,fun(X14,bool)),X15),s(X14,X16))),s(X14,X17))),s(bool,X21)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X18),s('type.num.num',X20),s(bool,sK23(X18,X19,X20,'const.arithmetic.NRC_4'(s(fun(X14,fun(X14,bool)),X15),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X14,X16),s(X14,X17)),X21,bool)),s(bool,X21))))) )),
  inference(superposition,[],[f303,f301])).
fof(f1871,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X8))),s(bool,X5),s(bool,X9)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X6),s('type.num.num',X8),s(bool,sK23(X6,X7,X8,'F_0',X9,bool)),s(bool,X9)))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f303,f255])).
fof(f1870,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s(bool,'F_0'),s(bool,X4)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num',X3),s(bool,sK23(X1,X2,X3,X0,X4,bool)),s(bool,X4)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f303,f255])).
fof(f1887,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s(X1,X4),s(X1,X5)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s(X1,sK23(X2,X3,'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X4,X5,X1)),s(X1,X5))))) )),
  inference(forward_demodulation,[],[f1868,f254])).
fof(f1868,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s(X1,X4),s(X1,X5)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X1,sK23(X2,X3,'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X4,X5,X1)),s(X1,X5))))) )),
  inference(superposition,[],[f303,f254])).
fof(f1867,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X3))),s(X1,X4),s(X1,X5)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num',X3),s(X1,sK23(X2,'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X3,X4,X5,X1)),s(X1,X5))))) )),
  inference(superposition,[],[f303,f254])).
fof(f1838,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun('type.num.num',fun('type.num.num',bool)),X41),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num',X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X40)))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun('type.num.num',fun('type.num.num',bool)),X41),s('type.num.num',X42),s('type.num.num',X44),s('type.num.num',sK23(X41,X42,X43,X44,'const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),'type.num.num')))))) )),
  inference(superposition,[],[f302,f254])).
fof(f1837,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X36),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38))),s(bool,X39),s(bool,'const.arithmetic.NRC_4'(s(fun(X32,fun(X32,bool)),X33),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X32,X34),s(X32,X35)))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X36),s('type.num.num',X37),s(bool,X39),s(bool,sK23(X36,X37,X38,X39,happ(s(fun(X32,bool),happ(s(fun(X32,fun(X32,bool)),X33),s(X32,X34))),s(X32,X35)),bool)))))) )),
  inference(superposition,[],[f302,f301])).
fof(f1836,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X28),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X30))),s(bool,X31),s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27)))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X28),s('type.num.num',X29),s(bool,X31),s(bool,sK23(X28,X29,X30,X31,happ(s(fun(X26,bool),X27),s(X26,'const.min.@_1'(s(fun(X26,bool),X27)))),bool)))))) )),
  inference(superposition,[],[f302,f289])).
fof(f1834,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X18),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))),s(bool,X21),s(bool,happ(s(fun(X14,bool),happ(s(fun(X14,fun(X14,bool)),X15),s(X14,X16))),s(X14,X17)))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X18),s('type.num.num',X19),s(bool,X21),s(bool,sK23(X18,X19,X20,X21,'const.arithmetic.NRC_4'(s(fun(X14,fun(X14,bool)),X15),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X14,X16),s(X14,X17)),bool)))))) )),
  inference(superposition,[],[f302,f301])).
fof(f1832,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X8))),s(bool,X9),s(bool,X5)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X6),s('type.num.num',X7),s(bool,X9),s(bool,sK23(X6,X7,X8,X9,'F_0',bool))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f302,f255])).
fof(f1831,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s(bool,X4),s(bool,'F_0')))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num',X2),s(bool,X4),s(bool,sK23(X1,X2,X3,X4,X0,bool))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f302,f255])).
fof(f1844,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun('type.num.num',fun('type.num.num',bool)),X41),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X40))),s('type.num.num',X44)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun('type.num.num',fun('type.num.num',bool)),X41),s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num',X40))),s('type.num.num',sK23(X41,X42,X43,'const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X44,'type.num.num')))))) )),
  inference(forward_demodulation,[],[f1830,f254])).
fof(f1830,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun('type.num.num',fun('type.num.num',bool)),X41),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X40))),s('type.num.num',X44)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun('type.num.num',fun('type.num.num',bool)),X41),s('type.num.num',X42),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',sK23(X41,X42,X43,'const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X44,'type.num.num')))))) )),
  inference(superposition,[],[f302,f254])).
fof(f1829,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X36),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38))),s(bool,'const.arithmetic.NRC_4'(s(fun(X32,fun(X32,bool)),X33),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X32,X34),s(X32,X35))),s(bool,X39)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X36),s('type.num.num',X37),s(bool,happ(s(fun(X32,bool),happ(s(fun(X32,fun(X32,bool)),X33),s(X32,X34))),s(X32,X35))),s(bool,sK23(X36,X37,X38,happ(s(fun(X32,bool),happ(s(fun(X32,fun(X32,bool)),X33),s(X32,X34))),s(X32,X35)),X39,bool)))))) )),
  inference(superposition,[],[f302,f301])).
fof(f1843,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X28),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X30))),s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27))),s(bool,X31)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X28),s('type.num.num',X29),s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27))),s(bool,sK23(X28,X29,X30,happ(s(fun(X26,bool),X27),s(X26,'const.min.@_1'(s(fun(X26,bool),X27)))),X31,bool)))))) )),
  inference(forward_demodulation,[],[f1828,f289])).
fof(f1828,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X28),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X30))),s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27))),s(bool,X31)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X28),s('type.num.num',X29),s(bool,happ(s(fun(X26,bool),X27),s(X26,'const.min.@_1'(s(fun(X26,bool),X27))))),s(bool,sK23(X28,X29,X30,happ(s(fun(X26,bool),X27),s(X26,'const.min.@_1'(s(fun(X26,bool),X27)))),X31,bool)))))) )),
  inference(superposition,[],[f302,f289])).
fof(f1826,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X18),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))),s(bool,happ(s(fun(X14,bool),happ(s(fun(X14,fun(X14,bool)),X15),s(X14,X16))),s(X14,X17))),s(bool,X21)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X18),s('type.num.num',X19),s(bool,'const.arithmetic.NRC_4'(s(fun(X14,fun(X14,bool)),X15),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X14,X16),s(X14,X17))),s(bool,sK23(X18,X19,X20,'const.arithmetic.NRC_4'(s(fun(X14,fun(X14,bool)),X15),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X14,X16),s(X14,X17)),X21,bool)))))) )),
  inference(superposition,[],[f302,f301])).
fof(f1824,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X8))),s(bool,X5),s(bool,X9)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X6),s('type.num.num',X7),s(bool,'F_0'),s(bool,sK23(X6,X7,X8,'F_0',X9,bool))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f302,f255])).
fof(f1823,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s(bool,'F_0'),s(bool,X4)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X1),s('type.num.num',X2),s(bool,X0),s(bool,sK23(X1,X2,X3,X0,X4,bool))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f302,f255])).
fof(f1821,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s(X1,X4),s(X1,X5)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num',X3),s(X1,X4),s(X1,sK23(X2,X3,'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X4,X5,X1)))))) )),
  inference(superposition,[],[f302,f254])).
fof(f1840,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X3))),s(X1,X4),s(X1,X5)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s(X1,X4),s(X1,sK23(X2,'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X3,X4,X5,X1)))))) )),
  inference(forward_demodulation,[],[f1820,f254])).
fof(f1820,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X3))),s(X1,X4),s(X1,X5)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X1,fun(X1,bool)),X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X1,X4),s(X1,sK23(X2,'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X3,X4,X5,X1)))))) )),
  inference(superposition,[],[f302,f254])).
fof(f1797,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (~sP9('const.arithmetic.NRC_4'(s(fun(X157,fun(X157,bool)),X158),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X157,X159),s(X157,X160)),X161,X162) | ~p(s(bool,happ(s(fun(X157,bool),happ(s(fun(X157,fun(X157,bool)),X158),s(X157,X159))),s(X157,X160)))) | sP10('const.arithmetic.NRC_4'(s(fun(X157,fun(X157,bool)),X158),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X157,X159),s(X157,X160)),X161,X162)) )),
  inference(superposition,[],[f424,f301])).
fof(f1796,plain,(
  ( ! [X156,X154,X152,X151,X155,X153] : (~sP11('const.arithmetic.NRC_4'(s(fun(X151,fun(X151,bool)),X152),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X151,X153),s(X151,X154)),X155,X156) | ~p(s(bool,happ(s(fun(X151,bool),happ(s(fun(X151,fun(X151,bool)),X152),s(X151,X153))),s(X151,X154)))) | sP12('const.arithmetic.NRC_4'(s(fun(X151,fun(X151,bool)),X152),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X151,X153),s(X151,X154)),X155,X156)) )),
  inference(superposition,[],[f423,f301])).
fof(f1791,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (~sP10('const.arithmetic.NRC_4'(s(fun(X121,fun(X121,bool)),X122),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X121,X123),s(X121,X124)),X126,X125) | p(s(bool,X125)) | p(s(bool,X126)) | ~p(s(bool,happ(s(fun(X121,bool),happ(s(fun(X121,fun(X121,bool)),X122),s(X121,X123))),s(X121,X124))))) )),
  inference(superposition,[],[f403,f301])).
fof(f1788,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (~sP9('const.arithmetic.NRC_4'(s(fun(X103,fun(X103,bool)),X104),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X103,X105),s(X103,X106)),X107,X108) | p(s(bool,X107)) | p(s(bool,X108)) | ~p(s(bool,happ(s(fun(X103,bool),happ(s(fun(X103,fun(X103,bool)),X104),s(X103,X105))),s(X103,X106))))) )),
  inference(superposition,[],[f396,f301])).
fof(f1783,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (~sP5(X77,'const.arithmetic.NRC_4'(s(fun(X73,fun(X73,bool)),X74),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X73,X75),s(X73,X76)),X78) | sP6(X77,X78,'const.arithmetic.NRC_4'(s(fun(X73,fun(X73,bool)),X74),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X73,X75),s(X73,X76))) | ~p(s(bool,happ(s(fun(X73,bool),happ(s(fun(X73,fun(X73,bool)),X74),s(X73,X75))),s(X73,X76)))) | ~sP4(X78,'const.arithmetic.NRC_4'(s(fun(X73,fun(X73,bool)),X74),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X73,X75),s(X73,X76)),X77) | ~sP3('const.arithmetic.NRC_4'(s(fun(X73,fun(X73,bool)),X74),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X73,X75),s(X73,X76)),X78,X77)) )),
  inference(superposition,[],[f353,f301])).
fof(f1777,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,happ(s(fun(X40,bool),happ(s(fun(X40,fun(X40,bool)),X41),s(X40,X42))),s(X40,X43)))) | sP0('const.arithmetic.NRC_4'(s(fun(X40,fun(X40,bool)),X41),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X40,X42),s(X40,X43)),X44)) )),
  inference(superposition,[],[f315,f301])).
fof(f1776,plain,(
  ( ! [X39,X37,X35,X38,X36] : (~p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) | sP0(X39,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) )),
  inference(superposition,[],[f314,f301])).
fof(f1773,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),happ(s(fun(X20,fun(X20,bool)),X21),s(X20,X22))),s(X20,X23)))) | p(s(bool,X24)) | s(bool,X24) = s(bool,'const.arithmetic.NRC_4'(s(fun(X20,fun(X20,bool)),X21),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X20,X22),s(X20,X23)))) )),
  inference(superposition,[],[f305,f301])).
fof(f1772,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),X17),s(X16,X18))),s(X16,X19)))) | s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(X16,fun(X16,bool)),X17),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X16,X18),s(X16,X19)))) )),
  inference(superposition,[],[f267,f301])).
fof(f1769,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X6))),s(X4,X7))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X4,fun(X4,bool)),X5),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X4,X6),s(X4,X7))))) )),
  inference(superposition,[],[f262,f301])).
fof(f1766,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (~sP9(happ(s(fun(X184,bool),happ(s(fun(X184,fun(X184,bool)),X185),s(X184,X186))),s(X184,X187)),X188,X189) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X184,fun(X184,bool)),X185),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X184,X186),s(X184,X187)))) | sP10(happ(s(fun(X184,bool),happ(s(fun(X184,fun(X184,bool)),X185),s(X184,X186))),s(X184,X187)),X188,X189)) )),
  inference(superposition,[],[f424,f301])).
fof(f1765,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (~sP11(happ(s(fun(X178,bool),happ(s(fun(X178,fun(X178,bool)),X179),s(X178,X180))),s(X178,X181)),X182,X183) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X178,fun(X178,bool)),X179),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X178,X180),s(X178,X181)))) | sP12(happ(s(fun(X178,bool),happ(s(fun(X178,fun(X178,bool)),X179),s(X178,X180))),s(X178,X181)),X182,X183)) )),
  inference(superposition,[],[f423,f301])).
fof(f1762,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X160,fun(X160,bool)),X161),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X160,X162),s(X160,X163)))) | p(s(bool,X164)) | sP12(happ(s(fun(X160,bool),happ(s(fun(X160,fun(X160,bool)),X161),s(X160,X162))),s(X160,X163)),X165,X164)) )),
  inference(superposition,[],[f418,f301])).
fof(f1760,plain,(
  ( ! [X152,X151,X149,X153,X150,X148] : (~sP10(happ(s(fun(X148,bool),happ(s(fun(X148,fun(X148,bool)),X149),s(X148,X150))),s(X148,X151)),X153,X152) | p(s(bool,X152)) | p(s(bool,X153)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X148,fun(X148,bool)),X149),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X148,X150),s(X148,X151))))) )),
  inference(superposition,[],[f403,f301])).
fof(f1757,plain,(
  ( ! [X134,X132,X130,X135,X133,X131] : (~sP9(happ(s(fun(X130,bool),happ(s(fun(X130,fun(X130,bool)),X131),s(X130,X132))),s(X130,X133)),X134,X135) | p(s(bool,X134)) | p(s(bool,X135)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X130,fun(X130,bool)),X131),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X130,X132),s(X130,X133))))) )),
  inference(superposition,[],[f396,f301])).
fof(f1752,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (~sP5(X104,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103)),X105) | sP6(X104,X105,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X100,fun(X100,bool)),X101),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X100,X102),s(X100,X103)))) | ~sP4(X105,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103)),X104) | ~sP3(happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103)),X105,X104)) )),
  inference(superposition,[],[f353,f301])).
fof(f1748,plain,(
  ( ! [X80,X78,X81,X79,X77] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X77,fun(X77,bool)),X78),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X77,X79),s(X77,X80)))) | p(s(bool,X81)) | sP1(X81,happ(s(fun(X77,bool),happ(s(fun(X77,fun(X77,bool)),X78),s(X77,X79))),s(X77,X80)))) )),
  inference(superposition,[],[f333,f301])).
fof(f1747,plain,(
  ( ! [X76,X74,X72,X75,X73] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X72,fun(X72,bool)),X73),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X72,X74),s(X72,X75)))) | p(s(bool,X76)) | sP1(happ(s(fun(X72,bool),happ(s(fun(X72,fun(X72,bool)),X73),s(X72,X74))),s(X72,X75)),X76)) )),
  inference(superposition,[],[f330,f301])).
fof(f1746,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X67,fun(X67,bool)),X68),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X67,X69),s(X67,X70)))) | sP0(happ(s(fun(X67,bool),happ(s(fun(X67,fun(X67,bool)),X68),s(X67,X69))),s(X67,X70)),X71)) )),
  inference(superposition,[],[f315,f301])).
fof(f1745,plain,(
  ( ! [X66,X64,X62,X65,X63] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X62,fun(X62,bool)),X63),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X62,X64),s(X62,X65)))) | sP0(X66,happ(s(fun(X62,bool),happ(s(fun(X62,fun(X62,bool)),X63),s(X62,X64))),s(X62,X65)))) )),
  inference(superposition,[],[f314,f301])).
fof(f1742,plain,(
  ( ! [X47,X50,X48,X51,X49] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X47,fun(X47,bool)),X48),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X47,X49),s(X47,X50)))) | p(s(bool,X51)) | s(bool,X51) = s(bool,happ(s(fun(X47,bool),happ(s(fun(X47,fun(X47,bool)),X48),s(X47,X49))),s(X47,X50)))) )),
  inference(superposition,[],[f305,f301])).
fof(f1741,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X43,fun(X43,bool)),X44),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X43,X45),s(X43,X46)))) | s(bool,'F_0') = s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X45))),s(X43,X46)))) )),
  inference(superposition,[],[f267,f301])).
fof(f1739,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X35,fun(X35,bool)),X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X35,X37),s(X35,X38)))) | s(bool,'T_0') = s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38)))) )),
  inference(superposition,[],[f263,f301])).
fof(f1738,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'T_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(X31,fun(X31,bool)),X32),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X31,X33),s(X31,X34))) | p(s(bool,happ(s(fun(X31,bool),happ(s(fun(X31,fun(X31,bool)),X32),s(X31,X33))),s(X31,X34))))) )),
  inference(superposition,[],[f262,f301])).
fof(f1736,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X22,fun(X22,bool)),X23),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X22,X24),s(X22,sK22(happ(s(fun(X22,fun(X22,bool)),X23),s(X22,X24)),X25,X22))))) | p(s(bool,X25)) | p(s(bool,happ(s(fun(X22,bool),happ(s(fun(X22,fun(X22,bool)),X23),s(X22,X24))),s(X22,X26))))) )),
  inference(superposition,[],[f300,f301])).
fof(f1735,plain,(
  ( ! [X21,X19,X17,X20,X18] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X17,fun(X17,bool)),X18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X17,X19),s(X17,sK21(happ(s(fun(X17,fun(X17,bool)),X18),s(X17,X19)),X20,X17))))) | p(s(bool,X20)) | p(s(bool,happ(s(fun(X17,bool),happ(s(fun(X17,fun(X17,bool)),X18),s(X17,X19))),s(X17,X21))))) )),
  inference(superposition,[],[f298,f301])).
fof(f1731,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))))))) )),
  inference(superposition,[],[f440,f301])).
fof(f1730,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X6))),s(X4,X7))) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(X4,fun(X4,bool)),X5),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X4,X6),s(X4,X7)))) )),
  inference(superposition,[],[f301,f255])).
fof(f1713,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(X7,fun(X7,bool)),X8),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X7,X9),s(X7,X10))) | s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X9))),s(X7,X10)))) )),
  inference(superposition,[],[f301,f255])).
fof(f1709,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('T_0','const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3)),happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) )),
  inference(unit_resulting_resolution,[],[f247,f301,f369])).
fof(f1711,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('T_0',happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)),'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X3)))) )),
  inference(unit_resulting_resolution,[],[f247,f301,f369])).
fof(f1677,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK22(X0,X1,bool))) )),
  inference(superposition,[],[f300,f255])).
fof(f1664,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK21(X0,X1,bool))) )),
  inference(superposition,[],[f298,f255])).
fof(f1653,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | ~p(s(bool,X0)) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK20(X0,X1,bool))) )),
  inference(superposition,[],[f295,f255])).
fof(f1646,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | ~p(s(bool,X0)) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK19(X0,X1,bool))) )),
  inference(superposition,[],[f293,f255])).
fof(f1633,plain,(
  ( ! [X103,X101,X102,X104] : (sP10(happ(s(fun(X101,bool),X102),s(X101,'const.min.@_1'(s(fun(X101,bool),X102)))),X103,X104) | ~sP9(happ(s(fun(X101,bool),X102),s(X101,'const.min.@_1'(s(fun(X101,bool),X102)))),X103,X104) | ~p(s(bool,'|dollar|exists_1'(s(fun(X101,bool),X102))))) )),
  inference(superposition,[],[f424,f289])).
fof(f1632,plain,(
  ( ! [X99,X97,X100,X98] : (sP12(happ(s(fun(X97,bool),X98),s(X97,'const.min.@_1'(s(fun(X97,bool),X98)))),X99,X100) | ~sP11(happ(s(fun(X97,bool),X98),s(X97,'const.min.@_1'(s(fun(X97,bool),X98)))),X99,X100) | ~p(s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))))) )),
  inference(superposition,[],[f423,f289])).
fof(f1631,plain,(
  ( ! [X94,X95,X93,X96] : (~sP12(happ(s(fun(X93,bool),X94),s(X93,'const.min.@_1'(s(fun(X93,bool),X94)))),X96,X95) | ~p(s(bool,X95)) | p(s(bool,X96)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X93,bool),X94))))) )),
  inference(superposition,[],[f420,f289])).
fof(f1630,plain,(
  ( ! [X92,X90,X91,X89] : (sP12(happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90)))),X91,X92) | ~p(s(bool,X91)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X89,bool),X90))))) )),
  inference(superposition,[],[f419,f289])).
fof(f1629,plain,(
  ( ! [X88,X87,X85,X86] : (sP12(happ(s(fun(X85,bool),X86),s(X85,'const.min.@_1'(s(fun(X85,bool),X86)))),X88,X87) | p(s(bool,X87)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X85,bool),X86))))) )),
  inference(superposition,[],[f418,f289])).
fof(f1627,plain,(
  ( ! [X80,X78,X79,X77] : (~sP10(happ(s(fun(X77,bool),X78),s(X77,'const.min.@_1'(s(fun(X77,bool),X78)))),X80,X79) | p(s(bool,X79)) | p(s(bool,X80)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X77,bool),X78))))) )),
  inference(superposition,[],[f403,f289])).
fof(f1626,plain,(
  ( ! [X76,X74,X75,X73] : (sP10(happ(s(fun(X73,bool),X74),s(X73,'const.min.@_1'(s(fun(X73,bool),X74)))),X75,X76) | ~p(s(bool,X75)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X73,bool),X74))))) )),
  inference(superposition,[],[f402,f289])).
fof(f1625,plain,(
  ( ! [X70,X72,X71,X69] : (sP10(happ(s(fun(X69,bool),X70),s(X69,'const.min.@_1'(s(fun(X69,bool),X70)))),X72,X71) | ~p(s(bool,X71)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X69,bool),X70))))) )),
  inference(superposition,[],[f401,f289])).
fof(f1624,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'|dollar|exists_1'(s(fun(X65,bool),X66)))) | p(s(bool,X67)) | p(s(bool,X68)) | ~sP9(happ(s(fun(X65,bool),X66),s(X65,'const.min.@_1'(s(fun(X65,bool),X66)))),X67,X68)) )),
  inference(superposition,[],[f396,f289])).
fof(f1619,plain,(
  ( ! [X47,X45,X48,X46] : (sP6(X47,X48,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X45,bool),X46)))) | ~sP5(X47,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X48) | ~sP4(X48,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X47) | ~sP3(happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X48,X47)) )),
  inference(superposition,[],[f353,f289])).
fof(f1617,plain,(
  ( ! [X39,X37,X38,X40] : (sP2(X39,happ(s(fun(X37,bool),X38),s(X37,'const.min.@_1'(s(fun(X37,bool),X38)))),X40) | ~p(s(bool,X39)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X37,bool),X38))))) )),
  inference(superposition,[],[f343,f289])).
fof(f1615,plain,(
  ( ! [X30,X31,X32] : (sP1(X32,happ(s(fun(X30,bool),X31),s(X30,'const.min.@_1'(s(fun(X30,bool),X31))))) | p(s(bool,X32)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X30,bool),X31))))) )),
  inference(superposition,[],[f333,f289])).
fof(f1614,plain,(
  ( ! [X28,X29,X27] : (sP1(happ(s(fun(X27,bool),X28),s(X27,'const.min.@_1'(s(fun(X27,bool),X28)))),X29) | p(s(bool,X29)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X27,bool),X28))))) )),
  inference(superposition,[],[f330,f289])).
fof(f1613,plain,(
  ( ! [X26,X24,X25] : (sP0(happ(s(fun(X24,bool),X25),s(X24,'const.min.@_1'(s(fun(X24,bool),X25)))),X26) | ~p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25))))) )),
  inference(superposition,[],[f315,f289])).
fof(f1612,plain,(
  ( ! [X23,X21,X22] : (sP0(X23,happ(s(fun(X21,bool),X22),s(X21,'const.min.@_1'(s(fun(X21,bool),X22))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X22))))) )),
  inference(superposition,[],[f314,f289])).
fof(f1602,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3)))) )),
  inference(superposition,[],[f289,f255])).
fof(f1601,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f289,f255])).
fof(f1598,plain,(
  ( ! [X0,X1] : (sP6('T_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f247,f289,f369])).
fof(f1600,plain,(
  ( ! [X0,X1] : (sP6('T_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f247,f289,f369])).
fof(f1185,plain,(
  sP6('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f452,f447,f444,f247,f353])).
fof(f894,plain,(
  sP8('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f453,f247,f247,f375])).
fof(f804,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f247,f487,f370])).
fof(f656,plain,(
  ( ! [X0] : (sP12('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f247,f247,f419])).
fof(f638,plain,(
  ( ! [X0] : (sP10('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f247,f247,f402])).
fof(f610,plain,(
  ( ! [X0] : (sP10('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f247,f401])).
fof(f571,plain,(
  ( ! [X0] : (sP2('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f247,f344])).
fof(f499,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f247,f247,f329])).
fof(f462,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f247,f426])).
fof(f443,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f247,f315])).
fof(f442,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f314])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f235])).
fof(f298,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0)))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0)))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22])],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X13,X17] : (! [X8] : (p(s(bool,X13)) | p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X8))))) <=> (p(s(bool,X13)) | ! [X8] : p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X8))))))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.bool.RIGHT_FORALL_OR_THM')).
fof(f300,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f189])).
fof(f3220,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197,f3199,f3201,f3203,f3205,f3207,f3209,f3211,f3213,f3215,f3217,f3219])).
fof(f3218,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197,f3199,f3201,f3203,f3205,f3207,f3209,f3211,f3213,f3215,f3217])).
fof(f3216,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197,f3199,f3201,f3203,f3205,f3207,f3209,f3211,f3213,f3215])).
fof(f3214,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197,f3199,f3201,f3203,f3205,f3207,f3209,f3211,f3213])).
fof(f3212,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197,f3199,f3201,f3203,f3205,f3207,f3209,f3211])).
fof(f3210,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197,f3199,f3201,f3203,f3205,f3207,f3209])).
fof(f3208,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197,f3199,f3201,f3203,f3205,f3207])).
fof(f3206,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197,f3199,f3201,f3203,f3205])).
fof(f3204,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197,f3199,f3201,f3203])).
fof(f3202,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197,f3199,f3201])).
fof(f3200,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197,f3199])).
fof(f3198,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195,f3197])).
fof(f3196,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193,f3195])).
fof(f3194,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191,f3193])).
fof(f3192,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189,f3191])).
fof(f3190,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187,f3189])).
fof(f3188,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185,f3187])).
fof(f3186,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183,f3185])).
fof(f3184,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181,f3183])).
fof(f3182,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179,f3181])).
fof(f3180,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177,f3179])).
fof(f3178,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175,f3177])).
fof(f3176,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173,f3175])).
fof(f3174,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171,f3173])).
fof(f3172,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169,f3171])).
fof(f3170,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167,f3169])).
fof(f3168,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165,f3167])).
fof(f3166,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163,f3165])).
fof(f3164,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161,f3163])).
fof(f3162,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159,f3161])).
fof(f3160,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157,f3159])).
fof(f3158,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155,f3157])).
fof(f3156,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153,f3155])).
fof(f3154,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151,f3153])).
fof(f3152,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149,f3151])).
fof(f3150,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147,f3149])).
fof(f3148,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145,f3147])).
fof(f3146,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143,f3145])).
fof(f3144,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141,f3143])).
fof(f3142,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139,f3141])).
fof(f3140,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065,f3139])).
fof(f3066,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063,f3065])).
fof(f3064,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061,f3063])).
fof(f3062,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059,f3061])).
fof(f3060,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057,f3059])).
fof(f3058,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055,f3057])).
fof(f3056,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053,f3055])).
fof(f3054,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051,f3053])).
fof(f3052,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049,f3051])).
fof(f3050,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047,f3049])).
fof(f3048,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045,f3047])).
fof(f3046,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043,f3045])).
fof(f3044,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041,f3043])).
fof(f3042,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039,f3041])).
fof(f3040,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037,f3039])).
fof(f3038,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035,f3037])).
fof(f3036,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033,f3035])).
fof(f3034,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031,f3033])).
fof(f3032,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029,f3031])).
fof(f3030,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027,f3029])).
fof(f3028,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025,f3027])).
fof(f3026,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993,f3025])).
fof(f3001,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3000,f247])).
fof(f3000,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(forward_demodulation,[],[f2999,f2645])).
fof(f2999,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X12,fun(X12,bool)),X13),s(X12,X14))),s(X12,X15))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2786,f2645])).
fof(f2786,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(X12,fun(X12,bool)),X13),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X12,X14),s(X12,X15))) | ~p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X12,fun(X12,bool)),X13),s(X12,X14))),s(X12,X15))))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1771])).
fof(f2998,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f2997,f247])).
fof(f2997,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(forward_demodulation,[],[f2996,f2645])).
fof(f2996,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X39,fun(X39,bool)),X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X39,X41),s(X39,X42))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2784,f2645])).
fof(f2784,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') != s(bool,happ(s(fun(X39,bool),happ(s(fun(X39,fun(X39,bool)),X40),s(X39,X41))),s(X39,X42))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X39,fun(X39,bool)),X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X39,X41),s(X39,X42))))) ) | $spl10),
  inference(backward_demodulation,[],[f2645,f1740])).
fof(f2995,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f2772,f247])).
fof(f2772,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(backward_demodulation,[],[f2645,f439])).
fof(f2994,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f246,f245,f244,f247,f439,f254,f255,f263,f262,f267,f266,f289,f440,f295,f293,f300,f298,f301,f304,f303,f302,f308,f305,f315,f314,f441,f333,f330,f329,f328,f338,f344,f343,f342,f426,f340,f339,f355,f354,f353,f352,f351,f350,f349,f359,f358,f357,f356,f363,f362,f361,f360,f367,f366,f365,f364,f371,f370,f369,f368,f425,f375,f374,f373,f372,f382,f381,f380,f379,f388,f385,f384,f383,f395,f394,f393,f424,f391,f390,f389,f399,f398,f397,f396,f403,f402,f401,f400,f412,f411,f410,f423,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f487,f498,f499,f502,f507,f516,f557,f562,f571,f578,f581,f586,f594,f600,f604,f610,f638,f644,f645,f648,f656,f711,f769,f786,f785,f806,f804,f835,f857,f869,f877,f890,f894,f908,f920,f972,f990,f1003,f1116,f1130,f1144,f1185,f1227,f1600,f1599,f1598,f1597,f1601,f1602,f1612,f1613,f1614,f1615,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1646,f1653,f1664,f1677,f1696,f1697,f1711,f1710,f1709,f1708,f1712,f1713,f1730,f1731,f1733,f1734,f1735,f1736,f1738,f1739,f1740,f1741,f1742,f1744,f1745,f1746,f1747,f1748,f1750,f1751,f1752,f1753,f1754,f1755,f1756,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1766,f1769,f1770,f1771,f1772,f1773,f1775,f1776,f1777,f1778,f1779,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1840,f1821,f1822,f1823,f1824,f1826,f1843,f1829,f1844,f1831,f1832,f1834,f1836,f1837,f1838,f1867,f1887,f1869,f1870,f1871,f1873,f1875,f1876,f1877,f1878,f1879,f1881,f1890,f1884,f1891,f1903,f1904,f1906,f1909,f1911,f1912,f1914,f1917,f1920,f2042,f2045,f2100,f2102,f2103,f2118,f2130,f2199,f2200,f2201,f2202,f2203,f2204,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2280,f2643,f2647,f2650,f2654,f2657,f2660,f2663,f2666,f2669,f2672,f2675,f2678,f2681,f2684,f2687,f2690,f2693,f2696,f2314,f2315,f2316,f2317,f2318,f2319,f2698,f2699,f2700,f2701,f2702,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2718,f2719,f2720,f2721,f2722,f2723,f2724,f2725,f2726,f2727,f2728,f2372,f2373,f2374,f2375,f2376,f2377,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2734,f2736,f2738,f2741,f2743,f2745,f2747,f2749,f2751,f2753,f2755,f2757,f2759,f2761,f2763,f2765,f2767,f2769,f2645,f2993])).
fof(f2697,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~sP9('T_0',X1,X2)) ) | $spl44),
  inference(cnf_transformation,[],[f2697_D])).
fof(f2697_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~sP9('T_0',X1,X2)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2769,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f2768])).
fof(f2768,plain,(
  ( ! [X107,X105,X106,X104] : (~sP9('T_0',X105,X106) | sP10('T_0',X105,X106) | s(bool,X104) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(subsumption_resolution,[],[f2416,f769])).
fof(f2416,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,X104)) | ~sP9('T_0',X105,X106) | sP10('T_0',X105,X106) | s(bool,X104) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f424,f711])).
fof(f2694,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~sP11('T_0',X1,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f2694_D])).
fof(f2694_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~sP11('T_0',X1,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2767,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f2766])).
fof(f2766,plain,(
  ( ! [X103,X101,X102,X100] : (~sP11('T_0',X101,X102) | sP12('T_0',X101,X102) | s(bool,X100) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(subsumption_resolution,[],[f2415,f769])).
fof(f2415,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,X100)) | ~sP11('T_0',X101,X102) | sP12('T_0',X101,X102) | s(bool,X100) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f423,f711])).
fof(f2691,plain,(
  ( ! [X2,X1] : (~sP12('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f2691_D])).
fof(f2691_D,plain,(
  ( ! [X2,X1] : (~sP12('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2765,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f2764])).
fof(f2764,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12('T_0',X98,X97) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f2414,f769])).
fof(f2414,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X96)) | ~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12('T_0',X98,X97) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f420,f711])).
fof(f2688,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl38),
  inference(cnf_transformation,[],[f2688_D])).
fof(f2688_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2763,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f2762])).
fof(f2762,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X93)) | sP12('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f2413,f769])).
fof(f2413,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~p(s(bool,X93)) | sP12('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f419,f711])).
fof(f2685,plain,(
  ( ! [X2,X1] : (sP12('T_0',X2,X1) | p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f2685_D])).
fof(f2685_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2761,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f2760])).
fof(f2760,plain,(
  ( ! [X90,X88,X91,X89] : (p(s(bool,X89)) | sP12('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f2412,f769])).
fof(f2412,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | p(s(bool,X89)) | sP12('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f418,f711])).
fof(f2682,plain,(
  ( ! [X2,X1] : (~sP11(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f2682_D])).
fof(f2682_D,plain,(
  ( ! [X2,X1] : (~sP11(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2759,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f2758])).
fof(f2758,plain,(
  ( ! [X87,X85,X86,X84] : (p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,'T_0') | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f2411,f769])).
fof(f2411,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,'T_0') | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f413,f711])).
fof(f2679,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f2679_D])).
fof(f2679_D,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2757,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f2756])).
fof(f2756,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X81)) | p(s(bool,X82)) | ~sP10('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f2410,f769])).
fof(f2410,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | p(s(bool,X82)) | ~sP10('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f403,f711])).
fof(f2676,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f2676_D])).
fof(f2676_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2755,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f2754])).
fof(f2754,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X77)) | sP10('T_0',X77,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f2409,f769])).
fof(f2409,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | ~p(s(bool,X77)) | sP10('T_0',X77,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f402,f711])).
fof(f2673,plain,(
  ( ! [X2,X1] : (sP10('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f2673_D])).
fof(f2673_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2753,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f2752])).
fof(f2752,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X73)) | sP10('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f2408,f769])).
fof(f2408,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | sP10('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f401,f711])).
fof(f2670,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f2670_D])).
fof(f2670_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2751,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2750])).
fof(f2750,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f2407,f769])).
fof(f2407,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f396,f711])).
fof(f2667,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f2667_D])).
fof(f2667_D,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2749,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f2748])).
fof(f2748,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,'T_0') | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f2406,f769])).
fof(f2406,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,'T_0') | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f388,f711])).
fof(f2664,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl22),
  inference(cnf_transformation,[],[f2664_D])).
fof(f2664_D,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2747,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f2746])).
fof(f2746,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,'T_0',X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f2405,f769])).
fof(f2405,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,'T_0',X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f379,f711])).
fof(f2661,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f2661_D])).
fof(f2661_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2745,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2744])).
fof(f2744,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,'T_0',X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f2404,f769])).
fof(f2404,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,'T_0',X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f360,f711])).
fof(f2658,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f2658_D])).
fof(f2658_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2743,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f2742])).
fof(f2742,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f2403,f769])).
fof(f2403,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f356,f711])).
fof(f2655,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | sP6(X1,X2,'T_0')) ) | $spl16),
  inference(cnf_transformation,[],[f2655_D])).
fof(f2655_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | sP6(X1,X2,'T_0')) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2741,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2740,f452])).
fof(f2740,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | ~sP3('T_0',X2,X1)) )),
  inference(condensation,[],[f2739])).
fof(f2739,plain,(
  ( ! [X50,X48,X51,X49] : (sP6(X49,X50,'T_0') | ~sP5(X49,'T_0',X50) | ~sP4(X50,'T_0',X49) | ~sP3('T_0',X50,X49) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f2402,f769])).
fof(f2402,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | sP6(X49,X50,'T_0') | ~sP5(X49,'T_0',X50) | ~sP4(X50,'T_0',X49) | ~sP3('T_0',X50,X49) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f353,f711])).
fof(f2651,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0') | ~p(s(bool,X1))) ) | $spl14),
  inference(cnf_transformation,[],[f2651_D])).
fof(f2651_D,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0') | ~p(s(bool,X1))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2738,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2(X1,X2,'T_0')) )),
  inference(condensation,[],[f2737])).
fof(f2737,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | sP2(X45,X46,'T_0') | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f2401,f769])).
fof(f2401,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | sP2(X45,X46,'T_0') | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f344,f711])).
fof(f2648,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl12),
  inference(cnf_transformation,[],[f2648_D])).
fof(f2648_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2736,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2735])).
fof(f2735,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f2398,f769])).
fof(f2398,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f333,f711])).
fof(f2644,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl8),
  inference(cnf_transformation,[],[f2644_D])).
fof(f2644_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2734,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2733])).
fof(f2733,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f2397,f769])).
fof(f2397,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f330,f711])).
fof(f2696,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f2695])).
fof(f2695,plain,(
  ( ! [X107,X105,X106,X104] : (~sP9('T_0',X105,X106) | sP10('T_0',X105,X106) | s(bool,X104) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(subsumption_resolution,[],[f2310,f769])).
fof(f2310,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,X104)) | ~sP9('T_0',X105,X106) | sP10('T_0',X105,X106) | s(bool,X104) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f424,f711])).
fof(f2693,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f2692])).
fof(f2692,plain,(
  ( ! [X103,X101,X102,X100] : (~sP11('T_0',X101,X102) | sP12('T_0',X101,X102) | s(bool,X100) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(subsumption_resolution,[],[f2309,f769])).
fof(f2309,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,X100)) | ~sP11('T_0',X101,X102) | sP12('T_0',X101,X102) | s(bool,X100) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f423,f711])).
fof(f2690,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f2689])).
fof(f2689,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12('T_0',X98,X97) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f2308,f769])).
fof(f2308,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X96)) | ~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12('T_0',X98,X97) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f420,f711])).
fof(f2687,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f2686])).
fof(f2686,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X93)) | sP12('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f2307,f769])).
fof(f2307,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~p(s(bool,X93)) | sP12('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f419,f711])).
fof(f2684,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f2683])).
fof(f2683,plain,(
  ( ! [X90,X88,X91,X89] : (p(s(bool,X89)) | sP12('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f2306,f769])).
fof(f2306,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | p(s(bool,X89)) | sP12('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f418,f711])).
fof(f2681,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f2680])).
fof(f2680,plain,(
  ( ! [X87,X85,X86,X84] : (p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,'T_0') | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f2305,f769])).
fof(f2305,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,'T_0') | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f413,f711])).
fof(f2678,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f2677])).
fof(f2677,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X81)) | p(s(bool,X82)) | ~sP10('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f2304,f769])).
fof(f2304,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | p(s(bool,X82)) | ~sP10('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f403,f711])).
fof(f2675,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f2674])).
fof(f2674,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X77)) | sP10('T_0',X77,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f2303,f769])).
fof(f2303,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | ~p(s(bool,X77)) | sP10('T_0',X77,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f402,f711])).
fof(f2672,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f2671])).
fof(f2671,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X73)) | sP10('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f2302,f769])).
fof(f2302,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | sP10('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f401,f711])).
fof(f2669,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2668])).
fof(f2668,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f2301,f769])).
fof(f2301,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f396,f711])).
fof(f2666,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f2665])).
fof(f2665,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,'T_0') | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f2300,f769])).
fof(f2300,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,'T_0') | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f388,f711])).
fof(f2663,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f2662])).
fof(f2662,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,'T_0',X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f2299,f769])).
fof(f2299,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,'T_0',X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f379,f711])).
fof(f2660,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2659])).
fof(f2659,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,'T_0',X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f2298,f769])).
fof(f2298,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,'T_0',X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f360,f711])).
fof(f2657,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f2656])).
fof(f2656,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f2297,f769])).
fof(f2297,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f356,f711])).
fof(f2654,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2653,f452])).
fof(f2653,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | ~sP3('T_0',X2,X1)) )),
  inference(condensation,[],[f2652])).
fof(f2652,plain,(
  ( ! [X50,X48,X51,X49] : (sP6(X49,X50,'T_0') | ~sP5(X49,'T_0',X50) | ~sP4(X50,'T_0',X49) | ~sP3('T_0',X50,X49) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f2296,f769])).
fof(f2296,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | sP6(X49,X50,'T_0') | ~sP5(X49,'T_0',X50) | ~sP4(X50,'T_0',X49) | ~sP3('T_0',X50,X49) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f353,f711])).
fof(f2650,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2(X1,X2,'T_0')) )),
  inference(condensation,[],[f2649])).
fof(f2649,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | sP2(X45,X46,'T_0') | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f2295,f769])).
fof(f2295,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | sP2(X45,X46,'T_0') | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f344,f711])).
fof(f2647,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2646])).
fof(f2646,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f2292,f769])).
fof(f2292,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f333,f711])).
fof(f2643,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2642])).
fof(f2642,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f2291,f769])).
fof(f2291,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f330,f711])).
fof(f432,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | $spl0),
  inference(cnf_transformation,[],[f432_D])).
fof(f432_D,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f438,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | $spl6),
  inference(cnf_transformation,[],[f438_D])).
fof(f438_D,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f244,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))))),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  (p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | (p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,sK18)))) & p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))))) & (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | ! [X5] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,X5)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,X5))),s(sK13,sK14))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16,sK17,sK18])],[f161])).
fof(f161,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) | ? [X6] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num',X3),s(X0,X2),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X1)))))) & (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) | ! [X5] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num',X3),s(X0,X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X1)))))))),
  inference(rectify,[],[f160])).
fof(f160,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) | ? [X5] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num',X3),s(X0,X2),s(X0,X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X1)))))) & (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) | ! [X5] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num',X3),s(X0,X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X1)))))))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) <~> ? [X5] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num',X3),s(X0,X2),s(X0,X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X1))))))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) <=> ? [X5] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num',X3),s(X0,X2),s(X0,X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X1))))))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s(X0,X2),s(X0,X1)))) <=> ? [X5] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X4),s('type.num.num',X3),s(X0,X2),s(X0,X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s(X6,X8),s(X6,X7)))) <=> ? [X11] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X10),s('type.num.num',X9),s(X6,X8),s(X6,X11)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X10),s(X6,X11))),s(X6,X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s(X6,X8),s(X6,X7)))) <=> ? [X11] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X10),s('type.num.num',X9),s(X6,X8),s(X6,X11)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X10),s(X6,X11))),s(X6,X7))))))),
  file('thm.arithmetic.NRC_SUC_RECURSE_LEFT','thm.arithmetic.NRC_SUC_RECURSE_LEFT')).
fof(f436,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | $spl4),
  inference(cnf_transformation,[],[f436_D])).
fof(f436_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f245,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,sK18))),s(sK13,sK14)))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))))),
  inference(cnf_transformation,[],[f162])).
fof(f433,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) | $spl1),
  inference(cnf_transformation,[],[f433_D])).
fof(f433_D,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f434,plain,(
  ( ! [X5] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,X5)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,X5))),s(sK13,sK14))))) ) | $spl2),
  inference(cnf_transformation,[],[f434_D])).
fof(f434_D,plain,(
  ( ! [X5] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,X5)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,X5))),s(sK13,sK14))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f246,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK17),s(sK13,X5))),s(sK13,sK14)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num',sK16),s(sK13,sK15),s(sK13,X5)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK13,fun(sK13,bool)),sK17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))),s(sK13,sK15),s(sK13,sK14))))) )),
  inference(cnf_transformation,[],[f162])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1079
% Passive clauses: 5482
% Generated clauses: 30180
% Final active clauses: 737
% Final passive clauses: 3537
% Input formulas: 43
% Initial clauses: 179
% 
% Duplicate literals: 433
% Trivial inequalities: 49
% Fw subsumption resolutions: 1418
% Fw demodulations: 6397
% Bw demodulations: 211
% Condensations: 493
% Global subsumptions: 66
% 
% Simple tautologies: 76
% Equational tautologies: 374
% Forward subsumptions: 14492
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 10415
% Factoring: 3
% Forward superposition: 6997
% Backward superposition: 3085
% Self superposition: 50
% Equality factoring: 28
% Equality resolution: 2
% 
% Split clauses: 100
% Split components: 205
% SAT solver clauses: 6550
% SAT solver unit clauses: 445
% SAT solver binary clauses: 1556
% 
% Sat splits: 100
% Sat splitting refutations: 800
% 
% Memory used [KB]: 28656
% Time elapsed: 1.934 s
% ------------------------------
% ------------------------------
% Success in time 3.264 s
