Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (20019)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1138
% Passive clauses: 8204
% Generated clauses: 31418
% Final active clauses: 955
% Final passive clauses: 6591
% Input formulas: 261
% Initial clauses: 520
% 
% Duplicate literals: 202
% Trivial inequalities: 69
% Fw subsumption resolutions: 2210
% Fw demodulations: 3442
% Bw demodulations: 172
% 
% Simple tautologies: 341
% Equational tautologies: 35
% Forward subsumptions: 16555
% 
% Binary resolution: 7540
% Factoring: 519
% Forward superposition: 8438
% Backward superposition: 7949
% Self superposition: 50
% Equality factoring: 2
% Equality resolution: 19
% 
% Split clauses: 529
% Split components: 1060
% SAT solver clauses: 708
% SAT solver unit clauses: 17
% SAT solver binary clauses: 89
% 
% Sat splits: 529
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 702
% TWLsolver calls for satisfiability: 411
% 
% Memory used [KB]: 14328
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (20040)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 772
% Passive clauses: 4877
% Generated clauses: 30808
% Final active clauses: 635
% Final passive clauses: 3387
% Input formulas: 261
% Initial clauses: 520
% 
% Duplicate literals: 98
% Trivial inequalities: 64
% Fw subsumption resolutions: 1252
% Fw demodulations: 5545
% Bw demodulations: 814
% Condensations: 70
% Global subsumptions: 25
% 
% Simple tautologies: 357
% Equational tautologies: 11
% Forward subsumptions: 12843
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 3491
% Factoring: 97
% Forward superposition: 7458
% Backward superposition: 6100
% Self superposition: 85
% Equality factoring: 2
% Equality resolution: 15
% 
% Split clauses: 32
% Split components: 65
% SAT solver clauses: 5939
% SAT solver unit clauses: 642
% SAT solver binary clauses: 166
% SAT solver learnt clauses: 50
% SAT solver learnt literals: 14
% 
% Sat splits: 32
% Sat splitting refutations: 29
% 
% TWLsolver clauses: 5839
% TWLsolver calls for satisfiability: 3527
% 
% Memory used [KB]: 27888
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (20078)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1367
% Passive clauses: 12272
% Generated clauses: 146679
% Final active clauses: 1120
% Final passive clauses: 5288
% Input formulas: 261
% Initial clauses: 520
% 
% Duplicate literals: 36
% Trivial inequalities: 142
% Fw subsumption resolutions: 2340
% Fw demodulations: 54689
% Bw demodulations: 3511
% Condensations: 104
% Global subsumptions: 4
% 
% Simple tautologies: 250
% Equational tautologies: 64
% Forward subsumptions: 74524
% Fw demodulations to eq. taut.: 190
% Bw demodulations to eq. taut.: 32
% 
% Unit resulting resolution: 54783
% Factoring: 8
% Forward superposition: 16581
% Backward superposition: 10840
% Self superposition: 134
% Equality factoring: 2
% Equality resolution: 23
% 
% Split clauses: 37
% Split components: 75
% SAT solver clauses: 14379
% SAT solver unit clauses: 3656
% SAT solver binary clauses: 6478
% 
% Sat splits: 37
% Sat splitting refutations: 2027
% 
% Memory used [KB]: 244473
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f69634,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f574,f881,f879,f877,f573,f572,f890,f1128,f1186,f1314,f1188,f1189,f1332,f1333,f1492,f1439,f1319,f1442,f1443,f1508,f1513,f1514,f1453,f1454,f1455,f1456,f1526,f1530,f1504,f2163,f2503,f2657,f2793,f3329,f3468,f3475,f3482,f3487,f3890,f3906,f4042,f4544,f3908,f4405,f4402,f4547,f4697,f5058,f5080,f5190,f5209,f2642,f2502,f5212,f5214,f5495,f5789,f5805,f5806,f5807,f5808,f5819,f5763,f5762,f5761,f5760,f5820,f5825,f5834,f6219,f6489,f6845,f6847,f6851,f6857,f6868,f6869,f6876,f6877,f6888,f7037,f7772,f8371,f8372,f8401,f8402,f8607,f8736,f8867,f8875,f8880,f8881,f9034,f9467,f9482,f9488,f9598,f9836,f9851,f9855,f9973,f9975,f9981,f4038,f9982,f9986,f9987,f9988,f10136,f10268,f10139,f1321,f10643,f10116,f10651,f9033,f10524,f1315,f10787,f10788,f10789,f10790,f10118,f10799,f10800,f10852,f10854,f10860,f10886,f10888,f10932,f10939,f8735,f11136,f11140,f11276,f11277,f11283,f878,f1502,f11284,f1010,f11422,f11423,f11430,f11431,f1313,f11635,f1483,f12066,f12072,f12076,f12037,f12038,f12039,f12040,f12078,f12538,f12745,f13503,f13504,f13534,f13535,f13706,f13967,f14149,f14286,f14288,f14423,f14434,f17160,f18747,f19049,f19215,f1510,f19219,f19184,f19185,f19186,f19187,f19221,f19392,f19394,f19396,f18545,f19416,f18529,f19417,f18515,f19418,f18501,f19419,f18488,f19420,f18475,f19421,f18462,f19422,f18449,f19423,f18436,f19424,f18423,f19425,f18394,f18382,f15724,f19433,f15704,f19434,f15684,f19435,f15665,f19436,f15647,f19437,f15628,f19438,f15608,f19439,f15590,f19440,f15555,f19441,f15537,f19426,f15500,f19427,f15325,f15307,f19432,f14146,f19448,f13553,f19443,f13542,f19444,f19223,f19218,f19174,f19706,f19707,f19748,f19749,f20147,f20492,f20497,f20499,f20502,f20495,f24040,f24041,f24043,f23656,f24045,f24050,f24051,f24056,f23662,f24060,f24064,f23665,f24065,f24068,f23669,f24069,f24075,f24078,f24077,f24138,f24144,f24143,f24333,f24334,f24363,f24364,f24734,f24737,f1500,f24739,f25038,f25039,f24058,f24745,f25041,f25122,f25123,f25153,f25154,f26906,f26907,f26910,f26911,f26912,f26913,f26914,f26915,f26916,f27501,f1509,f28690,f28704,f28709,f28714,f28715,f28718,f28725,f28727,f28705,f24066,f28731,f28733,f28710,f24071,f28773,f28774,f28777,f28720,f28836,f31579,f32040,f32045,f32729,f32730,f32731,f32732,f32733,f32727,f32051,f32737,f32740,f32743,f32741,f32745,f32748,f32746,f32752,f32757,f32756,f32821,f32822,f34546,f10797,f34674,f35006,f35788,f35789,f35790,f35791,f35792,f35786,f35012,f35794,f35799,f35802,f35800,f35804,f35807,f35805,f35811,f35816,f35813,f35885,f35886,f35971,f36478,f24046,f37514,f37515,f37516,f37517,f37518,f24057,f37512,f36484,f35796,f32735,f37522,f37525,f24052,f37523,f37527,f37530,f24070,f37528,f37534,f24061,f37535,f24079,f35815,f32754,f37620,f26909,f37621,f39553,f39561,f39562,f39563,f39564,f39565,f39566,f39580,f39581,f39575,f39668,f39669,f39670,f39681,f39895,f39896,f39897,f39908,f40011,f20493,f40019,f39576,f40020,f20498,f40021,f39577,f40022,f39578,f40023,f39579,f40024,f20500,f40034,f28837,f40035,f20503,f40033,f40359,f15291,f40377,f13045,f4043,f11759,f11281,f14964,f1316,f5340,f40213,f40380,f40750,f40751,f40890,f9004,f40891,f41381,f41385,f41383,f41538,f41541,f41553,f41694,f41695,f41540,f41847,f41864,f41865,f41873,f41874,f19447,f44644,f45149,f45209,f46583,f46584,f46585,f46586,f46587,f46588,f46589,f41544,f40362,f40212,f1501,f10272,f13834,f1323,f12674,f5211,f12537,f880,f46640,f1013,f46772,f46773,f46782,f46783,f5339,f46785,f4682,f46920,f1164,f46891,f1163,f1162,f47369,f1482,f5791,f47375,f47340,f47341,f47342,f47343,f47377,f47487,f1318,f47488,f47631,f2638,f47830,f4039,f4037,f47961,f4044,f48021,f2288,f1322,f48420,f2795,f1324,f49167,f49170,f49728,f49731,f49874,f49873,f49870,f49865,f49864,f49863,f49862,f49861,f49860,f49859,f49854,f49853,f49838,f49961,f49962,f49963,f49964,f49965,f49966,f49967,f49968,f49969,f49970,f49987,f49988,f50027,f50028,f50067,f50068,f50096,f50462,f50464,f50825,f50827,f50956,f51108,f51612,f51613,f51614,f50475,f2640,f49729,f51849,f51858,f51986,f45360,f45208,f9003,f52003,f52005,f51987,f9990,f52029,f52032,f52031,f52159,f52160,f52161,f52162,f52163,f52164,f52176,f52215,f52226,f52246,f52247,f52613,f52614,f52767,f52769,f52995,f14417,f52997,f53129,f10649,f53680,f53686,f40333,f40332,f40331,f40330,f53687,f3891,f36476,f36475,f36474,f36473,f35004,f35003,f35002,f35001,f32043,f32042,f32041,f53691,f53693,f53695,f53700,f18603,f18602,f18601,f18600,f53701,f24042,f53708,f1528,f18572,f18571,f18570,f1522,f18569,f53709,f53689,f53912,f54065,f54068,f1503,f54070,f19431,f54071,f54073,f54075,f54076,f5193,f54078,f54079,f19449,f54081,f14418,f1497,f52002,f9984,f52000,f5191,f9983,f54083,f51989,f54090,f54092,f50094,f19450,f54101,f18867,f26137,f40173,f54104,f5060,f54105,f54091,f54106,f54093,f26142,f26143,f47327,f26135,f5746,f12023,f1320,f18555,f54113,f18586,f54114,f40317,f1011,f26139,f19170,f50093,f9009,f50226,f50196,f10809,f1498,f54321,f9010,f55104,f55105,f55106,f55107,f55108,f55109,f55117,f52370,f10650,f55118,f36094,f55409,f3894,f3892,f55875,f55877,f3474,f56193,f56194,f56233,f56234,f9484,f56273,f56274,f56313,f56314,f5192,f56411,f56412,f3473,f56769,f14419,f49298,f56772,f3454,f56912,f9007,f14277,f57047,f57196,f57197,f57198,f57201,f41539,f57658,f57667,f57668,f57669,f57670,f57671,f57672,f57673,f57674,f57678,f9006,f10808,f57687,f41676,f58169,f57180,f58303,f58322,f50225,f41675,f58450,f9008,f41849,f59006,f60047,f60174,f60175,f60214,f60215,f5194,f60540,f60541,f60542,f60543,f62011,f62135,f10798,f62427,f62428,f62429,f62600,f47832,f19069,f13837,f8738,f876,f1481,f62612,f62613,f62614,f62615,f62616,f62617,f62618,f62619,f62620,f62621,f62622,f62623,f62624,f62625,f62626,f62627,f62628,f62629,f62630,f62631,f62632,f62633,f62634,f62635,f62636,f62637,f62638,f62639,f62640,f62641,f62642,f62643,f62644,f62857,f62858,f62876,f62877,f4681,f62881,f1487,f1485,f63042,f63043,f63062,f63063,f63075,f63076,f1317,f63130,f63131,f63300,f63301,f49730,f63315,f63316,f63328,f63329,f49297,f6846,f63333,f14422,f2287,f63463,f5210,f64098,f1480,f3455,f64488,f64494,f64649,f64490,f64651,f64492,f1524,f66487,f66488,f66489,f66511,f66649,f66647,f66970,f66653,f46607,f10271,f67156,f67446,f67598,f67588,f67599,f67663,f67654,f67664,f67742,f67731,f67743,f68207,f68197,f68208,f68323,f68313,f68324,f68395,f68384,f68396,f68462,f68450,f68463,f68581,f68569,f68582,f69027,f1516,f69015,f69028,f69260,f1520,f1518,f69295,f45357,f45355,f69606])).
fof(f69606,plain,(
  $false | ($spl2 | $spl4 | $spl27)),
  inference(subsumption_resolution,[],[f69605,f1156])).
fof(f1156,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl2),
  inference(forward_demodulation,[],[f1027,f844])).
fof(f844,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f377])).
fof(f377,plain,(
  ! [X0] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ABS_NEG')).
fof(f1027,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | $spl2),
  inference(resolution,[],[f878,f804])).
fof(f804,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f556])).
fof(f556,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))),
  inference(flattening,[],[f555])).
fof(f555,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))),
  inference(nnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))))))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ABS_LT|split|1')).
fof(f69605,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl4 | $spl27)),
  inference(subsumption_resolution,[],[f69456,f1481])).
fof(f69456,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl4),
  inference(resolution,[],[f46638,f808])).
fof(f808,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f558])).
fof(f558,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))))),
  inference(flattening,[],[f557])).
fof(f557,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f347])).
fof(f347,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ABS_LT|split|0')).
fof(f46638,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f880,f802])).
fof(f802,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f554])).
fof(f554,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1))))))),
  inference(flattening,[],[f553])).
fof(f553,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))))),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7))))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ABS_LT|split|2')).
fof(f45355,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | $spl331),
  inference(cnf_transformation,[],[f45355_D])).
fof(f45355_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) <=> ~$spl331),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl331])])).
fof(f45357,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | $spl332),
  inference(cnf_transformation,[],[f45357_D])).
fof(f45357_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f69295,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | $spl105),
  inference(resolution,[],[f9008,f809])).
fof(f809,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f560])).
fof(f560,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f559])).
fof(f559,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f176])).
fof(f176,axiom,(
  ! [X9,X11] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) | s('type.integer.int',X9) = s('type.integer.int',X11)))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_LE_LT')).
fof(f1518,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl46),
  inference(cnf_transformation,[],[f1518_D])).
fof(f1518_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1520,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | $spl48),
  inference(cnf_transformation,[],[f1520_D])).
fof(f1520_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f69260,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl31 | $spl40)),
  inference(subsumption_resolution,[],[f69223,f10329])).
fof(f10329,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl40),
  inference(resolution,[],[f1503,f810])).
fof(f810,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f560])).
fof(f69223,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl31),
  inference(superposition,[],[f1485,f835])).
fof(f835,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X0)))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26(X0))))) )),
  inference(cnf_transformation,[],[f564])).
fof(f564,plain,(
  ! [X0] : ((s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26(X0)))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26(X0))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X0)))))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27(X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26,sK27])],[f371])).
fof(f371,plain,(
  ! [X0] : (? [X2] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) | ? [X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f370])).
fof(f370,plain,(
  ! [X0] : (? [X2] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) | ? [X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : (? [X10] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) | ? [X10] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) | s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_NUM_CASES')).
fof(f69028,plain,(
  ( ! [X24,X25] : (sP7('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) | $spl392),
  inference(cnf_transformation,[],[f69028_D])).
fof(f69028_D,plain,(
  ( ! [X24,X25] : (sP7('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f69015,plain,(
  ( ! [X26,X27] : (sP7('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X26,X27) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl1),
  inference(superposition,[],[f1392,f835])).
fof(f1392,plain,(
  ( ! [X107,X106] : (sP7('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))),X106,X107)) ) | $spl1),
  inference(resolution,[],[f877,f640])).
fof(f640,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f479])).
fof(f479,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f478])).
fof(f478,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f477])).
fof(f477,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f435])).
fof(f435,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f1516,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | $spl45),
  inference(cnf_transformation,[],[f1516_D])).
fof(f1516_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f69027,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | sP7('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) | $spl1),
  inference(forward_demodulation,[],[f69014,f845])).
fof(f845,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f378])).
fof(f378,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9] : s('type.integer.int',X9) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_NEGNEG')).
fof(f69014,plain,(
  ( ! [X24,X25] : (sP7('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl1),
  inference(superposition,[],[f1392,f833])).
fof(f833,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f563])).
fof(f563,plain,(
  ! [X0] : ((s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f208])).
fof(f208,axiom,(
  ! [X6] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_NEG_SAME_EQ')).
fof(f68582,plain,(
  ( ! [X24,X25] : (sP7(X24,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25)) ) | $spl390),
  inference(cnf_transformation,[],[f68582_D])).
fof(f68582_D,plain,(
  ( ! [X24,X25] : (sP7(X24,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25)) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f68569,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X27) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl1),
  inference(superposition,[],[f1391,f835])).
fof(f1391,plain,(
  ( ! [X105,X104] : (sP7(X104,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))),X105)) ) | $spl1),
  inference(resolution,[],[f877,f639])).
fof(f639,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f479])).
fof(f68581,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | sP7(X24,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25)) ) | $spl1),
  inference(forward_demodulation,[],[f68568,f845])).
fof(f68568,plain,(
  ( ! [X24,X25] : (sP7(X24,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl1),
  inference(superposition,[],[f1391,f833])).
fof(f68463,plain,(
  ( ! [X24,X25] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) | $spl388),
  inference(cnf_transformation,[],[f68463_D])).
fof(f68463_D,plain,(
  ( ! [X24,X25] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f68450,plain,(
  ( ! [X26,X27] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X26,X27) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl1),
  inference(superposition,[],[f1389,f835])).
fof(f1389,plain,(
  ( ! [X101,X100] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))),X100,X101)) ) | $spl1),
  inference(resolution,[],[f877,f636])).
fof(f636,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f476])).
fof(f476,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f475])).
fof(f475,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f474])).
fof(f474,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f436])).
fof(f436,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f68462,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | sP8('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) | $spl1),
  inference(forward_demodulation,[],[f68449,f845])).
fof(f68449,plain,(
  ( ! [X24,X25] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl1),
  inference(superposition,[],[f1389,f833])).
fof(f68396,plain,(
  ( ! [X24,X25] : (sP8(X24,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25)) ) | $spl386),
  inference(cnf_transformation,[],[f68396_D])).
fof(f68396_D,plain,(
  ( ! [X24,X25] : (sP8(X24,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25)) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f68384,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X27) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl1),
  inference(superposition,[],[f1388,f835])).
fof(f1388,plain,(
  ( ! [X99,X98] : (sP8(X98,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))),X99)) ) | $spl1),
  inference(resolution,[],[f877,f635])).
fof(f635,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f476])).
fof(f68395,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | sP8(X24,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25)) ) | $spl1),
  inference(forward_demodulation,[],[f68383,f845])).
fof(f68383,plain,(
  ( ! [X24,X25] : (sP8(X24,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl1),
  inference(superposition,[],[f1388,f833])).
fof(f68324,plain,(
  ( ! [X24,X25] : (sP4('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) | $spl384),
  inference(cnf_transformation,[],[f68324_D])).
fof(f68324_D,plain,(
  ( ! [X24,X25] : (sP4('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f68313,plain,(
  ( ! [X26,X27] : (sP4('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X26,X27) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl1),
  inference(superposition,[],[f1376,f835])).
fof(f1376,plain,(
  ( ! [X74,X75] : (sP4('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))),X74,X75)) ) | $spl1),
  inference(resolution,[],[f877,f619])).
fof(f619,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f468])).
fof(f468,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f467])).
fof(f467,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f431])).
fof(f431,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f68323,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | sP4('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) | $spl1),
  inference(forward_demodulation,[],[f68312,f845])).
fof(f68312,plain,(
  ( ! [X24,X25] : (sP4('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl1),
  inference(superposition,[],[f1376,f833])).
fof(f68208,plain,(
  ( ! [X24,X25] : (sP4(X24,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25)) ) | $spl382),
  inference(cnf_transformation,[],[f68208_D])).
fof(f68208_D,plain,(
  ( ! [X24,X25] : (sP4(X24,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25)) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f68197,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X27) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl1),
  inference(superposition,[],[f1375,f835])).
fof(f1375,plain,(
  ( ! [X72,X73] : (sP4(X72,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))),X73)) ) | $spl1),
  inference(resolution,[],[f877,f618])).
fof(f618,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f468])).
fof(f68207,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | sP4(X24,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25)) ) | $spl1),
  inference(forward_demodulation,[],[f68196,f845])).
fof(f68196,plain,(
  ( ! [X24,X25] : (sP4(X24,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl1),
  inference(superposition,[],[f1375,f833])).
fof(f67743,plain,(
  ( ! [X24,X25] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) | $spl380),
  inference(cnf_transformation,[],[f67743_D])).
fof(f67743_D,plain,(
  ( ! [X24,X25] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f67731,plain,(
  ( ! [X26,X27] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X26,X27) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl1),
  inference(superposition,[],[f1362,f835])).
fof(f1362,plain,(
  ( ! [X47,X46] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))),X46,X47)) ) | $spl1),
  inference(resolution,[],[f877,f602])).
fof(f602,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f460])).
fof(f460,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f459])).
fof(f459,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f428])).
fof(f428,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f67742,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | sP2('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) | $spl1),
  inference(forward_demodulation,[],[f67730,f845])).
fof(f67730,plain,(
  ( ! [X24,X25] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl1),
  inference(superposition,[],[f1362,f833])).
fof(f67664,plain,(
  ( ! [X24,X25] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) | $spl378),
  inference(cnf_transformation,[],[f67664_D])).
fof(f67664_D,plain,(
  ( ! [X24,X25] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f67654,plain,(
  ( ! [X26,X27] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X26,X27) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl1),
  inference(superposition,[],[f1346,f835])).
fof(f1346,plain,(
  ( ! [X14,X15] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))),X14,X15)) ) | $spl1),
  inference(resolution,[],[f877,f585])).
fof(f585,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f453])).
fof(f453,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f452])).
fof(f452,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f451])).
fof(f451,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f425])).
fof(f425,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f67663,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | sP0('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25)) ) | $spl1),
  inference(forward_demodulation,[],[f67653,f845])).
fof(f67653,plain,(
  ( ! [X24,X25] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24,X25) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl1),
  inference(superposition,[],[f1346,f833])).
fof(f67599,plain,(
  ( ! [X24,X25] : (sP0(X24,X25,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl376),
  inference(cnf_transformation,[],[f67599_D])).
fof(f67599_D,plain,(
  ( ! [X24,X25] : (sP0(X24,X25,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f67588,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl1),
  inference(superposition,[],[f1345,f835])).
fof(f1345,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl1),
  inference(resolution,[],[f877,f583])).
fof(f583,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f453])).
fof(f67598,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | sP0(X24,X25,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl1),
  inference(forward_demodulation,[],[f67587,f845])).
fof(f67587,plain,(
  ( ! [X24,X25] : (sP0(X24,X25,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl1),
  inference(superposition,[],[f1345,f833])).
fof(f67446,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f67409,f10324])).
fof(f10324,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl40),
  inference(resolution,[],[f1503,f797])).
fof(f797,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f551])).
fof(f551,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f341])).
fof(f341,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f154])).
fof(f154,axiom,(
  ! [X9,X11] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_NOT_LT')).
fof(f67409,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(superposition,[],[f1337,f835])).
fof(f1337,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl1),
  inference(resolution,[],[f877,f798])).
fof(f798,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f552])).
fof(f552,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f343])).
fof(f343,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f230])).
fof(f230,axiom,(
  ! [X9,X11] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_NOT_LE')).
fof(f67156,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f67114,f1503])).
fof(f67114,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(superposition,[],[f877,f835])).
fof(f10271,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl116),
  inference(cnf_transformation,[],[f10271_D])).
fof(f10271_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f46607,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl335),
  inference(cnf_transformation,[],[f46607_D])).
fof(f46607_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl335),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl335])])).
fof(f66653,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl27),
  inference(forward_demodulation,[],[f66514,f856])).
fof(f856,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f385])).
fof(f385,plain,(
  ! [X0] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ABS_ABS')).
fof(f66514,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl27),
  inference(resolution,[],[f1481,f800])).
fof(f800,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f554])).
fof(f66970,plain,(
  $false | ($spl4 | $spl32 | $spl105)),
  inference(subsumption_resolution,[],[f66969,f880])).
fof(f66969,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | ($spl32 | $spl105)),
  inference(forward_demodulation,[],[f9008,f1487])).
fof(f66647,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl373),
  inference(cnf_transformation,[],[f66647_D])).
fof(f66647_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl373),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl373])])).
fof(f66649,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl374),
  inference(cnf_transformation,[],[f66649_D])).
fof(f66649_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f66511,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl27),
  inference(resolution,[],[f1481,f809])).
fof(f66489,plain,(
  $false | ($spl181 | $spl323)),
  inference(subsumption_resolution,[],[f66364,f41385])).
fof(f66364,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl181),
  inference(resolution,[],[f19448,f796])).
fof(f796,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f551])).
fof(f66488,plain,(
  $false | ($spl181 | $spl323)),
  inference(subsumption_resolution,[],[f66363,f41385])).
fof(f66363,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl181),
  inference(resolution,[],[f19448,f798])).
fof(f66487,plain,(
  $false | ($spl100 | $spl181)),
  inference(subsumption_resolution,[],[f66361,f9003])).
fof(f66361,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl181),
  inference(resolution,[],[f19448,f810])).
fof(f1524,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl53),
  inference(cnf_transformation,[],[f1524_D])).
fof(f1524_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f64492,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl370),
  inference(cnf_transformation,[],[f64492_D])).
fof(f64492_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f64651,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl11 | $spl32)),
  inference(forward_demodulation,[],[f64650,f1487])).
fof(f64650,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl11 | $spl32)),
  inference(subsumption_resolution,[],[f64428,f1162])).
fof(f64428,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl32),
  inference(superposition,[],[f837,f1487])).
fof(f837,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27(X0)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26(X0))))) )),
  inference(cnf_transformation,[],[f564])).
fof(f64490,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | $spl368),
  inference(cnf_transformation,[],[f64490_D])).
fof(f64490_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f64649,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | ($spl11 | $spl32)),
  inference(forward_demodulation,[],[f64648,f1487])).
fof(f64648,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl11 | $spl32)),
  inference(forward_demodulation,[],[f64647,f1487])).
fof(f64647,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl11 | $spl32)),
  inference(subsumption_resolution,[],[f64426,f1162])).
fof(f64426,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl32),
  inference(superposition,[],[f835,f1487])).
fof(f64494,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl11 | $spl32)),
  inference(forward_demodulation,[],[f64493,f1487])).
fof(f64493,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl11 | $spl32)),
  inference(subsumption_resolution,[],[f64245,f1162])).
fof(f64245,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl32),
  inference(superposition,[],[f1487,f837])).
fof(f64488,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | ($spl11 | $spl32)),
  inference(forward_demodulation,[],[f64487,f1487])).
fof(f64487,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl11 | $spl32)),
  inference(forward_demodulation,[],[f64486,f1487])).
fof(f64486,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl11 | $spl32)),
  inference(subsumption_resolution,[],[f64243,f1162])).
fof(f64243,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl32),
  inference(superposition,[],[f1487,f835])).
fof(f3455,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl67),
  inference(cnf_transformation,[],[f3455_D])).
fof(f3455_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f1480,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl26),
  inference(cnf_transformation,[],[f1480_D])).
fof(f1480_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f64098,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f808])).
fof(f5210,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl92),
  inference(cnf_transformation,[],[f5210_D])).
fof(f5210_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f63463,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl9),
  inference(forward_demodulation,[],[f63334,f856])).
fof(f63334,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int',sK13)))) | $spl9),
  inference(resolution,[],[f1012,f800])).
fof(f1012,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl9),
  inference(cnf_transformation,[],[f1012_D])).
fof(f1012_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f2287,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl57),
  inference(cnf_transformation,[],[f2287_D])).
fof(f2287_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f14422,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl130),
  inference(cnf_transformation,[],[f14422_D])).
fof(f14422_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f63333,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl9),
  inference(resolution,[],[f1012,f803])).
fof(f803,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))) )),
  inference(cnf_transformation,[],[f556])).
fof(f6846,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl98),
  inference(cnf_transformation,[],[f6846_D])).
fof(f6846_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f49297,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl336),
  inference(cnf_transformation,[],[f49297_D])).
fof(f49297_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f63329,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X1))) ) | $spl2),
  inference(subsumption_resolution,[],[f63318,f758])).
fof(f758,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f540])).
fof(f540,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f539])).
fof(f539,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X15] : ((p(s(bool,'F_0')) | p(s(bool,X15))) <=> p(s(bool,X15)))),
  file('thm.integer.INT_ABS_LT|split|3','thm.bool.OR_CLAUSES|split|2')).
fof(f63318,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X1))) ) | $spl2),
  inference(resolution,[],[f16740,f704])).
fof(f704,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f508])).
fof(f508,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f507])).
fof(f507,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f443])).
fof(f443,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f16740,plain,(
  ( ! [X145] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))),X145)) ) | $spl2),
  inference(resolution,[],[f3330,f706])).
fof(f706,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f508])).
fof(f3330,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl2),
  inference(resolution,[],[f1156,f810])).
fof(f63328,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f63317,f758])).
fof(f63317,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X0))) ) | $spl2),
  inference(resolution,[],[f16740,f708])).
fof(f708,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f510])).
fof(f510,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f509])).
fof(f509,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f444])).
fof(f444,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f408,f443])).
fof(f408,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f407])).
fof(f407,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X21,X8] : ((~(p(s(bool,X8)) | p(s(bool,X21))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X8)) => p(s(bool,'F_0'))) => (~p(s(bool,X21)) => p(s(bool,'F_0')))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.sat.sat_9_prove_ltb')).
fof(f63316,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f63305,f758])).
fof(f63305,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(resolution,[],[f16739,f704])).
fof(f16739,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl2),
  inference(resolution,[],[f3330,f705])).
fof(f705,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f508])).
fof(f63315,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f63304,f758])).
fof(f63304,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(resolution,[],[f16739,f708])).
fof(f49730,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl338),
  inference(cnf_transformation,[],[f49730_D])).
fof(f49730_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f63301,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | $spl7),
  inference(subsumption_resolution,[],[f63286,f758])).
fof(f63286,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | $spl7),
  inference(resolution,[],[f15438,f704])).
fof(f15438,plain,(
  ( ! [X145] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X145)) ) | $spl7),
  inference(resolution,[],[f2507,f706])).
fof(f2507,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl7),
  inference(resolution,[],[f1011,f798])).
fof(f63300,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | $spl7),
  inference(subsumption_resolution,[],[f63285,f758])).
fof(f63285,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | $spl7),
  inference(resolution,[],[f15438,f708])).
fof(f63131,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f63116,f758])).
fof(f63116,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))) ) | $spl7),
  inference(resolution,[],[f15437,f704])).
fof(f15437,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) ) | $spl7),
  inference(resolution,[],[f2507,f705])).
fof(f63130,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f63115,f758])).
fof(f63115,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))) ) | $spl7),
  inference(resolution,[],[f15437,f708])).
fof(f1317,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl18),
  inference(cnf_transformation,[],[f1317_D])).
fof(f1317_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f63076,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X1))) ) | $spl2),
  inference(subsumption_resolution,[],[f63065,f758])).
fof(f63065,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X1))) ) | $spl2),
  inference(resolution,[],[f3408,f704])).
fof(f3408,plain,(
  ( ! [X145] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))),X145)) ) | $spl2),
  inference(resolution,[],[f1156,f706])).
fof(f63075,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f63064,f758])).
fof(f63064,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X0))) ) | $spl2),
  inference(resolution,[],[f3408,f708])).
fof(f63063,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f63052,f758])).
fof(f63052,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(resolution,[],[f3407,f704])).
fof(f3407,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl2),
  inference(resolution,[],[f1156,f705])).
fof(f63062,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f63051,f758])).
fof(f63051,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(resolution,[],[f3407,f708])).
fof(f63043,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f63002,f1503])).
fof(f63002,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(superposition,[],[f877,f837])).
fof(f63042,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f63000,f1503])).
fof(f63000,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(superposition,[],[f877,f835])).
fof(f1485,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl31),
  inference(cnf_transformation,[],[f1485_D])).
fof(f1485_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f1487,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | $spl32),
  inference(cnf_transformation,[],[f1487_D])).
fof(f1487_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f62881,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(resolution,[],[f877,f809])).
fof(f4681,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl84),
  inference(cnf_transformation,[],[f4681_D])).
fof(f4681_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f62877,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,X1))) ) | $spl2),
  inference(subsumption_resolution,[],[f62862,f758])).
fof(f62862,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,X1))) ) | $spl2),
  inference(resolution,[],[f3039,f704])).
fof(f3039,plain,(
  ( ! [X145] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))),X145)) ) | $spl2),
  inference(resolution,[],[f1029,f706])).
fof(f1029,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl2),
  inference(resolution,[],[f878,f810])).
fof(f62876,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f62861,f758])).
fof(f62861,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,X0))) ) | $spl2),
  inference(resolution,[],[f3039,f708])).
fof(f62858,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f62843,f758])).
fof(f62843,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(resolution,[],[f3038,f704])).
fof(f3038,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl2),
  inference(resolution,[],[f1029,f705])).
fof(f62857,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f62842,f758])).
fof(f62842,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(resolution,[],[f3038,f708])).
fof(f62644,plain,(
  $false | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f1020,f1324])).
fof(f1020,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f878,f679])).
fof(f679,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f402])).
fof(f402,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f246])).
fof(f246,axiom,(
  ! [X9,X11] : ~(p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_LT_ANTISYM')).
fof(f62643,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3321,f1013])).
fof(f3321,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f679])).
fof(f62642,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3322,f1013])).
fof(f3322,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f679])).
fof(f62641,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3323,f1013])).
fof(f3323,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f734])).
fof(f734,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f412])).
fof(f412,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f194])).
fof(f194,axiom,(
  ! [X9,X11] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_LT_GT')).
fof(f62640,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2796,f1013])).
fof(f2796,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(forward_demodulation,[],[f2658,f844])).
fof(f2658,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1020,f806])).
fof(f806,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f558])).
fof(f62639,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3324,f1013])).
fof(f3324,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f734])).
fof(f62638,plain,(
  $false | ($spl2 | $spl126)),
  inference(subsumption_resolution,[],[f3325,f14416])).
fof(f14416,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl126),
  inference(cnf_transformation,[],[f14416_D])).
fof(f14416_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f3325,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f797])).
fof(f62637,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f6807,f1013])).
fof(f6807,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(forward_demodulation,[],[f6673,f844])).
fof(f6673,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1020,f806])).
fof(f62636,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f16661,f1013])).
fof(f16661,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f3330,f799])).
fof(f799,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f552])).
fof(f62635,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f16662,f1013])).
fof(f16662,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f3330,f797])).
fof(f62634,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f13698,f1013])).
fof(f13698,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f679])).
fof(f62633,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f13699,f1013])).
fof(f13699,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f679])).
fof(f62632,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f13700,f1013])).
fof(f13700,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f734])).
fof(f62631,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f13701,f1013])).
fof(f13701,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f734])).
fof(f62630,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8038,f1013])).
fof(f8038,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(forward_demodulation,[],[f7904,f844])).
fof(f7904,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1020,f806])).
fof(f62629,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8599,f1013])).
fof(f8599,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f679])).
fof(f62628,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8600,f1013])).
fof(f8600,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f679])).
fof(f62627,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8601,f1013])).
fof(f8601,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f734])).
fof(f62626,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8602,f1013])).
fof(f8602,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f734])).
fof(f62625,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f16494,f1013])).
fof(f16494,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f3325,f810])).
fof(f62624,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f13044,f1013])).
fof(f13044,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(forward_demodulation,[],[f12910,f844])).
fof(f12910,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1020,f806])).
fof(f62623,plain,(
  $false | ($spl6 | $spl20)),
  inference(subsumption_resolution,[],[f15154,f1010])).
fof(f15154,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f679])).
fof(f62622,plain,(
  $false | ($spl6 | $spl20)),
  inference(subsumption_resolution,[],[f15155,f1010])).
fof(f15155,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f679])).
fof(f62621,plain,(
  $false | ($spl6 | $spl20)),
  inference(subsumption_resolution,[],[f15156,f1010])).
fof(f15156,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f734])).
fof(f62620,plain,(
  $false | ($spl6 | $spl20)),
  inference(subsumption_resolution,[],[f15157,f1010])).
fof(f15157,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f734])).
fof(f62619,plain,(
  $false | ($spl20 | $spl58)),
  inference(subsumption_resolution,[],[f15158,f2637])).
fof(f2637,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl58),
  inference(cnf_transformation,[],[f2637_D])).
fof(f2637_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f15158,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f797])).
fof(f62618,plain,(
  $false | ($spl20 | $spl58)),
  inference(subsumption_resolution,[],[f15159,f2637])).
fof(f15159,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f799])).
fof(f62617,plain,(
  $false | ($spl6 | $spl20)),
  inference(subsumption_resolution,[],[f50328,f1010])).
fof(f50328,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f679])).
fof(f62616,plain,(
  $false | ($spl6 | $spl20)),
  inference(subsumption_resolution,[],[f50329,f1010])).
fof(f50329,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f734])).
fof(f62615,plain,(
  $false | ($spl6 | $spl20)),
  inference(subsumption_resolution,[],[f50330,f1010])).
fof(f50330,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f734])).
fof(f62614,plain,(
  $false | ($spl20 | $spl58)),
  inference(subsumption_resolution,[],[f50331,f2637])).
fof(f50331,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f797])).
fof(f62613,plain,(
  $false | ($spl20 | $spl58)),
  inference(subsumption_resolution,[],[f50332,f2637])).
fof(f50332,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f799])).
fof(f62612,plain,(
  $false | ($spl6 | $spl20)),
  inference(subsumption_resolution,[],[f50327,f1010])).
fof(f50327,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f679])).
fof(f1481,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl27),
  inference(cnf_transformation,[],[f1481_D])).
fof(f1481_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f876,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(cnf_transformation,[],[f876_D])).
fof(f876_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f8738,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl29),
  inference(resolution,[],[f1482,f803])).
fof(f13837,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl29),
  inference(resolution,[],[f1482,f803])).
fof(f19069,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl30),
  inference(resolution,[],[f1484,f809])).
fof(f1484,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl30),
  inference(cnf_transformation,[],[f1484_D])).
fof(f1484_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f47832,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl29),
  inference(resolution,[],[f1482,f803])).
fof(f62600,plain,(
  $false | ($spl3 | $spl4 | $spl26)),
  inference(subsumption_resolution,[],[f18758,f4686])).
fof(f4686,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f4685,f844])).
fof(f4685,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4684,f783])).
fof(f783,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f420])).
fof(f420,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f182])).
fof(f182,axiom,(
  ! [X15] : ((p(s(bool,X15)) => p(s(bool,'F_0'))) => ~p(s(bool,X15)))),
  file('thm.integer.INT_ABS_LT|split|3','thm.bool.IMP_F')).
fof(f4684,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f4683,f1279])).
fof(f1279,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))) | $spl4),
  inference(resolution,[],[f882,f752])).
fof(f752,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f536])).
fof(f536,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f535])).
fof(f535,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f415])).
fof(f415,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f310])).
fof(f310,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X15] : ((p(s(bool,X15)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X15))),
  file('thm.integer.INT_ABS_LT|split|3','thm.bool.IMP_F_EQ_F')).
fof(f882,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f880,f679])).
fof(f4683,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | $spl3),
  inference(forward_demodulation,[],[f4549,f817])).
fof(f817,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f354])).
fof(f354,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9,X11] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_LT_NEG')).
fof(f4549,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | $spl3),
  inference(resolution,[],[f879,f803])).
fof(f18758,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl26),
  inference(resolution,[],[f1480,f806])).
fof(f62429,plain,(
  $false | ($spl3 | $spl4 | $spl127)),
  inference(subsumption_resolution,[],[f62302,f4686])).
fof(f62302,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl127),
  inference(resolution,[],[f14417,f796])).
fof(f62428,plain,(
  $false | ($spl3 | $spl4 | $spl127)),
  inference(subsumption_resolution,[],[f62301,f4686])).
fof(f62301,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl127),
  inference(resolution,[],[f14417,f798])).
fof(f62427,plain,(
  $false | ($spl8 | $spl127)),
  inference(subsumption_resolution,[],[f62299,f1013])).
fof(f62299,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl127),
  inference(resolution,[],[f14417,f810])).
fof(f10798,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | $spl121),
  inference(cnf_transformation,[],[f10798_D])).
fof(f10798_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f62135,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl366),
  inference(cnf_transformation,[],[f62135_D])).
fof(f62135_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f62011,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | $spl118),
  inference(resolution,[],[f10649,f808])).
fof(f60543,plain,(
  $false | ($spl40 | $spl74)),
  inference(subsumption_resolution,[],[f60387,f1503])).
fof(f60387,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl74),
  inference(resolution,[],[f3892,f734])).
fof(f60542,plain,(
  $false | ($spl40 | $spl74)),
  inference(subsumption_resolution,[],[f60386,f1503])).
fof(f60386,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl74),
  inference(resolution,[],[f3892,f734])).
fof(f60541,plain,(
  $false | ($spl40 | $spl74)),
  inference(subsumption_resolution,[],[f60385,f1503])).
fof(f60385,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl74),
  inference(resolution,[],[f3892,f679])).
fof(f60540,plain,(
  $false | ($spl40 | $spl74)),
  inference(subsumption_resolution,[],[f60384,f1503])).
fof(f60384,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl74),
  inference(resolution,[],[f3892,f679])).
fof(f5194,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl90),
  inference(cnf_transformation,[],[f5194_D])).
fof(f5194_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f60215,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl40),
  inference(subsumption_resolution,[],[f60179,f758])).
fof(f60179,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl40),
  inference(resolution,[],[f33087,f704])).
fof(f33087,plain,(
  ( ! [X145] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X145)) ) | $spl40),
  inference(resolution,[],[f10329,f706])).
fof(f60214,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl40),
  inference(subsumption_resolution,[],[f60178,f758])).
fof(f60178,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl40),
  inference(resolution,[],[f33087,f708])).
fof(f60175,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl40),
  inference(subsumption_resolution,[],[f60139,f758])).
fof(f60139,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl40),
  inference(resolution,[],[f33086,f704])).
fof(f33086,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl40),
  inference(resolution,[],[f10329,f705])).
fof(f60174,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl40),
  inference(subsumption_resolution,[],[f60138,f758])).
fof(f60138,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl40),
  inference(resolution,[],[f33086,f708])).
fof(f60047,plain,(
  $false | ($spl41 | $spl111)),
  inference(subsumption_resolution,[],[f59897,f9984])).
fof(f59897,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl41),
  inference(resolution,[],[f3732,f853])).
fof(f853,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f567])).
fof(f567,plain,(
  ! [X0] : ((s('type.integer.int',X0) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ABS_EQ_ID')).
fof(f3732,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl41),
  inference(resolution,[],[f1504,f798])).
fof(f59006,plain,(
  $false | ($spl3 | $spl63)),
  inference(subsumption_resolution,[],[f2793,f4551])).
fof(f4551,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl3),
  inference(resolution,[],[f879,f798])).
fof(f41849,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl328),
  inference(cnf_transformation,[],[f41849_D])).
fof(f41849_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f9008,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | $spl105),
  inference(cnf_transformation,[],[f9008_D])).
fof(f9008_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f58450,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl323),
  inference(forward_demodulation,[],[f58323,f856])).
fof(f58323,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK13)))) | $spl323),
  inference(resolution,[],[f41385,f800])).
fof(f41675,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl327),
  inference(cnf_transformation,[],[f41675_D])).
fof(f41675_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl327),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl327])])).
fof(f50225,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl340),
  inference(cnf_transformation,[],[f50225_D])).
fof(f50225_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f58322,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl323),
  inference(resolution,[],[f41385,f803])).
fof(f58303,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl365),
  inference(cnf_transformation,[],[f58303_D])).
fof(f58303_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) <=> ~$spl365),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl365])])).
fof(f57180,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl362),
  inference(cnf_transformation,[],[f57180_D])).
fof(f57180_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f58169,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl107),
  inference(resolution,[],[f9009,f809])).
fof(f41676,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl326),
  inference(cnf_transformation,[],[f41676_D])).
fof(f41676_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f57687,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | $spl100),
  inference(resolution,[],[f9003,f808])).
fof(f10808,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | $spl122),
  inference(cnf_transformation,[],[f10808_D])).
fof(f10808_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f9006,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl102),
  inference(cnf_transformation,[],[f9006_D])).
fof(f9006_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f57678,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | $spl100),
  inference(resolution,[],[f9003,f803])).
fof(f57674,plain,(
  $false | ($spl100 | $spl322)),
  inference(subsumption_resolution,[],[f41545,f41384])).
fof(f41384,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl322),
  inference(cnf_transformation,[],[f41384_D])).
fof(f41384_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f41545,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl100),
  inference(resolution,[],[f9003,f679])).
fof(f57673,plain,(
  $false | ($spl100 | $spl322)),
  inference(subsumption_resolution,[],[f41546,f41384])).
fof(f41546,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl100),
  inference(resolution,[],[f9003,f679])).
fof(f57672,plain,(
  $false | ($spl100 | $spl322)),
  inference(subsumption_resolution,[],[f41547,f41384])).
fof(f41547,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl100),
  inference(resolution,[],[f9003,f734])).
fof(f57671,plain,(
  $false | ($spl100 | $spl322)),
  inference(subsumption_resolution,[],[f41548,f41384])).
fof(f41548,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl100),
  inference(resolution,[],[f9003,f734])).
fof(f57670,plain,(
  $false | ($spl102 | $spl106)),
  inference(subsumption_resolution,[],[f49461,f9010])).
fof(f49461,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl102),
  inference(resolution,[],[f9006,f679])).
fof(f57669,plain,(
  $false | ($spl102 | $spl106)),
  inference(subsumption_resolution,[],[f49462,f9010])).
fof(f49462,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl102),
  inference(resolution,[],[f9006,f679])).
fof(f57668,plain,(
  $false | ($spl102 | $spl106)),
  inference(subsumption_resolution,[],[f49463,f9010])).
fof(f49463,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl102),
  inference(resolution,[],[f9006,f734])).
fof(f57667,plain,(
  $false | ($spl102 | $spl106)),
  inference(subsumption_resolution,[],[f49464,f9010])).
fof(f49464,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl102),
  inference(resolution,[],[f9006,f734])).
fof(f57658,plain,(
  $false | ($spl101 | $spl322)),
  inference(subsumption_resolution,[],[f57527,f14144])).
fof(f14144,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | $spl101),
  inference(resolution,[],[f9004,f804])).
fof(f57527,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | $spl322),
  inference(resolution,[],[f41384,f806])).
fof(f41539,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl325),
  inference(cnf_transformation,[],[f41539_D])).
fof(f41539_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl325),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl325])])).
fof(f57201,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl103),
  inference(resolution,[],[f9005,f809])).
fof(f9005,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl103),
  inference(cnf_transformation,[],[f9005_D])).
fof(f9005_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f57198,plain,(
  $false | ($spl101 | $spl102)),
  inference(subsumption_resolution,[],[f49460,f9004])).
fof(f49460,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl102),
  inference(resolution,[],[f9006,f805])).
fof(f805,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) )),
  inference(cnf_transformation,[],[f556])).
fof(f57197,plain,(
  $false | ($spl101 | $spl102)),
  inference(subsumption_resolution,[],[f49594,f9004])).
fof(f49594,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl102),
  inference(forward_demodulation,[],[f49467,f817])).
fof(f49467,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl102),
  inference(resolution,[],[f9006,f802])).
fof(f57196,plain,(
  $false | ($spl101 | $spl102)),
  inference(subsumption_resolution,[],[f49595,f9004])).
fof(f49595,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl102),
  inference(forward_demodulation,[],[f49468,f844])).
fof(f49468,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl102),
  inference(resolution,[],[f9006,f804])).
fof(f57047,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl103),
  inference(resolution,[],[f9005,f809])).
fof(f14277,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl125),
  inference(cnf_transformation,[],[f14277_D])).
fof(f14277_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f9007,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | $spl104),
  inference(cnf_transformation,[],[f9007_D])).
fof(f9007_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f56912,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl101),
  inference(resolution,[],[f9004,f800])).
fof(f3454,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl66),
  inference(cnf_transformation,[],[f3454_D])).
fof(f3454_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f56772,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl3 | $spl4)),
  inference(resolution,[],[f4686,f800])).
fof(f49298,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl337),
  inference(cnf_transformation,[],[f49298_D])).
fof(f49298_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl337),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl337])])).
fof(f14419,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl128),
  inference(cnf_transformation,[],[f14419_D])).
fof(f14419_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f56769,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl3 | $spl4)),
  inference(resolution,[],[f4686,f809])).
fof(f3473,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl68),
  inference(cnf_transformation,[],[f3473_D])).
fof(f3473_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f56412,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14))))) ) | $spl110),
  inference(subsumption_resolution,[],[f56376,f758])).
fof(f56376,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14))))) ) | $spl110),
  inference(resolution,[],[f55577,f704])).
fof(f55577,plain,(
  ( ! [X11] : (sP12(X11,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) ) | $spl110),
  inference(superposition,[],[f21089,f9983])).
fof(f21089,plain,(
  ( ! [X223,X224] : (sP12(X223,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X224)))))) )),
  inference(resolution,[],[f854,f705])).
fof(f854,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))) )),
  inference(cnf_transformation,[],[f383])).
fof(f383,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ABS_POS')).
fof(f56411,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14))))) ) | $spl110),
  inference(subsumption_resolution,[],[f56375,f758])).
fof(f56375,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14))))) ) | $spl110),
  inference(resolution,[],[f55577,f708])).
fof(f5192,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl88),
  inference(cnf_transformation,[],[f5192_D])).
fof(f5192_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f56314,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl35),
  inference(subsumption_resolution,[],[f56278,f758])).
fof(f56278,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl35),
  inference(resolution,[],[f17916,f704])).
fof(f17916,plain,(
  ( ! [X145] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X145)) ) | $spl35),
  inference(resolution,[],[f3645,f706])).
fof(f3645,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl35),
  inference(forward_demodulation,[],[f3489,f841])).
fof(f841,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f374])).
fof(f374,plain,(
  ! [X0] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f243])).
fof(f243,axiom,(
  ! [X9] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_NEG_GE|48|')).
fof(f3489,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl35),
  inference(resolution,[],[f1498,f839])).
fof(f839,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f372])).
fof(f372,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))))),
  inference(rectify,[],[f171])).
fof(f171,axiom,(
  ! [X9] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_LE_NEGTOTAL')).
fof(f56313,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl35),
  inference(subsumption_resolution,[],[f56277,f758])).
fof(f56277,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl35),
  inference(resolution,[],[f17916,f708])).
fof(f56274,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(subsumption_resolution,[],[f56238,f758])).
fof(f56238,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(resolution,[],[f17915,f704])).
fof(f17915,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl35),
  inference(resolution,[],[f3645,f705])).
fof(f56273,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(subsumption_resolution,[],[f56237,f758])).
fof(f56237,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(resolution,[],[f17915,f708])).
fof(f9484,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl108),
  inference(cnf_transformation,[],[f9484_D])).
fof(f9484_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f56234,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl35),
  inference(subsumption_resolution,[],[f56198,f758])).
fof(f56198,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl35),
  inference(resolution,[],[f17262,f704])).
fof(f17262,plain,(
  ( ! [X145] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X145)) ) | $spl35),
  inference(resolution,[],[f3492,f706])).
fof(f3492,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl35),
  inference(resolution,[],[f1498,f798])).
fof(f56233,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl35),
  inference(subsumption_resolution,[],[f56197,f758])).
fof(f56197,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl35),
  inference(resolution,[],[f17262,f708])).
fof(f56194,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(subsumption_resolution,[],[f56158,f758])).
fof(f56158,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(resolution,[],[f17261,f704])).
fof(f17261,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl35),
  inference(resolution,[],[f3492,f705])).
fof(f56193,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(subsumption_resolution,[],[f56157,f758])).
fof(f56157,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(resolution,[],[f17261,f708])).
fof(f3474,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl69),
  inference(cnf_transformation,[],[f3474_D])).
fof(f3474_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f55877,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl360),
  inference(cnf_transformation,[],[f55877_D])).
fof(f55877_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f55875,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl110),
  inference(forward_demodulation,[],[f55730,f9983])).
fof(f55730,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl110),
  inference(superposition,[],[f864,f9983])).
fof(f864,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29(X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f571])).
fof(f571,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK29])],[f424])).
fof(f424,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(ennf_transformation,[],[f393])).
fof(f393,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) => ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f233])).
fof(f233,axiom,(
  ! [X17] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X17)))) => ? [X10] : s('type.integer.int',X17) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.NUM_POSINT_EXISTS')).
fof(f3892,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl74),
  inference(cnf_transformation,[],[f3892_D])).
fof(f3892_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3894,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl77),
  inference(cnf_transformation,[],[f3894_D])).
fof(f3894_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f55409,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl41),
  inference(forward_demodulation,[],[f55408,f821])).
fof(f821,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0] : s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9] : s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_NEG_LT|48|')).
fof(f55408,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl41),
  inference(forward_demodulation,[],[f55251,f846])).
fof(f846,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f123])).
fof(f123,axiom,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_NEG_|48|')).
fof(f55251,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl41),
  inference(resolution,[],[f1504,f800])).
fof(f36094,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl290),
  inference(cnf_transformation,[],[f36094_D])).
fof(f36094_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f55118,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | $spl343),
  inference(resolution,[],[f52030,f800])).
fof(f52030,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl343),
  inference(cnf_transformation,[],[f52030_D])).
fof(f52030_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) <=> ~$spl343),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl343])])).
fof(f10650,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl119),
  inference(cnf_transformation,[],[f10650_D])).
fof(f10650_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f52370,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl344),
  inference(cnf_transformation,[],[f52370_D])).
fof(f52370_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f55117,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl343),
  inference(resolution,[],[f52030,f803])).
fof(f55109,plain,(
  $false | $spl342),
  inference(subsumption_resolution,[],[f54983,f868])).
fof(f868,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(rectify,[],[f238])).
fof(f238,axiom,(
  ! [X9] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X9))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_LE_REFL')).
fof(f54983,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl342),
  inference(resolution,[],[f52031,f799])).
fof(f55108,plain,(
  $false | $spl342),
  inference(subsumption_resolution,[],[f54982,f868])).
fof(f54982,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl342),
  inference(resolution,[],[f52031,f797])).
fof(f55107,plain,(
  $false | $spl342),
  inference(subsumption_resolution,[],[f54981,f52031])).
fof(f54981,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl342),
  inference(resolution,[],[f52031,f734])).
fof(f55106,plain,(
  $false | $spl342),
  inference(subsumption_resolution,[],[f54980,f52031])).
fof(f54980,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl342),
  inference(resolution,[],[f52031,f734])).
fof(f55105,plain,(
  $false | $spl342),
  inference(subsumption_resolution,[],[f54979,f52031])).
fof(f54979,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl342),
  inference(resolution,[],[f52031,f679])).
fof(f55104,plain,(
  $false | $spl342),
  inference(subsumption_resolution,[],[f54978,f52031])).
fof(f54978,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl342),
  inference(resolution,[],[f52031,f679])).
fof(f9010,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl106),
  inference(cnf_transformation,[],[f9010_D])).
fof(f9010_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f54321,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | $spl101),
  inference(resolution,[],[f14144,f800])).
fof(f1498,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl35),
  inference(cnf_transformation,[],[f1498_D])).
fof(f1498_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f10809,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | $spl123),
  inference(cnf_transformation,[],[f10809_D])).
fof(f10809_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f50196,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl323),
  inference(superposition,[],[f41385,f853])).
fof(f50226,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl323),
  inference(forward_demodulation,[],[f50097,f856])).
fof(f50097,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK13)))) | $spl323),
  inference(resolution,[],[f41385,f800])).
fof(f9009,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl107),
  inference(cnf_transformation,[],[f9009_D])).
fof(f9009_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f50093,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | $spl323),
  inference(resolution,[],[f41385,f808])).
fof(f19170,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl30),
  inference(superposition,[],[f1484,f853])).
fof(f26139,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl30),
  inference(superposition,[],[f1484,f853])).
fof(f1011,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl7),
  inference(cnf_transformation,[],[f1011_D])).
fof(f1011_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f40317,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl0),
  inference(superposition,[],[f5638,f853])).
fof(f5638,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl0),
  inference(resolution,[],[f876,f679])).
fof(f54114,plain,(
  ( ! [X32] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X32)) ) | $spl358),
  inference(cnf_transformation,[],[f54114_D])).
fof(f54114_D,plain,(
  ( ! [X32] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X32)) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f18586,plain,(
  ( ! [X0] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl0),
  inference(superposition,[],[f5725,f853])).
fof(f5725,plain,(
  ( ! [X145] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))),X145)) ) | $spl0),
  inference(resolution,[],[f876,f706])).
fof(f54113,plain,(
  ( ! [X31] : (sP12(X31,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) ) | $spl356),
  inference(cnf_transformation,[],[f54113_D])).
fof(f54113_D,plain,(
  ( ! [X31] : (sP12(X31,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f18555,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl0),
  inference(superposition,[],[f5724,f853])).
fof(f5724,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl0),
  inference(resolution,[],[f876,f705])).
fof(f1320,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl20),
  inference(cnf_transformation,[],[f1320_D])).
fof(f1320_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f12023,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl0),
  inference(superposition,[],[f876,f853])).
fof(f5746,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl0),
  inference(superposition,[],[f876,f853])).
fof(f26135,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl0),
  inference(superposition,[],[f876,f853])).
fof(f47327,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl0),
  inference(superposition,[],[f876,f853])).
fof(f26143,plain,(
  ( ! [X32] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X32) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl0),
  inference(superposition,[],[f5725,f853])).
fof(f26142,plain,(
  ( ! [X31] : (sP12(X31,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl0),
  inference(superposition,[],[f5724,f853])).
fof(f54093,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) ) | $spl352),
  inference(cnf_transformation,[],[f54093_D])).
fof(f54093_D,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f54106,plain,(
  ( ! [X31] : (sP12(X31,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | ($spl0 | $spl60)),
  inference(forward_demodulation,[],[f26142,f2640])).
fof(f54091,plain,(
  ( ! [X0] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)),X0)) ) | $spl350),
  inference(cnf_transformation,[],[f54091_D])).
fof(f54091_D,plain,(
  ( ! [X0] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)),X0)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f54105,plain,(
  ( ! [X32] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)),X32) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | ($spl0 | $spl60)),
  inference(forward_demodulation,[],[f26143,f2640])).
fof(f5060,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl86),
  inference(cnf_transformation,[],[f5060_D])).
fof(f5060_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f54104,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl42),
  inference(forward_demodulation,[],[f40036,f846])).
fof(f40036,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl42),
  inference(resolution,[],[f1510,f803])).
fof(f40173,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl42),
  inference(superposition,[],[f1510,f853])).
fof(f26137,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl26),
  inference(superposition,[],[f1480,f853])).
fof(f18867,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl26),
  inference(superposition,[],[f1480,f853])).
fof(f54101,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl355),
  inference(cnf_transformation,[],[f54101_D])).
fof(f54101_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) <=> ~$spl355),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl355])])).
fof(f19450,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl182),
  inference(cnf_transformation,[],[f19450_D])).
fof(f19450_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f50094,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl323),
  inference(resolution,[],[f41385,f809])).
fof(f54092,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | ($spl0 | $spl60)),
  inference(forward_demodulation,[],[f18555,f2640])).
fof(f54090,plain,(
  ( ! [X0] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | ($spl0 | $spl60)),
  inference(forward_demodulation,[],[f18586,f2640])).
fof(f51989,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl60),
  inference(forward_demodulation,[],[f51835,f841])).
fof(f51835,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl60),
  inference(superposition,[],[f33229,f2640])).
fof(f33229,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f839,f841])).
fof(f54083,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | ($spl11 | $spl60)),
  inference(subsumption_resolution,[],[f54082,f1162])).
fof(f54082,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl60),
  inference(forward_demodulation,[],[f51990,f2640])).
fof(f51990,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl60),
  inference(forward_demodulation,[],[f51836,f821])).
fof(f51836,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl60),
  inference(superposition,[],[f37776,f2640])).
fof(f37776,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f786,f822])).
fof(f822,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f360])).
fof(f360,plain,(
  ! [X0] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_NEG_GT|48|')).
fof(f786,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X9] : (s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_LT_NEGTOTAL')).
fof(f9983,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl110),
  inference(cnf_transformation,[],[f9983_D])).
fof(f9983_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f5191,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl89),
  inference(cnf_transformation,[],[f5191_D])).
fof(f5191_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f52000,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl60),
  inference(forward_demodulation,[],[f51823,f841])).
fof(f51823,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl60),
  inference(superposition,[],[f853,f2640])).
fof(f9984,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl111),
  inference(cnf_transformation,[],[f9984_D])).
fof(f9984_D,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f52002,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl60),
  inference(forward_demodulation,[],[f51821,f841])).
fof(f51821,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl60),
  inference(superposition,[],[f852,f2640])).
fof(f852,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f567])).
fof(f1497,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl34),
  inference(cnf_transformation,[],[f1497_D])).
fof(f1497_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f14418,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl129),
  inference(cnf_transformation,[],[f14418_D])).
fof(f14418_D,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f54081,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl64),
  inference(forward_demodulation,[],[f54080,f2795])).
fof(f54080,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl64),
  inference(forward_demodulation,[],[f52584,f844])).
fof(f52584,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | $spl64),
  inference(superposition,[],[f852,f2795])).
fof(f19449,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl183),
  inference(cnf_transformation,[],[f19449_D])).
fof(f19449_D,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) <=> ~$spl183),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl183])])).
fof(f54079,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl64),
  inference(forward_demodulation,[],[f52585,f841])).
fof(f52585,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl64),
  inference(superposition,[],[f852,f2795])).
fof(f54078,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl64),
  inference(forward_demodulation,[],[f54077,f2795])).
fof(f54077,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl64),
  inference(forward_demodulation,[],[f52586,f844])).
fof(f52586,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | $spl64),
  inference(superposition,[],[f853,f2795])).
fof(f5193,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl91),
  inference(cnf_transformation,[],[f5193_D])).
fof(f5193_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f54076,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl64),
  inference(forward_demodulation,[],[f52587,f841])).
fof(f52587,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl64),
  inference(superposition,[],[f853,f2795])).
fof(f54075,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl348),
  inference(cnf_transformation,[],[f54075_D])).
fof(f54075_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f54073,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl64),
  inference(forward_demodulation,[],[f52592,f2795])).
fof(f52592,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl64),
  inference(superposition,[],[f864,f2795])).
fof(f54071,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl64),
  inference(forward_demodulation,[],[f52599,f841])).
fof(f52599,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl64),
  inference(superposition,[],[f33229,f2795])).
fof(f19431,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl156),
  inference(cnf_transformation,[],[f19431_D])).
fof(f19431_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f54070,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl64),
  inference(forward_demodulation,[],[f54069,f2795])).
fof(f54069,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f52600,f821])).
fof(f52600,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(superposition,[],[f37776,f2795])).
fof(f1503,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl40),
  inference(cnf_transformation,[],[f1503_D])).
fof(f1503_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f54068,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl64),
  inference(forward_demodulation,[],[f54067,f2795])).
fof(f54067,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f52601,f822])).
fof(f52601,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(superposition,[],[f37776,f2795])).
fof(f54065,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl11 | $spl60)),
  inference(subsumption_resolution,[],[f54064,f1162])).
fof(f54064,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl60),
  inference(forward_demodulation,[],[f53682,f2640])).
fof(f53682,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl60),
  inference(forward_demodulation,[],[f51837,f822])).
fof(f51837,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl60),
  inference(superposition,[],[f37776,f2640])).
fof(f53912,plain,(
  $false | ($spl4 | $spl9 | $spl35 | $spl60)),
  inference(subsumption_resolution,[],[f53713,f880])).
fof(f53713,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | ($spl9 | $spl35 | $spl60)),
  inference(backward_demodulation,[],[f52001,f1012])).
fof(f52001,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ($spl35 | $spl60)),
  inference(subsumption_resolution,[],[f52000,f3645])).
fof(f53689,plain,(
  ( ! [X427] : (sP12(X427,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl346),
  inference(cnf_transformation,[],[f53689_D])).
fof(f53689_D,plain,(
  ( ! [X427] : (sP12(X427,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f53709,plain,(
  ( ! [X16] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | sP12(X16,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f18568,f845])).
fof(f18568,plain,(
  ( ! [X16] : (sP12(X16,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl0),
  inference(superposition,[],[f5724,f833])).
fof(f18569,plain,(
  ( ! [X17] : (sP12(X17,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl0),
  inference(superposition,[],[f5724,f835])).
fof(f1522,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl51),
  inference(cnf_transformation,[],[f1522_D])).
fof(f1522_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f18570,plain,(
  ( ! [X18] : (sP12(X18,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl0),
  inference(superposition,[],[f5724,f836])).
fof(f836,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X0)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26(X0))) )),
  inference(cnf_transformation,[],[f564])).
fof(f18571,plain,(
  ( ! [X19] : (sP12(X19,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl0),
  inference(superposition,[],[f5724,f837])).
fof(f18572,plain,(
  ( ! [X20] : (sP12(X20,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl0),
  inference(superposition,[],[f5724,f838])).
fof(f838,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27(X0)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26(X0))) )),
  inference(cnf_transformation,[],[f564])).
fof(f1528,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl55),
  inference(cnf_transformation,[],[f1528_D])).
fof(f1528_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f53708,plain,(
  ( ! [X21] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP12(X21,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f53707,f856])).
fof(f53707,plain,(
  ( ! [X21] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP12(X21,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f18573,f844])).
fof(f18573,plain,(
  ( ! [X21] : (sP12(X21,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl0),
  inference(superposition,[],[f5724,f848])).
fof(f848,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f565])).
fof(f565,plain,(
  ! [X0] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(nnf_transformation,[],[f380])).
fof(f380,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ABS_LE|48|')).
fof(f24042,plain,(
  ( ! [X391] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X391)) ) | $spl194),
  inference(cnf_transformation,[],[f24042_D])).
fof(f24042_D,plain,(
  ( ! [X391] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X391)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f53701,plain,(
  ( ! [X16] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X16)) ) | $spl0),
  inference(forward_demodulation,[],[f18599,f845])).
fof(f18599,plain,(
  ( ! [X16] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X16) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl0),
  inference(superposition,[],[f5725,f833])).
fof(f18600,plain,(
  ( ! [X17] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X17) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl0),
  inference(superposition,[],[f5725,f835])).
fof(f18601,plain,(
  ( ! [X18] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X18) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl0),
  inference(superposition,[],[f5725,f836])).
fof(f18602,plain,(
  ( ! [X19] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X19) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl0),
  inference(superposition,[],[f5725,f837])).
fof(f18603,plain,(
  ( ! [X20] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl0),
  inference(superposition,[],[f5725,f838])).
fof(f53700,plain,(
  ( ! [X21] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X21)) ) | $spl0),
  inference(forward_demodulation,[],[f53699,f856])).
fof(f53699,plain,(
  ( ! [X21] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X21)) ) | $spl0),
  inference(forward_demodulation,[],[f18604,f844])).
fof(f18604,plain,(
  ( ! [X21] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X21) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl0),
  inference(superposition,[],[f5725,f848])).
fof(f53695,plain,(
  ( ! [X404] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X404)) ) | $spl0),
  inference(forward_demodulation,[],[f53694,f856])).
fof(f53694,plain,(
  ( ! [X404] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X404)) ) | $spl0),
  inference(forward_demodulation,[],[f28161,f844])).
fof(f28161,plain,(
  ( ! [X404] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X404) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl0),
  inference(superposition,[],[f5725,f848])).
fof(f53693,plain,(
  ( ! [X405] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP12(X405,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f53692,f856])).
fof(f53692,plain,(
  ( ! [X405] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP12(X405,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f28162,f844])).
fof(f28162,plain,(
  ( ! [X405] : (sP12(X405,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl0),
  inference(superposition,[],[f5724,f848])).
fof(f53691,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f53690,f856])).
fof(f53690,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f28163,f844])).
fof(f28163,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(superposition,[],[f876,f848])).
fof(f32041,plain,(
  ( ! [X420] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X420) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl0),
  inference(superposition,[],[f5725,f837])).
fof(f32042,plain,(
  ( ! [X421] : (sP12(X421,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl0),
  inference(superposition,[],[f5724,f837])).
fof(f32043,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(superposition,[],[f876,f837])).
fof(f35001,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl30),
  inference(superposition,[],[f1484,f836])).
fof(f35002,plain,(
  ( ! [X425] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X425) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl0),
  inference(superposition,[],[f5725,f836])).
fof(f35003,plain,(
  ( ! [X426] : (sP12(X426,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl0),
  inference(superposition,[],[f5724,f836])).
fof(f35004,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(superposition,[],[f876,f836])).
fof(f36473,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl30),
  inference(superposition,[],[f1484,f835])).
fof(f36474,plain,(
  ( ! [X426] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X426) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl0),
  inference(superposition,[],[f5725,f835])).
fof(f36475,plain,(
  ( ! [X427] : (sP12(X427,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))) ) | $spl0),
  inference(superposition,[],[f5724,f835])).
fof(f36476,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(superposition,[],[f876,f835])).
fof(f3891,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl75),
  inference(cnf_transformation,[],[f3891_D])).
fof(f3891_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f53687,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl0),
  inference(forward_demodulation,[],[f40329,f845])).
fof(f40329,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl0),
  inference(superposition,[],[f5638,f833])).
fof(f40330,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(superposition,[],[f5638,f835])).
fof(f40331,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(superposition,[],[f5638,f836])).
fof(f40332,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(superposition,[],[f5638,f837])).
fof(f40333,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(superposition,[],[f5638,f838])).
fof(f53686,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl0),
  inference(forward_demodulation,[],[f53685,f856])).
fof(f53685,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl0),
  inference(forward_demodulation,[],[f40334,f844])).
fof(f40334,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(superposition,[],[f5638,f848])).
fof(f53680,plain,(
  $false | ($spl35 | $spl40 | $spl60)),
  inference(subsumption_resolution,[],[f52001,f32848])).
fof(f32848,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl40),
  inference(resolution,[],[f10324,f852])).
fof(f10649,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl118),
  inference(cnf_transformation,[],[f10649_D])).
fof(f10649_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f53129,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl9 | $spl60)),
  inference(forward_demodulation,[],[f53128,f856])).
fof(f53128,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int',sK13)))) | ($spl9 | $spl60)),
  inference(forward_demodulation,[],[f52998,f2640])).
fof(f52998,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int',sK13)))) | $spl9),
  inference(resolution,[],[f1012,f800])).
fof(f52997,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl9),
  inference(resolution,[],[f1012,f803])).
fof(f14417,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl127),
  inference(cnf_transformation,[],[f14417_D])).
fof(f14417_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f52995,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl9),
  inference(resolution,[],[f1012,f809])).
fof(f52769,plain,(
  $false | ($spl24 | $spl64)),
  inference(subsumption_resolution,[],[f52768,f868])).
fof(f52768,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ($spl24 | $spl64)),
  inference(forward_demodulation,[],[f50629,f2795])).
fof(f50629,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl24),
  inference(resolution,[],[f1324,f797])).
fof(f52767,plain,(
  $false | ($spl24 | $spl64)),
  inference(subsumption_resolution,[],[f52766,f868])).
fof(f52766,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ($spl24 | $spl64)),
  inference(forward_demodulation,[],[f50630,f2795])).
fof(f50630,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl24),
  inference(resolution,[],[f1324,f799])).
fof(f52614,plain,(
  $false | ($spl64 | $spl85)),
  inference(subsumption_resolution,[],[f52390,f868])).
fof(f52390,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ($spl64 | $spl85)),
  inference(backward_demodulation,[],[f2795,f4682])).
fof(f52613,plain,(
  $false | ($spl24 | $spl64 | $spl105)),
  inference(subsumption_resolution,[],[f52379,f45206])).
fof(f45206,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | $spl105),
  inference(resolution,[],[f9008,f802])).
fof(f52379,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ($spl24 | $spl64)),
  inference(backward_demodulation,[],[f2795,f1324])).
fof(f52247,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | ($spl7 | $spl64)),
  inference(resolution,[],[f9611,f800])).
fof(f9611,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ($spl7 | $spl64)),
  inference(backward_demodulation,[],[f4877,f1011])).
fof(f4877,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl64),
  inference(superposition,[],[f845,f2795])).
fof(f52246,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl7 | $spl64)),
  inference(resolution,[],[f9611,f803])).
fof(f52226,plain,(
  $false | ($spl7 | $spl8)),
  inference(subsumption_resolution,[],[f2501,f1013])).
fof(f2501,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl7),
  inference(resolution,[],[f1011,f807])).
fof(f807,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f558])).
fof(f52215,plain,(
  $false | ($spl7 | $spl8)),
  inference(subsumption_resolution,[],[f47629,f1013])).
fof(f47629,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl7),
  inference(resolution,[],[f1011,f807])).
fof(f52176,plain,(
  $false | ($spl7 | $spl8)),
  inference(subsumption_resolution,[],[f12536,f1013])).
fof(f12536,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl7),
  inference(resolution,[],[f1011,f807])).
fof(f52164,plain,(
  $false | ($spl6 | $spl60)),
  inference(subsumption_resolution,[],[f52038,f868])).
fof(f52038,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ($spl6 | $spl60)),
  inference(resolution,[],[f51624,f799])).
fof(f51624,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ($spl6 | $spl60)),
  inference(backward_demodulation,[],[f2640,f1010])).
fof(f52163,plain,(
  $false | ($spl6 | $spl60)),
  inference(subsumption_resolution,[],[f52037,f868])).
fof(f52037,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ($spl6 | $spl60)),
  inference(resolution,[],[f51624,f797])).
fof(f52162,plain,(
  $false | ($spl6 | $spl60)),
  inference(subsumption_resolution,[],[f52036,f51624])).
fof(f52036,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ($spl6 | $spl60)),
  inference(resolution,[],[f51624,f734])).
fof(f52161,plain,(
  $false | ($spl6 | $spl60)),
  inference(subsumption_resolution,[],[f52035,f51624])).
fof(f52035,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ($spl6 | $spl60)),
  inference(resolution,[],[f51624,f734])).
fof(f52160,plain,(
  $false | ($spl6 | $spl60)),
  inference(subsumption_resolution,[],[f52034,f51624])).
fof(f52034,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ($spl6 | $spl60)),
  inference(resolution,[],[f51624,f679])).
fof(f52159,plain,(
  $false | ($spl6 | $spl60)),
  inference(subsumption_resolution,[],[f52033,f51624])).
fof(f52033,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ($spl6 | $spl60)),
  inference(resolution,[],[f51624,f679])).
fof(f52031,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl342),
  inference(cnf_transformation,[],[f52031_D])).
fof(f52031_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f52032,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) ) | ($spl6 | $spl60)),
  inference(subsumption_resolution,[],[f52024,f758])).
fof(f52024,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) ) | ($spl6 | $spl60)),
  inference(resolution,[],[f51629,f704])).
fof(f51629,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) ) | ($spl6 | $spl60)),
  inference(backward_demodulation,[],[f2640,f51203])).
fof(f51203,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) ) | $spl6),
  inference(resolution,[],[f1010,f705])).
fof(f52029,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) ) | ($spl6 | $spl60)),
  inference(subsumption_resolution,[],[f52023,f758])).
fof(f52023,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) ) | ($spl6 | $spl60)),
  inference(resolution,[],[f51629,f708])).
fof(f9990,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl112),
  inference(cnf_transformation,[],[f9990_D])).
fof(f9990_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f51987,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl60),
  inference(forward_demodulation,[],[f51828,f2640])).
fof(f51828,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl60),
  inference(superposition,[],[f864,f2640])).
fof(f52005,plain,(
  $false | ($spl35 | $spl101 | $spl180)),
  inference(subsumption_resolution,[],[f14274,f19447])).
fof(f14274,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl35 | $spl101)),
  inference(subsumption_resolution,[],[f14146,f3488])).
fof(f3488,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl35),
  inference(resolution,[],[f1498,f852])).
fof(f52003,plain,(
  $false | ($spl35 | $spl182)),
  inference(subsumption_resolution,[],[f19450,f3488])).
fof(f9003,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl100),
  inference(cnf_transformation,[],[f9003_D])).
fof(f9003_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f45208,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl105),
  inference(resolution,[],[f9008,f800])).
fof(f45360,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl105),
  inference(forward_demodulation,[],[f45211,f845])).
fof(f45211,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl105),
  inference(resolution,[],[f9008,f803])).
fof(f51986,plain,(
  $false | ($spl35 | $spl60 | $spl111)),
  inference(subsumption_resolution,[],[f51985,f3645])).
fof(f51985,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl60 | $spl111)),
  inference(forward_demodulation,[],[f51984,f841])).
fof(f51984,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ($spl60 | $spl111)),
  inference(subsumption_resolution,[],[f51823,f9984])).
fof(f51858,plain,(
  $false | ($spl60 | $spl65)),
  inference(subsumption_resolution,[],[f51672,f2794])).
fof(f2794,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl65),
  inference(cnf_transformation,[],[f2794_D])).
fof(f2794_D,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f51672,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl60),
  inference(superposition,[],[f845,f2640])).
fof(f51849,plain,(
  $false | ($spl13 | $spl60 | $spl102)),
  inference(subsumption_resolution,[],[f51626,f1163])).
fof(f51626,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | ($spl60 | $spl102)),
  inference(backward_demodulation,[],[f2640,f9006])).
fof(f49729,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl339),
  inference(cnf_transformation,[],[f49729_D])).
fof(f49729_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl339),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl339])])).
fof(f2640,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl60),
  inference(cnf_transformation,[],[f2640_D])).
fof(f2640_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f50475,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl21),
  inference(resolution,[],[f1319,f809])).
fof(f51614,plain,(
  $false | ($spl21 | $spl59)),
  inference(subsumption_resolution,[],[f51483,f1319])).
fof(f51483,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl59),
  inference(resolution,[],[f2638,f796])).
fof(f51613,plain,(
  $false | ($spl21 | $spl59)),
  inference(subsumption_resolution,[],[f51482,f1319])).
fof(f51482,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl59),
  inference(resolution,[],[f2638,f798])).
fof(f51612,plain,(
  $false | ($spl6 | $spl59)),
  inference(subsumption_resolution,[],[f51480,f1010])).
fof(f51480,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl59),
  inference(resolution,[],[f2638,f810])).
fof(f51108,plain,(
  $false | ($spl7 | $spl21 | $spl59)),
  inference(subsumption_resolution,[],[f51107,f1319])).
fof(f51107,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ($spl7 | $spl59)),
  inference(subsumption_resolution,[],[f50979,f15739])).
fof(f15739,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl59),
  inference(resolution,[],[f2638,f811])).
fof(f811,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) != s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f560])).
fof(f50979,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl7),
  inference(resolution,[],[f2507,f809])).
fof(f50956,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl7),
  inference(forward_demodulation,[],[f50828,f856])).
fof(f50828,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int',sK13)))) | $spl7),
  inference(resolution,[],[f2501,f800])).
fof(f50827,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl7),
  inference(resolution,[],[f2501,f803])).
fof(f50825,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl7),
  inference(resolution,[],[f2501,f809])).
fof(f50464,plain,(
  $false | ($spl3 | $spl4 | $spl20)),
  inference(subsumption_resolution,[],[f50463,f4686])).
fof(f50463,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl20),
  inference(forward_demodulation,[],[f50333,f817])).
fof(f50333,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl20),
  inference(resolution,[],[f1320,f802])).
fof(f50462,plain,(
  $false | ($spl3 | $spl4 | $spl20)),
  inference(subsumption_resolution,[],[f50326,f4686])).
fof(f50326,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl20),
  inference(resolution,[],[f1320,f805])).
fof(f50096,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl323),
  inference(resolution,[],[f41385,f803])).
fof(f50068,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl40),
  inference(subsumption_resolution,[],[f50032,f758])).
fof(f50032,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl40),
  inference(resolution,[],[f33087,f704])).
fof(f50067,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl40),
  inference(subsumption_resolution,[],[f50031,f758])).
fof(f50031,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl40),
  inference(resolution,[],[f33087,f708])).
fof(f50028,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl40),
  inference(subsumption_resolution,[],[f49992,f758])).
fof(f49992,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl40),
  inference(resolution,[],[f33086,f704])).
fof(f50027,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl40),
  inference(subsumption_resolution,[],[f49991,f758])).
fof(f49991,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl40),
  inference(resolution,[],[f33086,f708])).
fof(f49988,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | $spl7),
  inference(subsumption_resolution,[],[f49973,f758])).
fof(f49973,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | $spl7),
  inference(resolution,[],[f15438,f704])).
fof(f49987,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | $spl7),
  inference(subsumption_resolution,[],[f49972,f758])).
fof(f49972,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | $spl7),
  inference(resolution,[],[f15438,f708])).
fof(f49970,plain,(
  $false | $spl98),
  inference(subsumption_resolution,[],[f49868,f6846])).
fof(f49868,plain,(
  ( ! [X304,X305] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X304),s('type.integer.int',X305))))) ) | $spl98),
  inference(resolution,[],[f6846,f797])).
fof(f49969,plain,(
  $false | $spl98),
  inference(subsumption_resolution,[],[f49867,f6846])).
fof(f49867,plain,(
  ( ! [X302,X303] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X302),s('type.integer.int',X303))))) ) | $spl98),
  inference(resolution,[],[f6846,f799])).
fof(f49968,plain,(
  $false | $spl98),
  inference(subsumption_resolution,[],[f49844,f6846])).
fof(f49844,plain,(
  ( ! [X275,X274] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X274),s('type.integer.int',X275))))) ) | $spl98),
  inference(resolution,[],[f6846,f799])).
fof(f49967,plain,(
  $false | $spl98),
  inference(subsumption_resolution,[],[f49843,f6846])).
fof(f49843,plain,(
  ( ! [X273,X272] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X272),s('type.integer.int',X273))))) ) | $spl98),
  inference(resolution,[],[f6846,f797])).
fof(f49966,plain,(
  $false | $spl98),
  inference(subsumption_resolution,[],[f49842,f6846])).
fof(f49842,plain,(
  ( ! [X271,X270] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X270),s('type.integer.int',X271))))) ) | $spl98),
  inference(resolution,[],[f6846,f734])).
fof(f49965,plain,(
  $false | $spl98),
  inference(subsumption_resolution,[],[f49841,f6846])).
fof(f49841,plain,(
  ( ! [X269,X268] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X268),s('type.integer.int',X269))))) ) | $spl98),
  inference(resolution,[],[f6846,f734])).
fof(f49964,plain,(
  $false | $spl98),
  inference(subsumption_resolution,[],[f49840,f6846])).
fof(f49840,plain,(
  ( ! [X267,X266] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X266),s('type.integer.int',X267))))) ) | $spl98),
  inference(resolution,[],[f6846,f679])).
fof(f49963,plain,(
  $false | $spl98),
  inference(subsumption_resolution,[],[f49839,f6846])).
fof(f49839,plain,(
  ( ! [X265,X264] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X264),s('type.integer.int',X265))))) ) | $spl98),
  inference(resolution,[],[f6846,f679])).
fof(f49962,plain,(
  $false | $spl98),
  inference(subsumption_resolution,[],[f49836,f6846])).
fof(f49836,plain,(
  ( ! [X262] : (~p(s(bool,X262))) ) | $spl98),
  inference(resolution,[],[f6846,f783])).
fof(f49961,plain,(
  $false | $spl98),
  inference(subsumption_resolution,[],[f49830,f6846])).
fof(f49830,plain,(
  ( ! [X256] : (~p(s(bool,X256))) ) | $spl98),
  inference(resolution,[],[f6846,f750])).
fof(f750,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f534])).
fof(f534,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f533])).
fof(f533,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f414])).
fof(f414,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f181])).
fof(f181,axiom,(
  ! [X15] : ((p(s(bool,X15)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X15)))),
  file('thm.integer.INT_ABS_LT|split|3','thm.bool.IMP_CLAUSES|split|4')).
fof(f49838,plain,(
  $false | $spl98),
  inference(resolution,[],[f6846,f788])).
fof(f788,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_ABS_LT|split|3',aHL_FALSITY)).
fof(f49853,plain,(
  $false | $spl98),
  inference(resolution,[],[f6846,f819])).
fof(f819,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(flattening,[],[f356])).
fof(f356,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ABS_LT|48|')).
fof(f49854,plain,(
  $false | ($spl29 | $spl98)),
  inference(resolution,[],[f6846,f1482])).
fof(f49859,plain,(
  $false | ($spl98 | $spl105)),
  inference(resolution,[],[f6846,f9008])).
fof(f49860,plain,(
  $false | ($spl7 | $spl98)),
  inference(resolution,[],[f6846,f1011])).
fof(f49861,plain,(
  $false | ($spl35 | $spl98)),
  inference(resolution,[],[f6846,f3490])).
fof(f3490,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl35),
  inference(resolution,[],[f1498,f810])).
fof(f49862,plain,(
  $false | ($spl3 | $spl4 | $spl98)),
  inference(resolution,[],[f6846,f4686])).
fof(f49863,plain,(
  $false | ($spl3 | $spl98)),
  inference(resolution,[],[f6846,f879])).
fof(f49864,plain,(
  $false | ($spl98 | $spl105)),
  inference(resolution,[],[f6846,f45206])).
fof(f49865,plain,(
  $false | ($spl13 | $spl98)),
  inference(resolution,[],[f6846,f1163])).
fof(f49870,plain,(
  $false | ($spl35 | $spl98)),
  inference(resolution,[],[f6846,f1498])).
fof(f49873,plain,(
  $false | ($spl4 | $spl98)),
  inference(resolution,[],[f6846,f886])).
fof(f886,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f880,f797])).
fof(f49874,plain,(
  $false | ($spl85 | $spl98)),
  inference(resolution,[],[f6846,f4682])).
fof(f49731,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f49714,f758])).
fof(f49714,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))) ) | $spl7),
  inference(resolution,[],[f15437,f704])).
fof(f49728,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f49713,f758])).
fof(f49713,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))) ) | $spl7),
  inference(resolution,[],[f15437,f708])).
fof(f49170,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl3 | $spl4)),
  inference(resolution,[],[f4686,f800])).
fof(f49167,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl3 | $spl4)),
  inference(resolution,[],[f4686,f809])).
fof(f1324,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl24),
  inference(cnf_transformation,[],[f1324_D])).
fof(f1324_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2795,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl64),
  inference(cnf_transformation,[],[f2795_D])).
fof(f2795_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f48420,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl3),
  inference(resolution,[],[f4551,f809])).
fof(f1322,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl23),
  inference(cnf_transformation,[],[f1322_D])).
fof(f1322_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f2288,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl56),
  inference(cnf_transformation,[],[f2288_D])).
fof(f2288_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f48021,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl38),
  inference(resolution,[],[f1501,f808])).
fof(f4044,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl82),
  inference(cnf_transformation,[],[f4044_D])).
fof(f4044_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f47961,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl29),
  inference(forward_demodulation,[],[f47833,f856])).
fof(f47833,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK14)))) | $spl29),
  inference(resolution,[],[f1482,f800])).
fof(f4037,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl79),
  inference(cnf_transformation,[],[f4037_D])).
fof(f4037_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f4039,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl80),
  inference(cnf_transformation,[],[f4039_D])).
fof(f4039_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f47830,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl29),
  inference(resolution,[],[f1482,f809])).
fof(f2638,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl59),
  inference(cnf_transformation,[],[f2638_D])).
fof(f2638_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f47631,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl7),
  inference(resolution,[],[f1011,f809])).
fof(f47488,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl13),
  inference(resolution,[],[f1163,f800])).
fof(f1318,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl19),
  inference(cnf_transformation,[],[f1318_D])).
fof(f1318_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f47487,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl13),
  inference(resolution,[],[f1163,f803])).
fof(f47377,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f47376,f856])).
fof(f47376,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f47344,f844])).
fof(f47344,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(superposition,[],[f876,f848])).
fof(f47343,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(superposition,[],[f876,f838])).
fof(f47342,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(superposition,[],[f876,f837])).
fof(f47341,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(superposition,[],[f876,f836])).
fof(f47340,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(superposition,[],[f876,f835])).
fof(f47375,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f47339,f845])).
fof(f47339,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl0),
  inference(superposition,[],[f876,f833])).
fof(f5791,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl96),
  inference(cnf_transformation,[],[f5791_D])).
fof(f5791_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f1482,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl29),
  inference(cnf_transformation,[],[f1482_D])).
fof(f1482_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f47369,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(forward_demodulation,[],[f47227,f845])).
fof(f47227,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK14)))) | $spl0),
  inference(resolution,[],[f876,f808])).
fof(f1162,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl11),
  inference(cnf_transformation,[],[f1162_D])).
fof(f1162_D,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1163,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl13),
  inference(cnf_transformation,[],[f1163_D])).
fof(f1163_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f46891,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(superposition,[],[f879,f834])).
fof(f834,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f563])).
fof(f1164,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl12),
  inference(cnf_transformation,[],[f1164_D])).
fof(f1164_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f46920,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl3),
  inference(forward_demodulation,[],[f46919,f844])).
fof(f46919,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | $spl3),
  inference(forward_demodulation,[],[f46787,f817])).
fof(f46787,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | $spl3),
  inference(resolution,[],[f879,f803])).
fof(f4682,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl85),
  inference(cnf_transformation,[],[f4682_D])).
fof(f4682_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f46785,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl3),
  inference(resolution,[],[f879,f809])).
fof(f5339,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | $spl94),
  inference(cnf_transformation,[],[f5339_D])).
fof(f5339_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f46783,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | p(s(bool,X1))) ) | $spl4),
  inference(subsumption_resolution,[],[f46775,f758])).
fof(f46775,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | p(s(bool,X1))) ) | $spl4),
  inference(resolution,[],[f1748,f704])).
fof(f1748,plain,(
  ( ! [X145] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X145)) ) | $spl4),
  inference(resolution,[],[f891,f706])).
fof(f891,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f880,f810])).
fof(f46782,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f46774,f758])).
fof(f46774,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | p(s(bool,X0))) ) | $spl4),
  inference(resolution,[],[f1748,f708])).
fof(f46773,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) ) | $spl4),
  inference(subsumption_resolution,[],[f46765,f758])).
fof(f46765,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) ) | $spl4),
  inference(resolution,[],[f1747,f704])).
fof(f1747,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) | $spl4),
  inference(resolution,[],[f891,f705])).
fof(f46772,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) ) | $spl4),
  inference(subsumption_resolution,[],[f46764,f758])).
fof(f46764,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) ) | $spl4),
  inference(resolution,[],[f1747,f708])).
fof(f1013,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl8),
  inference(cnf_transformation,[],[f1013_D])).
fof(f1013_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f46640,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f880,f808])).
fof(f880,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | $spl4),
  inference(cnf_transformation,[],[f880_D])).
fof(f880_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f12537,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl7),
  inference(resolution,[],[f1011,f800])).
fof(f5211,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl93),
  inference(cnf_transformation,[],[f5211_D])).
fof(f5211_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f12674,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl7),
  inference(forward_demodulation,[],[f12540,f845])).
fof(f12540,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl7),
  inference(resolution,[],[f1011,f803])).
fof(f1323,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl25),
  inference(cnf_transformation,[],[f1323_D])).
fof(f1323_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f13834,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl29),
  inference(resolution,[],[f1482,f808])).
fof(f10272,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl117),
  inference(cnf_transformation,[],[f10272_D])).
fof(f10272_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f1501,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl38),
  inference(cnf_transformation,[],[f1501_D])).
fof(f1501_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f40212,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(resolution,[],[f5638,f800])).
fof(f40362,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl0),
  inference(forward_demodulation,[],[f40215,f845])).
fof(f40215,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl0),
  inference(resolution,[],[f5638,f803])).
fof(f41544,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | $spl100),
  inference(resolution,[],[f9003,f803])).
fof(f46589,plain,(
  $false | ($spl16 | $spl26)),
  inference(subsumption_resolution,[],[f46459,f868])).
fof(f46459,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl16 | $spl26)),
  inference(resolution,[],[f40403,f799])).
fof(f40403,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl16 | $spl26)),
  inference(backward_demodulation,[],[f1316,f1480])).
fof(f46588,plain,(
  $false | ($spl16 | $spl26)),
  inference(subsumption_resolution,[],[f46458,f868])).
fof(f46458,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl16 | $spl26)),
  inference(resolution,[],[f40403,f797])).
fof(f46587,plain,(
  $false | ($spl16 | $spl26)),
  inference(subsumption_resolution,[],[f46457,f40403])).
fof(f46457,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl16 | $spl26)),
  inference(resolution,[],[f40403,f734])).
fof(f46586,plain,(
  $false | ($spl16 | $spl26)),
  inference(subsumption_resolution,[],[f46456,f40403])).
fof(f46456,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl16 | $spl26)),
  inference(resolution,[],[f40403,f734])).
fof(f46585,plain,(
  $false | ($spl16 | $spl26)),
  inference(subsumption_resolution,[],[f46455,f40403])).
fof(f46455,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl16 | $spl26)),
  inference(resolution,[],[f40403,f679])).
fof(f46584,plain,(
  $false | ($spl16 | $spl26)),
  inference(subsumption_resolution,[],[f46454,f40403])).
fof(f46454,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl16 | $spl26)),
  inference(resolution,[],[f40403,f679])).
fof(f46583,plain,(
  $false | ($spl16 | $spl26 | $spl105)),
  inference(subsumption_resolution,[],[f46452,f9008])).
fof(f46452,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | ($spl16 | $spl26)),
  inference(resolution,[],[f40403,f807])).
fof(f45209,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | $spl105),
  inference(resolution,[],[f9008,f809])).
fof(f45149,plain,(
  $false | ($spl35 | $spl86)),
  inference(subsumption_resolution,[],[f5060,f3490])).
fof(f44644,plain,(
  $false | ($spl16 | $spl100 | $spl119)),
  inference(subsumption_resolution,[],[f40528,f9003])).
fof(f40528,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl16 | $spl119)),
  inference(backward_demodulation,[],[f1316,f10650])).
fof(f19447,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl180),
  inference(cnf_transformation,[],[f19447_D])).
fof(f19447_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f41874,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f41867,f758])).
fof(f41867,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | ($spl2 | $spl16)),
  inference(resolution,[],[f40565,f704])).
fof(f40565,plain,(
  ( ! [X145] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))),X145)) ) | ($spl2 | $spl16)),
  inference(backward_demodulation,[],[f1316,f16740])).
fof(f41873,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f41866,f758])).
fof(f41866,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | ($spl2 | $spl16)),
  inference(resolution,[],[f40565,f708])).
fof(f41865,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f41858,f758])).
fof(f41858,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | ($spl2 | $spl16)),
  inference(resolution,[],[f40564,f704])).
fof(f40564,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) ) | ($spl2 | $spl16)),
  inference(backward_demodulation,[],[f1316,f16739])).
fof(f41864,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f41857,f758])).
fof(f41857,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | ($spl2 | $spl16)),
  inference(resolution,[],[f40564,f708])).
fof(f41847,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl16 | $spl29)),
  inference(forward_demodulation,[],[f41720,f856])).
fof(f41720,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK13)))) | ($spl16 | $spl29)),
  inference(resolution,[],[f10947,f800])).
fof(f10947,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ($spl16 | $spl29)),
  inference(backward_demodulation,[],[f1316,f1482])).
fof(f41540,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl324),
  inference(cnf_transformation,[],[f41540_D])).
fof(f41540_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f41695,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f41684,f758])).
fof(f41684,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | ($spl2 | $spl16)),
  inference(resolution,[],[f40448,f704])).
fof(f40448,plain,(
  ( ! [X145] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X145)) ) | ($spl2 | $spl16)),
  inference(backward_demodulation,[],[f1316,f3039])).
fof(f41694,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f41683,f758])).
fof(f41683,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | ($spl2 | $spl16)),
  inference(resolution,[],[f40448,f708])).
fof(f41553,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | $spl100),
  inference(resolution,[],[f9003,f808])).
fof(f41541,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f41528,f758])).
fof(f41528,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))) ) | ($spl2 | $spl16)),
  inference(resolution,[],[f40447,f704])).
fof(f40447,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) ) | ($spl2 | $spl16)),
  inference(backward_demodulation,[],[f1316,f3038])).
fof(f41538,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f41527,f758])).
fof(f41527,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))) ) | ($spl2 | $spl16)),
  inference(resolution,[],[f40447,f708])).
fof(f41383,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl320),
  inference(cnf_transformation,[],[f41383_D])).
fof(f41383_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f41385,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl323),
  inference(cnf_transformation,[],[f41385_D])).
fof(f41385_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) <=> ~$spl323),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl323])])).
fof(f41381,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl0 | $spl16)),
  inference(forward_demodulation,[],[f41241,f845])).
fof(f41241,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK13)))) | ($spl0 | $spl16)),
  inference(resolution,[],[f10945,f808])).
fof(f10945,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl0 | $spl16)),
  inference(backward_demodulation,[],[f1316,f876])).
fof(f40891,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | ($spl5 | $spl16)),
  inference(resolution,[],[f11259,f800])).
fof(f11259,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ($spl5 | $spl16)),
  inference(forward_demodulation,[],[f881,f1316])).
fof(f9004,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl101),
  inference(cnf_transformation,[],[f9004_D])).
fof(f9004_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f40890,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl5 | $spl16)),
  inference(resolution,[],[f11259,f803])).
fof(f40751,plain,(
  $false | ($spl2 | $spl16 | $spl101)),
  inference(subsumption_resolution,[],[f40400,f9004])).
fof(f40400,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl2 | $spl16)),
  inference(backward_demodulation,[],[f1316,f1156])).
fof(f40750,plain,(
  $false | ($spl2 | $spl16 | $spl101)),
  inference(subsumption_resolution,[],[f40384,f9004])).
fof(f40384,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl2 | $spl16)),
  inference(backward_demodulation,[],[f1316,f1019])).
fof(f1019,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(resolution,[],[f878,f805])).
fof(f40380,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl319),
  inference(cnf_transformation,[],[f40380_D])).
fof(f40380_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) <=> ~$spl319),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl319])])).
fof(f40213,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl0),
  inference(resolution,[],[f5638,f809])).
fof(f5340,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | $spl95),
  inference(cnf_transformation,[],[f5340_D])).
fof(f5340_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f1316,plain,(
  s('type.integer.int',sK13) = s('type.integer.int',sK14) | $spl16),
  inference(cnf_transformation,[],[f1316_D])).
fof(f1316_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int',sK14) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f14964,plain,(
  s('type.integer.int',sK13) = s('type.integer.int',sK14) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | $spl5),
  inference(subsumption_resolution,[],[f14713,f750])).
fof(f14713,plain,(
  p(s(bool,'F_0')) | s('type.integer.int',sK13) = s('type.integer.int',sK14) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | $spl5),
  inference(superposition,[],[f809,f11374])).
fof(f11374,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | $spl5),
  inference(resolution,[],[f881,f752])).
fof(f11281,plain,(
  s('type.integer.int',sK13) = s('type.integer.int',sK14) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | $spl5),
  inference(resolution,[],[f881,f809])).
fof(f11759,plain,(
  s('type.integer.int',sK13) = s('type.integer.int',sK14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl5),
  inference(resolution,[],[f5216,f809])).
fof(f5216,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl5),
  inference(resolution,[],[f881,f798])).
fof(f4043,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl83),
  inference(cnf_transformation,[],[f4043_D])).
fof(f4043_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f13045,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(forward_demodulation,[],[f12911,f845])).
fof(f12911,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(resolution,[],[f1020,f803])).
fof(f40377,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl316),
  inference(cnf_transformation,[],[f40377_D])).
fof(f40377_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f15291,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl20),
  inference(forward_demodulation,[],[f15162,f845])).
fof(f15162,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl20),
  inference(resolution,[],[f1320,f808])).
fof(f40359,plain,(
  $false | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f40210,f1164])).
fof(f40210,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl0),
  inference(resolution,[],[f5638,f802])).
fof(f40033,plain,(
  ( ! [X3] : (sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3)) ) | $spl314),
  inference(cnf_transformation,[],[f40033_D])).
fof(f40033_D,plain,(
  ( ! [X3] : (sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3)) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f20503,plain,(
  ( ! [X14] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X14)))) ) | $spl192),
  inference(cnf_transformation,[],[f20503_D])).
fof(f20503_D,plain,(
  ( ! [X14] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X14)))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f40035,plain,(
  ( ! [X37,X36] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X36))) | sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X37)) )),
  inference(forward_demodulation,[],[f40026,f856])).
fof(f40026,plain,(
  ( ! [X37,X36] : (sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X37) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X36)))))) )),
  inference(superposition,[],[f21090,f850])).
fof(f850,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f566])).
fof(f566,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ABS_EQ|48|')).
fof(f21090,plain,(
  ( ! [X225,X226] : (sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X225)))),X226)) )),
  inference(resolution,[],[f854,f706])).
fof(f28837,plain,(
  ( ! [X404] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X404))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl246),
  inference(cnf_transformation,[],[f28837_D])).
fof(f28837_D,plain,(
  ( ! [X404] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X404))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f40034,plain,(
  ( ! [X35,X34] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X34))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X35)) )),
  inference(forward_demodulation,[],[f40025,f856])).
fof(f40025,plain,(
  ( ! [X35,X34] : (sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X35) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X34))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f21090,f848])).
fof(f20500,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X13)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X13))))) ) | $spl190),
  inference(cnf_transformation,[],[f20500_D])).
fof(f20500_D,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X13)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X13))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f40024,plain,(
  ( ! [X33,X32] : (sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X33) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X32)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X32))))) )),
  inference(superposition,[],[f21090,f838])).
fof(f39579,plain,(
  ( ! [X30] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X30)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X30))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X30))))))) ) | $spl312),
  inference(cnf_transformation,[],[f39579_D])).
fof(f39579_D,plain,(
  ( ! [X30] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X30)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X30))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X30))))))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f40023,plain,(
  ( ! [X30,X31] : (sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X31) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X30)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X30))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X30))))))) )),
  inference(superposition,[],[f21090,f837])).
fof(f39578,plain,(
  ( ! [X28] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X28))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X28)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X28))))) ) | $spl310),
  inference(cnf_transformation,[],[f39578_D])).
fof(f39578_D,plain,(
  ( ! [X28] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X28))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X28)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X28))))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f40022,plain,(
  ( ! [X28,X29] : (sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X29) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X28))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X28)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X28))))) )),
  inference(superposition,[],[f21090,f836])).
fof(f39577,plain,(
  ( ! [X26] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X26))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X26)))))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X26))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X26))))))) ) | $spl308),
  inference(cnf_transformation,[],[f39577_D])).
fof(f39577_D,plain,(
  ( ! [X26] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X26))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X26)))))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X26))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X26))))))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f40021,plain,(
  ( ! [X26,X27] : (sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X27) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X26))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X26)))))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X26))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X26))))))) )),
  inference(superposition,[],[f21090,f835])).
fof(f20498,plain,(
  ( ! [X12] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X12)))))) ) | $spl188),
  inference(cnf_transformation,[],[f20498_D])).
fof(f20498_D,plain,(
  ( ! [X12] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X12)))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f40020,plain,(
  ( ! [X24,X25] : (sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X24))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X24)))))) )),
  inference(superposition,[],[f21090,f833])).
fof(f39576,plain,(
  ( ! [X22] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X22)))))) ) | $spl306),
  inference(cnf_transformation,[],[f39576_D])).
fof(f39576_D,plain,(
  ( ! [X22] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X22)))))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f40019,plain,(
  ( ! [X23,X22] : (sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X22)))))) )),
  inference(superposition,[],[f21090,f831])).
fof(f831,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f562])).
fof(f562,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f368])).
fof(f368,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) <=> s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_NEG_EQ|48|')).
fof(f20493,plain,(
  ( ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl184),
  inference(cnf_transformation,[],[f20493_D])).
fof(f20493_D,plain,(
  ( ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f40011,plain,(
  ( ! [X2,X3] : (sP12('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3) | s('type.integer.int',X2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f21090,f851])).
fof(f851,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f566])).
fof(f39908,plain,(
  $false | $spl192),
  inference(trivial_inequality_removal,[],[f39831])).
fof(f39831,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl192),
  inference(superposition,[],[f39804,f846])).
fof(f39804,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl192),
  inference(trivial_inequality_removal,[],[f39780])).
fof(f39780,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl192),
  inference(superposition,[],[f20503,f851])).
fof(f39897,plain,(
  $false | $spl192),
  inference(trivial_inequality_removal,[],[f39857])).
fof(f39857,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl192),
  inference(superposition,[],[f39804,f843])).
fof(f843,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f376])).
fof(f376,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))),
  inference(rectify,[],[f153])).
fof(f153,axiom,(
  ! [X9] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int',X9)))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ADD_LINV')).
fof(f39896,plain,(
  $false | $spl192),
  inference(trivial_inequality_removal,[],[f39862])).
fof(f39862,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl192),
  inference(superposition,[],[f39804,f842])).
fof(f842,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f375])).
fof(f375,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f227])).
fof(f227,axiom,(
  ! [X9] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ADD_RINV')).
fof(f39895,plain,(
  $false | $spl192),
  inference(equality_resolution,[],[f39804])).
fof(f39681,plain,(
  $false | $spl184),
  inference(trivial_inequality_removal,[],[f39604])).
fof(f39604,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl184),
  inference(superposition,[],[f20493,f846])).
fof(f39670,plain,(
  $false | $spl184),
  inference(trivial_inequality_removal,[],[f39630])).
fof(f39630,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl184),
  inference(superposition,[],[f20493,f843])).
fof(f39669,plain,(
  $false | $spl184),
  inference(trivial_inequality_removal,[],[f39635])).
fof(f39635,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl184),
  inference(superposition,[],[f20493,f842])).
fof(f39668,plain,(
  $false | $spl184),
  inference(equality_resolution,[],[f20493])).
fof(f39575,plain,(
  ( ! [X3] : (sP12(X3,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl304),
  inference(cnf_transformation,[],[f39575_D])).
fof(f39575_D,plain,(
  ( ! [X3] : (sP12(X3,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f39581,plain,(
  ( ! [X37,X36] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X36))) | sP12(X37,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f39568,f856])).
fof(f39568,plain,(
  ( ! [X37,X36] : (sP12(X37,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X36)))))) )),
  inference(superposition,[],[f21089,f850])).
fof(f39580,plain,(
  ( ! [X35,X34] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X34))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP12(X35,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f39567,f856])).
fof(f39567,plain,(
  ( ! [X35,X34] : (sP12(X35,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X34))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f21089,f848])).
fof(f39566,plain,(
  ( ! [X33,X32] : (sP12(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X32)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X32))))) )),
  inference(superposition,[],[f21089,f838])).
fof(f39565,plain,(
  ( ! [X30,X31] : (sP12(X31,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X30)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X30))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X30))))))) )),
  inference(superposition,[],[f21089,f837])).
fof(f39564,plain,(
  ( ! [X28,X29] : (sP12(X29,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X28))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X28)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X28))))) )),
  inference(superposition,[],[f21089,f836])).
fof(f39563,plain,(
  ( ! [X26,X27] : (sP12(X27,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X26))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X26)))))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X26))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X26))))))) )),
  inference(superposition,[],[f21089,f835])).
fof(f39562,plain,(
  ( ! [X24,X25] : (sP12(X25,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X24))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X24)))))) )),
  inference(superposition,[],[f21089,f833])).
fof(f39561,plain,(
  ( ! [X23,X22] : (sP12(X23,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X22)))))) )),
  inference(superposition,[],[f21089,f831])).
fof(f39553,plain,(
  ( ! [X2,X3] : (sP12(X3,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f21089,f851])).
fof(f37621,plain,(
  ( ! [X504] : (s('type.integer.int',X504) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X504)))))) | s('type.integer.int',X504) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26(X504))))) ) | $spl302),
  inference(cnf_transformation,[],[f37621_D])).
fof(f37621_D,plain,(
  ( ! [X504] : (s('type.integer.int',X504) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X504)))))) | s('type.integer.int',X504) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26(X504))))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f26909,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl230),
  inference(cnf_transformation,[],[f26909_D])).
fof(f26909_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f37620,plain,(
  ( ! [X504] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X504) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X504)))))) | s('type.integer.int',X504) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26(X504))))) )),
  inference(subsumption_resolution,[],[f36883,f868])).
fof(f36883,plain,(
  ( ! [X504] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X504),s('type.integer.int',X504)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X504) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X504)))))) | s('type.integer.int',X504) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26(X504))))) )),
  inference(superposition,[],[f864,f835])).
fof(f32754,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl262),
  inference(cnf_transformation,[],[f32754_D])).
fof(f32754_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f35815,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | $spl286),
  inference(cnf_transformation,[],[f35815_D])).
fof(f35815_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f24079,plain,(
  ( ! [X463,X464] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X463),s('type.integer.int',X464))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X463),s('type.integer.int',X464)))) ) | $spl224),
  inference(cnf_transformation,[],[f24079_D])).
fof(f24079_D,plain,(
  ( ! [X463,X464] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X463),s('type.integer.int',X464))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X463),s('type.integer.int',X464)))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f37535,plain,(
  ( ! [X563,X564] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X563),s('type.integer.int',X564))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X563),s('type.integer.int',X564))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(forward_demodulation,[],[f36578,f867])).
fof(f867,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f244])).
fof(f244,axiom,(
  ! [X9] : s('type.integer.int',X9) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ADD_RID')).
fof(f36578,plain,(
  ( ! [X563,X564] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X563),s('type.integer.int',X564))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X563),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X564))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f812,f835])).
fof(f812,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int',X0)))),
  inference(rectify,[],[f255])).
fof(f255,axiom,(
  ! [X11,X9] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int',X11)))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_LT_LE|49|')).
fof(f24061,plain,(
  ( ! [X406] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X406))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X406)))) ) | $spl210),
  inference(cnf_transformation,[],[f24061_D])).
fof(f24061_D,plain,(
  ( ! [X406] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X406))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X406)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f37534,plain,(
  ( ! [X562] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X562))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X562))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(forward_demodulation,[],[f36577,f846])).
fof(f36577,plain,(
  ( ! [X562] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X562))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X562))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f840,f835])).
fof(f840,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f373])).
fof(f373,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int',X0)))),
  inference(rectify,[],[f257])).
fof(f257,axiom,(
  ! [X9] : s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int',X9)))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_NEG_MINUS|49|')).
fof(f37528,plain,(
  ( ! [X449] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X449) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X449))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X449)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X449))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X449))))))))) ) | $spl300),
  inference(cnf_transformation,[],[f37528_D])).
fof(f37528_D,plain,(
  ( ! [X449] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X449) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X449))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X449)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X449))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X449))))))))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f24070,plain,(
  ( ! [X414] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X414))))) ) | $spl218),
  inference(cnf_transformation,[],[f24070_D])).
fof(f24070_D,plain,(
  ( ! [X414] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X414))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f37530,plain,(
  ( ! [X455,X456] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X456)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X455) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X455))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X455)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X455))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X455))))))) )),
  inference(forward_demodulation,[],[f37529,f815])).
fof(f815,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f352])).
fof(f352,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f156])).
fof(f156,axiom,(
  ! [X10,X12] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_LT_CALCULATE|split|0')).
fof(f37529,plain,(
  ( ! [X455,X456] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X456)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X455) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X455))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X455)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X455))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X455))))))) )),
  inference(forward_demodulation,[],[f36492,f821])).
fof(f36492,plain,(
  ( ! [X455,X456] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X456))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X455) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X455))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X455)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X455))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X455))))))) )),
  inference(superposition,[],[f723,f835])).
fof(f723,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f516])).
fof(f516,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) & ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))))),
  inference(flattening,[],[f515])).
fof(f515,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) <=> (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X10,X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X10) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_LT_CALCULATE|split|2')).
fof(f37527,plain,(
  ( ! [X449,X450] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X450)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X449) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X449))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X449)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X449))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X449))))))) )),
  inference(forward_demodulation,[],[f37526,f815])).
fof(f37526,plain,(
  ( ! [X449,X450] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X450)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X449) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X449))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X449)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X449))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X449))))))) )),
  inference(forward_demodulation,[],[f36489,f846])).
fof(f36489,plain,(
  ( ! [X449,X450] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X450)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X449) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X449))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X449)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X449))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X449))))))) )),
  inference(superposition,[],[f722,f835])).
fof(f722,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f516])).
fof(f37523,plain,(
  ( ! [X445] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X445) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X445)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X445))))))))) ) | $spl298),
  inference(cnf_transformation,[],[f37523_D])).
fof(f37523_D,plain,(
  ( ! [X445] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X445) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X445)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X445))))))))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f24052,plain,(
  ( ! [X401] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X401)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X401)) ) | $spl202),
  inference(cnf_transformation,[],[f24052_D])).
fof(f24052_D,plain,(
  ( ! [X401] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X401)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X401)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f37525,plain,(
  ( ! [X447,X448] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X448)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X447) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X448) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X447))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X447)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X447))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X447))))))) )),
  inference(forward_demodulation,[],[f37524,f815])).
fof(f37524,plain,(
  ( ! [X447,X448] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X448)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X447) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X448) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X447))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X447)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X447))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X447))))))) )),
  inference(forward_demodulation,[],[f36488,f821])).
fof(f36488,plain,(
  ( ! [X447,X448] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X448))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X447) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X448) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X447))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X447)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X447))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X447))))))) )),
  inference(superposition,[],[f721,f835])).
fof(f721,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f516])).
fof(f37522,plain,(
  ( ! [X445,X446] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X446)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X446) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X445) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X445)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X445))))))) )),
  inference(forward_demodulation,[],[f37521,f815])).
fof(f37521,plain,(
  ( ! [X445,X446] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X446)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X446) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X445) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X445)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X445))))))) )),
  inference(forward_demodulation,[],[f36487,f846])).
fof(f36487,plain,(
  ( ! [X445,X446] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X446)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X446) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X445) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X445)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X445))))))) )),
  inference(superposition,[],[f721,f835])).
fof(f32735,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | $spl254),
  inference(cnf_transformation,[],[f32735_D])).
fof(f32735_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f35796,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) | $spl278),
  inference(cnf_transformation,[],[f35796_D])).
fof(f35796_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f36484,plain,(
  ( ! [X441] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X441))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X441))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) )),
  inference(superposition,[],[f840,f835])).
fof(f37512,plain,(
  ( ! [X429] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))))))) ) | $spl292),
  inference(cnf_transformation,[],[f37512_D])).
fof(f37512_D,plain,(
  ( ! [X429] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))))))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f24057,plain,(
  ( ! [X405] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X405)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X405)) ) | $spl204),
  inference(cnf_transformation,[],[f24057_D])).
fof(f24057_D,plain,(
  ( ! [X405] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X405)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X405)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f37518,plain,(
  ( ! [X440,X439] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X440)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X440) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X439))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X439)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X439))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X439))))))))) )),
  inference(forward_demodulation,[],[f36483,f815])).
fof(f36483,plain,(
  ( ! [X440,X439] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X440)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X440) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X439))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X439)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X439))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X439))))))))) )),
  inference(superposition,[],[f723,f835])).
fof(f37517,plain,(
  ( ! [X437] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X437) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437))))))))))) ) | $spl296),
  inference(cnf_transformation,[],[f37517_D])).
fof(f37517_D,plain,(
  ( ! [X437] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X437) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437))))))))))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f37516,plain,(
  ( ! [X438,X437] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X438)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X437) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437))))))))) )),
  inference(forward_demodulation,[],[f36482,f815])).
fof(f36482,plain,(
  ( ! [X438,X437] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X438)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X437) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437))))))))) )),
  inference(superposition,[],[f722,f835])).
fof(f37515,plain,(
  ( ! [X435] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X435) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435))))))))))) ) | $spl294),
  inference(cnf_transformation,[],[f37515_D])).
fof(f37515_D,plain,(
  ( ! [X435] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X435) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435))))))))))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f37514,plain,(
  ( ! [X436,X435] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X436)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X436) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X435) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435))))))))) )),
  inference(forward_demodulation,[],[f36481,f815])).
fof(f36481,plain,(
  ( ! [X436,X435] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X436) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X435) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435))))))))) )),
  inference(superposition,[],[f721,f835])).
fof(f24046,plain,(
  ( ! [X395] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X395))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl198),
  inference(cnf_transformation,[],[f24046_D])).
fof(f24046_D,plain,(
  ( ! [X395] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X395))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f36478,plain,(
  ( ! [X430,X429] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X430))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))))))) )),
  inference(superposition,[],[f818,f835])).
fof(f818,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f355])).
fof(f355,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10,X12] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_LT_CALCULATE|split|3')).
fof(f35971,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | $spl119),
  inference(resolution,[],[f34541,f800])).
fof(f34541,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl119),
  inference(resolution,[],[f10650,f804])).
fof(f35886,plain,(
  ( ! [X501] : (s('type.integer.int',X501) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X501)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26(X501))) ) | $spl288),
  inference(cnf_transformation,[],[f35886_D])).
fof(f35886_D,plain,(
  ( ! [X501] : (s('type.integer.int',X501) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X501)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26(X501))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f35885,plain,(
  ( ! [X501] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X501) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X501)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26(X501))) )),
  inference(subsumption_resolution,[],[f35407,f868])).
fof(f35407,plain,(
  ( ! [X501] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X501),s('type.integer.int',X501)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X501) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X501)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26(X501))) )),
  inference(superposition,[],[f864,f836])).
fof(f35813,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl285),
  inference(cnf_transformation,[],[f35813_D])).
fof(f35813_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl285),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl285])])).
fof(f35816,plain,(
  ( ! [X561,X560] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X560),s('type.integer.int',X561))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X560),s('type.integer.int',X561))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) )),
  inference(forward_demodulation,[],[f35104,f867])).
fof(f35104,plain,(
  ( ! [X561,X560] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X560),s('type.integer.int',X561))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X560),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X561))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) )),
  inference(superposition,[],[f812,f836])).
fof(f35811,plain,(
  ( ! [X559] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X559))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X559))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) )),
  inference(forward_demodulation,[],[f35103,f846])).
fof(f35103,plain,(
  ( ! [X559] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X559))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X559))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) )),
  inference(superposition,[],[f840,f836])).
fof(f35805,plain,(
  ( ! [X447] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X447) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X447)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X447))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X447))))))))) ) | $spl282),
  inference(cnf_transformation,[],[f35805_D])).
fof(f35805_D,plain,(
  ( ! [X447] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X447) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X447)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X447))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X447))))))))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f35807,plain,(
  ( ! [X453,X454] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X454)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X453) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X453))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X453)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X453))))) )),
  inference(forward_demodulation,[],[f35806,f815])).
fof(f35806,plain,(
  ( ! [X453,X454] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X454)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X453) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X453))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X453)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X453))))) )),
  inference(forward_demodulation,[],[f35019,f821])).
fof(f35019,plain,(
  ( ! [X453,X454] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X454))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X453) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X453))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X453)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X453))))) )),
  inference(superposition,[],[f723,f836])).
fof(f35804,plain,(
  ( ! [X447,X448] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X448)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X447) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X447))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X447)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X447))))) )),
  inference(forward_demodulation,[],[f35803,f815])).
fof(f35803,plain,(
  ( ! [X447,X448] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X448)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X447) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X447))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X447)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X447))))) )),
  inference(forward_demodulation,[],[f35016,f846])).
fof(f35016,plain,(
  ( ! [X447,X448] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X448)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X447) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X447))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X447)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X447))))) )),
  inference(superposition,[],[f722,f836])).
fof(f35800,plain,(
  ( ! [X443] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X443) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X443)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X443))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X443))))))))) ) | $spl280),
  inference(cnf_transformation,[],[f35800_D])).
fof(f35800_D,plain,(
  ( ! [X443] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X443) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X443)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X443))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X443))))))))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f35802,plain,(
  ( ! [X445,X446] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X446)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X445) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X446) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X445)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X445))))) )),
  inference(forward_demodulation,[],[f35801,f815])).
fof(f35801,plain,(
  ( ! [X445,X446] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X446)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X445) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X446) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X445)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X445))))) )),
  inference(forward_demodulation,[],[f35015,f821])).
fof(f35015,plain,(
  ( ! [X445,X446] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X446))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X445) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X446) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X445))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X445)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X445))))) )),
  inference(superposition,[],[f721,f836])).
fof(f35799,plain,(
  ( ! [X443,X444] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X444)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X444) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X443) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X443))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X443)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X443))))) )),
  inference(forward_demodulation,[],[f35798,f815])).
fof(f35798,plain,(
  ( ! [X443,X444] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X444)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X444) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X443) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X443))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X443)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X443))))) )),
  inference(forward_demodulation,[],[f35014,f846])).
fof(f35014,plain,(
  ( ! [X443,X444] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X444)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X444) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X443) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X443))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X443)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X443))))) )),
  inference(superposition,[],[f721,f836])).
fof(f35794,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl277),
  inference(cnf_transformation,[],[f35794_D])).
fof(f35794_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ~$spl277),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl277])])).
fof(f35012,plain,(
  ( ! [X440] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X440))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X440))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f840,f836])).
fof(f35786,plain,(
  ( ! [X428] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X428))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X428)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X428))))))) ) | $spl270),
  inference(cnf_transformation,[],[f35786_D])).
fof(f35786_D,plain,(
  ( ! [X428] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X428))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X428)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X428))))))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f35792,plain,(
  ( ! [X438,X439] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X439)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X439) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X438))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X438)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X438))))))) )),
  inference(forward_demodulation,[],[f35011,f815])).
fof(f35011,plain,(
  ( ! [X438,X439] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X439)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X439) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X438))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X438)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X438))))))) )),
  inference(superposition,[],[f723,f836])).
fof(f35791,plain,(
  ( ! [X436] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X436) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436))))))))))) ) | $spl274),
  inference(cnf_transformation,[],[f35791_D])).
fof(f35791_D,plain,(
  ( ! [X436] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X436) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436))))))))))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f35790,plain,(
  ( ! [X436,X437] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X437)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X436) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436))))))) )),
  inference(forward_demodulation,[],[f35010,f815])).
fof(f35010,plain,(
  ( ! [X436,X437] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X437)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X436) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X436))))))) )),
  inference(superposition,[],[f722,f836])).
fof(f35789,plain,(
  ( ! [X434] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X434) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434))))))))))) ) | $spl272),
  inference(cnf_transformation,[],[f35789_D])).
fof(f35789_D,plain,(
  ( ! [X434] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X434) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434))))))))))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f35788,plain,(
  ( ! [X434,X435] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X435)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X435) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X434) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434))))))) )),
  inference(forward_demodulation,[],[f35009,f815])).
fof(f35009,plain,(
  ( ! [X434,X435] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X435)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X435) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X434) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434))))))) )),
  inference(superposition,[],[f721,f836])).
fof(f35006,plain,(
  ( ! [X428,X429] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X428))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X428)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X428))))))) )),
  inference(superposition,[],[f818,f836])).
fof(f34674,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl269),
  inference(cnf_transformation,[],[f34674_D])).
fof(f34674_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl269),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl269])])).
fof(f10797,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | $spl120),
  inference(cnf_transformation,[],[f10797_D])).
fof(f10797_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f34546,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl119),
  inference(resolution,[],[f10650,f800])).
fof(f32822,plain,(
  ( ! [X455] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27(X455)) | s('type.integer.int',X455) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26(X455))))) ) | $spl266),
  inference(cnf_transformation,[],[f32822_D])).
fof(f32822_D,plain,(
  ( ! [X455] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27(X455)) | s('type.integer.int',X455) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26(X455))))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f32821,plain,(
  ( ! [X455] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27(X455)) | s('type.integer.int',X455) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26(X455))))) )),
  inference(subsumption_resolution,[],[f32407,f868])).
fof(f32407,plain,(
  ( ! [X455] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X455),s('type.integer.int',X455)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27(X455)) | s('type.integer.int',X455) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26(X455))))) )),
  inference(superposition,[],[f864,f837])).
fof(f32756,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl265),
  inference(cnf_transformation,[],[f32756_D])).
fof(f32756_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl265),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl265])])).
fof(f32757,plain,(
  ( ! [X530,X529] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X529),s('type.integer.int',X530))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X529),s('type.integer.int',X530))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(forward_demodulation,[],[f32125,f867])).
fof(f32125,plain,(
  ( ! [X530,X529] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X529),s('type.integer.int',X530))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X529),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X530))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f812,f837])).
fof(f32752,plain,(
  ( ! [X528] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X528))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X528))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(forward_demodulation,[],[f32124,f846])).
fof(f32124,plain,(
  ( ! [X528] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X528))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X528))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f840,f837])).
fof(f32746,plain,(
  ( ! [X442] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X442) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X442))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X442)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X442))))) ) | $spl260),
  inference(cnf_transformation,[],[f32746_D])).
fof(f32746_D,plain,(
  ( ! [X442] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X442) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X442))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X442)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X442))))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f32748,plain,(
  ( ! [X449,X448] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X449)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X448) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X448)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X448))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X448))))))) )),
  inference(forward_demodulation,[],[f32747,f815])).
fof(f32747,plain,(
  ( ! [X449,X448] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X449)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X448) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X448)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X448))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X448))))))) )),
  inference(forward_demodulation,[],[f32058,f821])).
fof(f32058,plain,(
  ( ! [X449,X448] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X449))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X448) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X448)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X448))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X448))))))) )),
  inference(superposition,[],[f723,f837])).
fof(f32745,plain,(
  ( ! [X443,X442] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X443)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X442) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X442)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X442))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X442))))))) )),
  inference(forward_demodulation,[],[f32744,f815])).
fof(f32744,plain,(
  ( ! [X443,X442] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X443)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X442) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X442)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X442))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X442))))))) )),
  inference(forward_demodulation,[],[f32055,f846])).
fof(f32055,plain,(
  ( ! [X443,X442] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X443)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X442) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X442)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X442))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X442))))))) )),
  inference(superposition,[],[f722,f837])).
fof(f32741,plain,(
  ( ! [X438] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X438) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X438))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X438)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X438))))) ) | $spl258),
  inference(cnf_transformation,[],[f32741_D])).
fof(f32741_D,plain,(
  ( ! [X438] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X438) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X438))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X438)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X438))))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f32743,plain,(
  ( ! [X441,X440] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X441)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X440) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X441) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X440)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X440))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X440))))))) )),
  inference(forward_demodulation,[],[f32742,f815])).
fof(f32742,plain,(
  ( ! [X441,X440] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X441)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X440) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X441) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X440)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X440))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X440))))))) )),
  inference(forward_demodulation,[],[f32054,f821])).
fof(f32054,plain,(
  ( ! [X441,X440] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X441))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X440) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X441) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X440)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X440))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X440))))))) )),
  inference(superposition,[],[f721,f837])).
fof(f32740,plain,(
  ( ! [X438,X439] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X439)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X439) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X438) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X438)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X438))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X438))))))) )),
  inference(forward_demodulation,[],[f32739,f815])).
fof(f32739,plain,(
  ( ! [X438,X439] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X439)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X439) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X438) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X438)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X438))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X438))))))) )),
  inference(forward_demodulation,[],[f32053,f846])).
fof(f32053,plain,(
  ( ! [X438,X439] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X439)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X439) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X438) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num',X438)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X438))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_of_num_1'(s('type.num.num',X438))))))) )),
  inference(superposition,[],[f721,f837])).
fof(f32737,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl257),
  inference(cnf_transformation,[],[f32737_D])).
fof(f32737_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f32051,plain,(
  ( ! [X435] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X435))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X435))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) )),
  inference(superposition,[],[f840,f837])).
fof(f32727,plain,(
  ( ! [X423] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X423)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X423))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X423))))))))) ) | $spl248),
  inference(cnf_transformation,[],[f32727_D])).
fof(f32727_D,plain,(
  ( ! [X423] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X423)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X423))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X423))))))))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f32733,plain,(
  ( ! [X434,X433] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X434)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X434) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X433)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X433))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X433))))))))) )),
  inference(forward_demodulation,[],[f32050,f815])).
fof(f32050,plain,(
  ( ! [X434,X433] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X434)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X434) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X433)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X433))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X433))))))))) )),
  inference(superposition,[],[f723,f837])).
fof(f32732,plain,(
  ( ! [X431] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X431) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X431))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X431)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X431))))))) ) | $spl252),
  inference(cnf_transformation,[],[f32732_D])).
fof(f32732_D,plain,(
  ( ! [X431] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X431) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X431))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X431)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X431))))))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f32731,plain,(
  ( ! [X432,X431] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X432)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X431) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X431)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X431))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X431))))))))) )),
  inference(forward_demodulation,[],[f32049,f815])).
fof(f32049,plain,(
  ( ! [X432,X431] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X432)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X431) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X431)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X431))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X431))))))))) )),
  inference(superposition,[],[f722,f837])).
fof(f32730,plain,(
  ( ! [X429] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X429) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))))) ) | $spl250),
  inference(cnf_transformation,[],[f32730_D])).
fof(f32730_D,plain,(
  ( ! [X429] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X429) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f32729,plain,(
  ( ! [X430,X429] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X430)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X430) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X429) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))))))) )),
  inference(forward_demodulation,[],[f32048,f815])).
fof(f32048,plain,(
  ( ! [X430,X429] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X430)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X430) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X429) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X429))))))))) )),
  inference(superposition,[],[f721,f837])).
fof(f32045,plain,(
  ( ! [X424,X423] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X424))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X423)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X423))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X423))))))))) )),
  inference(superposition,[],[f818,f837])).
fof(f32040,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl30),
  inference(superposition,[],[f1484,f837])).
fof(f31579,plain,(
  $false | ($spl2 | $spl115)),
  inference(subsumption_resolution,[],[f10268,f6490])).
fof(f6490,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(resolution,[],[f1019,f810])).
fof(f28836,plain,(
  ( ! [X404] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X404))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(subsumption_resolution,[],[f28484,f868])).
fof(f28484,plain,(
  ( ! [X404] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X404),s('type.integer.int',X404)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X404))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f864,f848])).
fof(f28720,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.num.|48|_0')))) | $spl243),
  inference(cnf_transformation,[],[f28720_D])).
fof(f28720_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f28777,plain,(
  ( ! [X490,X491] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.num.|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X490),s('type.integer.int',X491))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X490),s('type.integer.int',X491)))) )),
  inference(forward_demodulation,[],[f28776,f862])).
fof(f862,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f391])).
fof(f391,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f212])).
fof(f212,axiom,(
  ! [X12,X10] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',X10))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_LE')).
fof(f28776,plain,(
  ( ! [X490,X491] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X490),s('type.integer.int',X491))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X490),s('type.integer.int',X491)))) )),
  inference(forward_demodulation,[],[f28775,f855])).
fof(f855,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f384])).
fof(f384,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X10] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ABS_NUM')).
fof(f28775,plain,(
  ( ! [X490,X491] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X490),s('type.integer.int',X491))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X490),s('type.integer.int',X491))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f28228,f867])).
fof(f28228,plain,(
  ( ! [X490,X491] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X490),s('type.integer.int',X491))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X490),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X491))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f812,f848])).
fof(f28774,plain,(
  ( ! [X489] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X489))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X489)))))) ) | $spl244),
  inference(cnf_transformation,[],[f28774_D])).
fof(f28774_D,plain,(
  ( ! [X489] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X489))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X489)))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f28773,plain,(
  ( ! [X489] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X489))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X489)))))) )),
  inference(forward_demodulation,[],[f28772,f862])).
fof(f28772,plain,(
  ( ! [X489] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X489))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X489)))))) )),
  inference(forward_demodulation,[],[f28771,f855])).
fof(f28771,plain,(
  ( ! [X489] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X489))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X489))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f28227,f828])).
fof(f828,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))),
  inference(rectify,[],[f164])).
fof(f164,axiom,(
  ! [X9,X11] : s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X11))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int',X11)))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_NEG_LMUL')).
fof(f28227,plain,(
  ( ! [X489] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X489))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X489))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f840,f848])).
fof(f24071,plain,(
  ( ! [X420] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X420))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl220),
  inference(cnf_transformation,[],[f24071_D])).
fof(f24071_D,plain,(
  ( ! [X420] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X420))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f28710,plain,(
  ( ! [X415] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X415),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X415)) ) | $spl238),
  inference(cnf_transformation,[],[f28710_D])).
fof(f28710_D,plain,(
  ( ! [X415] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X415),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X415)) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f28733,plain,(
  ( ! [X432,X433] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X432),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X433))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X432)) )),
  inference(forward_demodulation,[],[f28732,f862])).
fof(f28732,plain,(
  ( ! [X432,X433] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X432))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X433))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X432)) )),
  inference(forward_demodulation,[],[f28178,f855])).
fof(f28178,plain,(
  ( ! [X432,X433] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X433))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X432) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X432))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f723,f848])).
fof(f28731,plain,(
  ( ! [X426,X427] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X426),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X427)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X426)) )),
  inference(forward_demodulation,[],[f28730,f862])).
fof(f28730,plain,(
  ( ! [X426,X427] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X426))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X427)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X426)) )),
  inference(forward_demodulation,[],[f28729,f855])).
fof(f28729,plain,(
  ( ! [X426,X427] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X427)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X426) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X426))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f28728,f815])).
fof(f28728,plain,(
  ( ! [X426,X427] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X427)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X426) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X426))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f28175,f846])).
fof(f28175,plain,(
  ( ! [X426,X427] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X427)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X426) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X426))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f722,f848])).
fof(f24066,plain,(
  ( ! [X412] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X412))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X412)) ) | $spl214),
  inference(cnf_transformation,[],[f24066_D])).
fof(f24066_D,plain,(
  ( ! [X412] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X412))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X412)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f28705,plain,(
  ( ! [X413] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X413),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X413)) ) | $spl236),
  inference(cnf_transformation,[],[f28705_D])).
fof(f28705_D,plain,(
  ( ! [X413] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X413),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X413)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f28727,plain,(
  ( ! [X424,X425] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X424),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X425))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X424) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X425)) )),
  inference(forward_demodulation,[],[f28726,f862])).
fof(f28726,plain,(
  ( ! [X424,X425] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X424))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X425))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X424) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X425)) )),
  inference(forward_demodulation,[],[f28174,f855])).
fof(f28174,plain,(
  ( ! [X424,X425] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X425))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X424) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X425) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X424))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f721,f848])).
fof(f28725,plain,(
  ( ! [X422,X423] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X422),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X423)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X423) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X422)) )),
  inference(forward_demodulation,[],[f28724,f862])).
fof(f28724,plain,(
  ( ! [X422,X423] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X422))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X423)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X423) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X422)) )),
  inference(forward_demodulation,[],[f28723,f855])).
fof(f28723,plain,(
  ( ! [X422,X423] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X423)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X423) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X422) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X422))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f28722,f815])).
fof(f28722,plain,(
  ( ! [X422,X423] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X423)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X423) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X422) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X422))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f28173,f846])).
fof(f28173,plain,(
  ( ! [X422,X423] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X423)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X423) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X422) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X422))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f721,f848])).
fof(f28718,plain,(
  ( ! [X419] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X419))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X419)))) )),
  inference(forward_demodulation,[],[f28717,f862])).
fof(f28717,plain,(
  ( ! [X419] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X419))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X419)))) )),
  inference(forward_demodulation,[],[f28716,f855])).
fof(f28716,plain,(
  ( ! [X419] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X419))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X419)))) )),
  inference(forward_demodulation,[],[f28171,f844])).
fof(f28171,plain,(
  ( ! [X419] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X419))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X419))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f840,f848])).
fof(f28715,plain,(
  ( ! [X417] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X417),s('type.num.num','const.num.|48|_0'))))) ) | $spl240),
  inference(cnf_transformation,[],[f28715_D])).
fof(f28715_D,plain,(
  ( ! [X417] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X417),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f28714,plain,(
  ( ! [X418,X417] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X417),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X418)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X418)) )),
  inference(forward_demodulation,[],[f28713,f862])).
fof(f28713,plain,(
  ( ! [X418,X417] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X417))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X418)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X418)) )),
  inference(forward_demodulation,[],[f28712,f855])).
fof(f28712,plain,(
  ( ! [X418,X417] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X417))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X418)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X418)) )),
  inference(forward_demodulation,[],[f28711,f844])).
fof(f28711,plain,(
  ( ! [X418,X417] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X418)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X418) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X417))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f28170,f815])).
fof(f28170,plain,(
  ( ! [X418,X417] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X418)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X418) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X417))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f723,f848])).
fof(f28709,plain,(
  ( ! [X416,X415] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X415),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X416)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X415)) )),
  inference(forward_demodulation,[],[f28708,f862])).
fof(f28708,plain,(
  ( ! [X416,X415] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X415))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X416)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X415)) )),
  inference(forward_demodulation,[],[f28707,f855])).
fof(f28707,plain,(
  ( ! [X416,X415] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X415))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X416)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X415)) )),
  inference(forward_demodulation,[],[f28706,f844])).
fof(f28706,plain,(
  ( ! [X416,X415] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X416)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X415) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X415))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f28169,f815])).
fof(f28169,plain,(
  ( ! [X416,X415] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X416)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X415) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X415))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f722,f848])).
fof(f28704,plain,(
  ( ! [X414,X413] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X413),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X414)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X414) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X413)) )),
  inference(forward_demodulation,[],[f28703,f862])).
fof(f28703,plain,(
  ( ! [X414,X413] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X413))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X414)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X414) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X413)) )),
  inference(forward_demodulation,[],[f28702,f855])).
fof(f28702,plain,(
  ( ! [X414,X413] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X413))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X414)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X414) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X413)) )),
  inference(forward_demodulation,[],[f28701,f844])).
fof(f28701,plain,(
  ( ! [X414,X413] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X414)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X414) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X413) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X413))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f28168,f815])).
fof(f28168,plain,(
  ( ! [X414,X413] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X414)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X414) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X413) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X413))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f721,f848])).
fof(f28690,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(forward_demodulation,[],[f28689,f856])).
fof(f28689,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(forward_demodulation,[],[f28160,f844])).
fof(f28160,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(superposition,[],[f1484,f848])).
fof(f1509,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl43),
  inference(cnf_transformation,[],[f1509_D])).
fof(f1509_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f27501,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl35),
  inference(forward_demodulation,[],[f27355,f846])).
fof(f27355,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl35),
  inference(resolution,[],[f3490,f803])).
fof(f26916,plain,(
  ( ! [X35] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f26647,f868])).
fof(f26647,plain,(
  ( ! [X35] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X35))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X35)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f864,f851])).
fof(f26915,plain,(
  ( ! [X34] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X34)))) )),
  inference(subsumption_resolution,[],[f26646,f868])).
fof(f26646,plain,(
  ( ! [X34] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X34),s('type.integer.int',X34)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X34)))) )),
  inference(superposition,[],[f864,f850])).
fof(f26914,plain,(
  ( ! [X33] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27(X33)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26(X33))) ) | $spl234),
  inference(cnf_transformation,[],[f26914_D])).
fof(f26914_D,plain,(
  ( ! [X33] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27(X33)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26(X33))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f26913,plain,(
  ( ! [X33] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27(X33)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26(X33))) )),
  inference(subsumption_resolution,[],[f26645,f868])).
fof(f26645,plain,(
  ( ! [X33] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X33),s('type.integer.int',X33)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27(X33)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26(X33))) )),
  inference(superposition,[],[f864,f838])).
fof(f26912,plain,(
  ( ! [X32] : (s('type.integer.int',X32) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))) ) | $spl232),
  inference(cnf_transformation,[],[f26912_D])).
fof(f26912_D,plain,(
  ( ! [X32] : (s('type.integer.int',X32) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f26911,plain,(
  ( ! [X32] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))) )),
  inference(subsumption_resolution,[],[f26644,f868])).
fof(f26644,plain,(
  ( ! [X32] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X32),s('type.integer.int',X32)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))) )),
  inference(superposition,[],[f864,f833])).
fof(f26910,plain,(
  ( ! [X31] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X31) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f26643,f868])).
fof(f26643,plain,(
  ( ! [X31] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X31) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f864,f832])).
fof(f832,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f562])).
fof(f26907,plain,(
  ( ! [X30] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30)))) ) | $spl228),
  inference(cnf_transformation,[],[f26907_D])).
fof(f26907_D,plain,(
  ( ! [X30] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30)))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f26906,plain,(
  ( ! [X30] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30)))) )),
  inference(subsumption_resolution,[],[f26642,f868])).
fof(f26642,plain,(
  ( ! [X30] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X30),s('type.integer.int',X30)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30)))) )),
  inference(superposition,[],[f864,f831])).
fof(f25154,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl35),
  inference(subsumption_resolution,[],[f25127,f758])).
fof(f25127,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl35),
  inference(resolution,[],[f17916,f704])).
fof(f25153,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl35),
  inference(subsumption_resolution,[],[f25126,f758])).
fof(f25126,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl35),
  inference(resolution,[],[f17916,f708])).
fof(f25123,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(subsumption_resolution,[],[f25096,f758])).
fof(f25096,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(resolution,[],[f17915,f704])).
fof(f25122,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(subsumption_resolution,[],[f25095,f758])).
fof(f25095,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(resolution,[],[f17915,f708])).
fof(f25041,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl227),
  inference(cnf_transformation,[],[f25041_D])).
fof(f25041_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl227),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl227])])).
fof(f24745,plain,(
  ( ! [X78] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X78))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X78))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f840,f832])).
fof(f24058,plain,(
  ( ! [X404] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X404)))) ) | $spl206),
  inference(cnf_transformation,[],[f24058_D])).
fof(f24058_D,plain,(
  ( ! [X404] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X404)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f25039,plain,(
  ( ! [X76,X77] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X77)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X77) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76)))) )),
  inference(forward_demodulation,[],[f24744,f815])).
fof(f24744,plain,(
  ( ! [X76,X77] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X77)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X77) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76)))) )),
  inference(superposition,[],[f723,f832])).
fof(f25038,plain,(
  ( ! [X72,X73] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X73)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X73) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X72)))) )),
  inference(forward_demodulation,[],[f25037,f815])).
fof(f25037,plain,(
  ( ! [X72,X73] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X73)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X73) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X72)))) )),
  inference(subsumption_resolution,[],[f24742,f860])).
fof(f860,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f570])).
fof(f570,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f390])).
fof(f390,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f210])).
fof(f210,axiom,(
  ! [X12,X10] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))) <=> s('type.num.num',X10) = s('type.num.num',X12))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_INJ')).
fof(f24742,plain,(
  ( ! [X72,X73] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X73)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X73) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X72) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X72)))) )),
  inference(superposition,[],[f721,f832])).
fof(f24739,plain,(
  ( ! [X66,X67] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X67))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X66)))) )),
  inference(superposition,[],[f818,f832])).
fof(f1500,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl37),
  inference(cnf_transformation,[],[f1500_D])).
fof(f1500_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f24737,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f876,f832])).
fof(f24734,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl30),
  inference(superposition,[],[f1484,f832])).
fof(f24364,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl35),
  inference(subsumption_resolution,[],[f24338,f758])).
fof(f24338,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl35),
  inference(resolution,[],[f17262,f704])).
fof(f24363,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl35),
  inference(subsumption_resolution,[],[f24337,f758])).
fof(f24337,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl35),
  inference(resolution,[],[f17262,f708])).
fof(f24334,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(subsumption_resolution,[],[f24308,f758])).
fof(f24308,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(resolution,[],[f17261,f704])).
fof(f24333,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(subsumption_resolution,[],[f24307,f758])).
fof(f24307,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(resolution,[],[f17261,f708])).
fof(f24143,plain,(
  $false | ($spl187 | $spl218)),
  inference(resolution,[],[f24070,f20495])).
fof(f24144,plain,(
  $false | $spl218),
  inference(resolution,[],[f24070,f20496])).
fof(f20496,plain,(
  ( ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f20437,f815])).
fof(f20437,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f819,f855])).
fof(f24138,plain,(
  $false | $spl200),
  inference(equality_resolution,[],[f24051])).
fof(f24077,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl223),
  inference(cnf_transformation,[],[f24077_D])).
fof(f24077_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> ~$spl223),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl223])])).
fof(f24078,plain,(
  ( ! [X463,X464] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X463),s('type.integer.int',X464))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X463),s('type.integer.int',X464))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(forward_demodulation,[],[f23701,f867])).
fof(f23701,plain,(
  ( ! [X463,X464] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X463),s('type.integer.int',X464))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X463),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X464))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f812,f831])).
fof(f24075,plain,(
  ( ! [X462] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X462))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X462))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(forward_demodulation,[],[f23700,f846])).
fof(f23700,plain,(
  ( ! [X462] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X462))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X462))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f840,f831])).
fof(f24069,plain,(
  ( ! [X413] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X413) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X413)))))) ) | $spl216),
  inference(cnf_transformation,[],[f24069_D])).
fof(f24069_D,plain,(
  ( ! [X413] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X413) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X413)))))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f23669,plain,(
  ( ! [X420,X419] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X420))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X419) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X419)))))) )),
  inference(superposition,[],[f723,f831])).
fof(f24068,plain,(
  ( ! [X414,X413] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X414)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X413) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X413)))))) )),
  inference(forward_demodulation,[],[f24067,f815])).
fof(f24067,plain,(
  ( ! [X414,X413] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X414)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X413) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X413)))))) )),
  inference(forward_demodulation,[],[f23666,f846])).
fof(f23666,plain,(
  ( ! [X414,X413] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X414)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X413) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X413)))))) )),
  inference(superposition,[],[f722,f831])).
fof(f24065,plain,(
  ( ! [X409] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X409) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X409)))))) ) | $spl212),
  inference(cnf_transformation,[],[f24065_D])).
fof(f24065_D,plain,(
  ( ! [X409] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X409) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X409)))))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f23665,plain,(
  ( ! [X412,X411] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X412))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X411) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X412) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X411)))))) )),
  inference(superposition,[],[f721,f831])).
fof(f24064,plain,(
  ( ! [X410,X409] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X410)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X410) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X409) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X409)))))) )),
  inference(forward_demodulation,[],[f24063,f815])).
fof(f24063,plain,(
  ( ! [X410,X409] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X410)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X410) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X409) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X409)))))) )),
  inference(forward_demodulation,[],[f23664,f846])).
fof(f23664,plain,(
  ( ! [X410,X409] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X410)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X410) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X409) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X409)))))) )),
  inference(superposition,[],[f721,f831])).
fof(f24060,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | $spl209),
  inference(cnf_transformation,[],[f24060_D])).
fof(f24060_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) <=> ~$spl209),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl209])])).
fof(f23662,plain,(
  ( ! [X406] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X406))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X406))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(superposition,[],[f840,f831])).
fof(f24056,plain,(
  ( ! [X405,X404] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X404))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X405)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X405)) )),
  inference(forward_demodulation,[],[f24055,f845])).
fof(f24055,plain,(
  ( ! [X405,X404] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X405)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X405) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X404)))))))) )),
  inference(forward_demodulation,[],[f23661,f815])).
fof(f23661,plain,(
  ( ! [X405,X404] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X405)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X405) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X404)))))))) )),
  inference(superposition,[],[f723,f831])).
fof(f24051,plain,(
  ( ! [X400] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X400)) ) | $spl200),
  inference(cnf_transformation,[],[f24051_D])).
fof(f24051_D,plain,(
  ( ! [X400] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X400)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f24050,plain,(
  ( ! [X401,X400] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X401)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X401) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X400)) )),
  inference(subsumption_resolution,[],[f24049,f861])).
fof(f861,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f570])).
fof(f24049,plain,(
  ( ! [X401,X400] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X400))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X401)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X401) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X400)) )),
  inference(forward_demodulation,[],[f24048,f845])).
fof(f24048,plain,(
  ( ! [X401,X400] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X401)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X401) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X400) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X400)))))))) )),
  inference(forward_demodulation,[],[f23659,f815])).
fof(f23659,plain,(
  ( ! [X401,X400] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X401)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X401) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X400) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X400)))))))) )),
  inference(superposition,[],[f721,f831])).
fof(f24045,plain,(
  ( ! [X394] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X394)))))))) ) | $spl196),
  inference(cnf_transformation,[],[f24045_D])).
fof(f24045_D,plain,(
  ( ! [X394] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X394)))))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f23656,plain,(
  ( ! [X395,X394] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X395))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X394)))))))) )),
  inference(superposition,[],[f818,f831])).
fof(f24043,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f23654,f845])).
fof(f23654,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl0),
  inference(superposition,[],[f876,f831])).
fof(f24041,plain,(
  ( ! [X391] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X391)) ) | $spl0),
  inference(forward_demodulation,[],[f23652,f845])).
fof(f23652,plain,(
  ( ! [X391] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X391) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) ) | $spl0),
  inference(superposition,[],[f5725,f831])).
fof(f24040,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(forward_demodulation,[],[f23651,f845])).
fof(f23651,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl30),
  inference(superposition,[],[f1484,f831])).
fof(f20495,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl187),
  inference(cnf_transformation,[],[f20495_D])).
fof(f20495_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f20502,plain,(
  ( ! [X14] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f20501,f856])).
fof(f20501,plain,(
  ( ! [X14] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X14)))))) )),
  inference(forward_demodulation,[],[f20444,f815])).
fof(f20444,plain,(
  ( ! [X14] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X14)))))) )),
  inference(superposition,[],[f819,f850])).
fof(f20499,plain,(
  ( ! [X13] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X13)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X13))))) )),
  inference(forward_demodulation,[],[f20443,f815])).
fof(f20443,plain,(
  ( ! [X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',X13)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.ABS_1'(s('type.integer.int',X13))))) )),
  inference(superposition,[],[f819,f838])).
fof(f20497,plain,(
  ( ! [X12] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X12)))))) )),
  inference(forward_demodulation,[],[f20442,f815])).
fof(f20442,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X12)))))) )),
  inference(superposition,[],[f819,f833])).
fof(f20492,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f20434,f815])).
fof(f20434,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f819,f851])).
fof(f20147,plain,(
  $false | ($spl5 | $spl35 | $spl36 | $spl101)),
  inference(subsumption_resolution,[],[f20138,f9004])).
fof(f20138,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl5 | $spl35 | $spl36)),
  inference(resolution,[],[f19782,f11410])).
fof(f11410,plain,(
  ( ! [X143] : (~sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X143) | p(s(bool,X143))) ) | $spl5),
  inference(subsumption_resolution,[],[f11363,f758])).
fof(f11363,plain,(
  ( ! [X143] : (p(s(bool,'F_0')) | p(s(bool,X143)) | ~sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X143)) ) | $spl5),
  inference(resolution,[],[f881,f704])).
fof(f19782,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) ) | ($spl35 | $spl36)),
  inference(backward_demodulation,[],[f1499,f17261])).
fof(f1499,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(cnf_transformation,[],[f1499_D])).
fof(f1499_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f19749,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl35),
  inference(subsumption_resolution,[],[f19712,f758])).
fof(f19712,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl35),
  inference(resolution,[],[f17262,f704])).
fof(f19748,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl35),
  inference(subsumption_resolution,[],[f19711,f758])).
fof(f19711,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl35),
  inference(resolution,[],[f17262,f708])).
fof(f19707,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(subsumption_resolution,[],[f19670,f758])).
fof(f19670,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(resolution,[],[f17261,f704])).
fof(f19706,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(subsumption_resolution,[],[f19669,f758])).
fof(f19669,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl35),
  inference(resolution,[],[f17261,f708])).
fof(f19174,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl30),
  inference(superposition,[],[f1484,f832])).
fof(f19218,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(forward_demodulation,[],[f19182,f845])).
fof(f19182,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl30),
  inference(superposition,[],[f1484,f831])).
fof(f19223,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(forward_demodulation,[],[f19222,f856])).
fof(f19222,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(forward_demodulation,[],[f19189,f844])).
fof(f19189,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl30),
  inference(superposition,[],[f1484,f850])).
fof(f19444,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) | $spl178),
  inference(cnf_transformation,[],[f19444_D])).
fof(f19444_D,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f13542,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6567,f853])).
fof(f6567,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) ) | $spl2),
  inference(resolution,[],[f1019,f705])).
fof(f19443,plain,(
  ( ! [X0] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0)) ) | $spl176),
  inference(cnf_transformation,[],[f19443_D])).
fof(f19443_D,plain,(
  ( ! [X0] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f13553,plain,(
  ( ! [X0] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6568,f853])).
fof(f6568,plain,(
  ( ! [X145] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))),X145)) ) | $spl2),
  inference(resolution,[],[f1019,f706])).
fof(f19448,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl181),
  inference(cnf_transformation,[],[f19448_D])).
fof(f19448_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f14146,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl101),
  inference(resolution,[],[f9004,f809])).
fof(f19432,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl157),
  inference(cnf_transformation,[],[f19432_D])).
fof(f19432_D,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f15307,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(superposition,[],[f15240,f834])).
fof(f15240,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) ) | $spl20),
  inference(resolution,[],[f1320,f705])).
fof(f15325,plain,(
  ( ! [X0] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(superposition,[],[f15241,f834])).
fof(f15241,plain,(
  ( ! [X145] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X145)) ) | $spl20),
  inference(resolution,[],[f1320,f706])).
fof(f19427,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) | $spl154),
  inference(cnf_transformation,[],[f19427_D])).
fof(f19427_D,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f15500,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f15437,f834])).
fof(f19426,plain,(
  ( ! [X0] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0)) ) | $spl152),
  inference(cnf_transformation,[],[f19426_D])).
fof(f19426_D,plain,(
  ( ! [X0] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f15537,plain,(
  ( ! [X0] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f15438,f834])).
fof(f19441,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) ) | $spl174),
  inference(cnf_transformation,[],[f19441_D])).
fof(f19441_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f15555,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f2515,f834])).
fof(f2515,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) ) | $spl7),
  inference(resolution,[],[f1011,f583])).
fof(f19440,plain,(
  ( ! [X0,X1] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1)) ) | $spl172),
  inference(cnf_transformation,[],[f19440_D])).
fof(f19440_D,plain,(
  ( ! [X0,X1] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f15590,plain,(
  ( ! [X0,X1] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f2516,f834])).
fof(f2516,plain,(
  ( ! [X14,X15] : (sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)),X14,X15)) ) | $spl7),
  inference(resolution,[],[f1011,f585])).
fof(f19439,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1)) ) | $spl170),
  inference(cnf_transformation,[],[f19439_D])).
fof(f19439_D,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f15608,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f2532,f834])).
fof(f2532,plain,(
  ( ! [X47,X46] : (sP2('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)),X46,X47)) ) | $spl7),
  inference(resolution,[],[f1011,f602])).
fof(f19438,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X1)) ) | $spl168),
  inference(cnf_transformation,[],[f19438_D])).
fof(f19438_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X1)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f15628,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f2545,f834])).
fof(f2545,plain,(
  ( ! [X72,X73] : (sP4(X72,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)),X73)) ) | $spl7),
  inference(resolution,[],[f1011,f618])).
fof(f19437,plain,(
  ( ! [X0,X1] : (sP4('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1)) ) | $spl166),
  inference(cnf_transformation,[],[f19437_D])).
fof(f19437_D,plain,(
  ( ! [X0,X1] : (sP4('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f15647,plain,(
  ( ! [X0,X1] : (sP4('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f2546,f834])).
fof(f2546,plain,(
  ( ! [X74,X75] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)),X74,X75)) ) | $spl7),
  inference(resolution,[],[f1011,f619])).
fof(f19436,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X1)) ) | $spl164),
  inference(cnf_transformation,[],[f19436_D])).
fof(f19436_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X1)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f15665,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f2558,f834])).
fof(f2558,plain,(
  ( ! [X99,X98] : (sP8(X98,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)),X99)) ) | $spl7),
  inference(resolution,[],[f1011,f635])).
fof(f19435,plain,(
  ( ! [X0,X1] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1)) ) | $spl162),
  inference(cnf_transformation,[],[f19435_D])).
fof(f19435_D,plain,(
  ( ! [X0,X1] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f15684,plain,(
  ( ! [X0,X1] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f2559,f834])).
fof(f2559,plain,(
  ( ! [X101,X100] : (sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)),X100,X101)) ) | $spl7),
  inference(resolution,[],[f1011,f636])).
fof(f19434,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X1)) ) | $spl160),
  inference(cnf_transformation,[],[f19434_D])).
fof(f19434_D,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X1)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f15704,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f2561,f834])).
fof(f2561,plain,(
  ( ! [X105,X104] : (sP7(X104,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)),X105)) ) | $spl7),
  inference(resolution,[],[f1011,f639])).
fof(f19433,plain,(
  ( ! [X0,X1] : (sP7('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1)) ) | $spl158),
  inference(cnf_transformation,[],[f19433_D])).
fof(f19433_D,plain,(
  ( ! [X0,X1] : (sP7('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f15724,plain,(
  ( ! [X0,X1] : (sP7('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0,X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f2562,f834])).
fof(f2562,plain,(
  ( ! [X107,X106] : (sP7('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)),X106,X107)) ) | $spl7),
  inference(resolution,[],[f1011,f640])).
fof(f18382,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f18325,f853])).
fof(f18325,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) ) | $spl2),
  inference(resolution,[],[f6490,f705])).
fof(f18394,plain,(
  ( ! [X0] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f18326,f853])).
fof(f18326,plain,(
  ( ! [X145] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))),X145)) ) | $spl2),
  inference(resolution,[],[f6490,f706])).
fof(f19425,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1)) ) | $spl150),
  inference(cnf_transformation,[],[f19425_D])).
fof(f19425_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f18423,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6498,f853])).
fof(f6498,plain,(
  ( ! [X14,X15] : (sP0(X14,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))),X15)) ) | $spl2),
  inference(resolution,[],[f1019,f584])).
fof(f584,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f453])).
fof(f19424,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) | $spl148),
  inference(cnf_transformation,[],[f19424_D])).
fof(f19424_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f18436,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6511,f853])).
fof(f6511,plain,(
  ( ! [X41,X40] : (sP2(X40,X41,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) ) | $spl2),
  inference(resolution,[],[f1019,f600])).
fof(f600,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f461])).
fof(f19423,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1)) ) | $spl146),
  inference(cnf_transformation,[],[f19423_D])).
fof(f19423_D,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f18449,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6512,f853])).
fof(f6512,plain,(
  ( ! [X43,X42] : (sP2(X42,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))),X43)) ) | $spl2),
  inference(resolution,[],[f1019,f601])).
fof(f601,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f461])).
fof(f19422,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) | $spl144),
  inference(cnf_transformation,[],[f19422_D])).
fof(f19422_D,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f18462,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6528,f853])).
fof(f6528,plain,(
  ( ! [X74,X75] : (sP4(X74,X75,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) ) | $spl2),
  inference(resolution,[],[f1019,f617])).
fof(f617,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f468])).
fof(f19421,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) | $spl142),
  inference(cnf_transformation,[],[f19421_D])).
fof(f19421_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f18475,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6541,f853])).
fof(f6541,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) ) | $spl2),
  inference(resolution,[],[f1019,f634])).
fof(f634,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f476])).
fof(f19420,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) | $spl140),
  inference(cnf_transformation,[],[f19420_D])).
fof(f19420_D,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f18488,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6544,f853])).
fof(f6544,plain,(
  ( ! [X107,X106] : (sP7(X106,X107,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) ) | $spl2),
  inference(resolution,[],[f1019,f638])).
fof(f638,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f479])).
fof(f19419,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) | $spl138),
  inference(cnf_transformation,[],[f19419_D])).
fof(f19419_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f18501,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6545,f853])).
fof(f6545,plain,(
  ( ! [X109,X108] : (sP6(X108,X109,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) ) | $spl2),
  inference(resolution,[],[f1019,f642])).
fof(f642,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f481])).
fof(f481,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f480])).
fof(f480,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f434])).
fof(f434,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f19418,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1)) ) | $spl136),
  inference(cnf_transformation,[],[f19418_D])).
fof(f19418_D,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f18515,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6546,f853])).
fof(f6546,plain,(
  ( ! [X111,X110] : (sP6(X110,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))),X111)) ) | $spl2),
  inference(resolution,[],[f1019,f643])).
fof(f643,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f481])).
fof(f19417,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0,X1)) ) | $spl134),
  inference(cnf_transformation,[],[f19417_D])).
fof(f19417_D,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0,X1)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f18529,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0,X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6547,f853])).
fof(f6547,plain,(
  ( ! [X112,X113] : (sP6('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))),X112,X113)) ) | $spl2),
  inference(resolution,[],[f1019,f644])).
fof(f644,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f481])).
fof(f19416,plain,(
  ( ! [X0,X1] : (sP10('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0,X1)) ) | $spl132),
  inference(cnf_transformation,[],[f19416_D])).
fof(f19416_D,plain,(
  ( ! [X0,X1] : (sP10('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0,X1)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f18545,plain,(
  ( ! [X0,X1] : (sP10('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0,X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6558,f853])).
fof(f6558,plain,(
  ( ! [X134,X135] : (sP10('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))),X134,X135)) ) | $spl2),
  inference(resolution,[],[f1019,f662])).
fof(f662,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X2,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f489])).
fof(f489,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) | ~p(s(bool,X2))) & ~p(s(bool,X0))) | sP10(X0,X2,X1)) & (~sP10(X0,X2,X1) | (p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))))),
  inference(flattening,[],[f488])).
fof(f488,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) | ~p(s(bool,X2))) & ~p(s(bool,X0))) | sP10(X0,X2,X1)) & (~sP10(X0,X2,X1) | ((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f440])).
fof(f440,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))) <=> sP10(X0,X2,X1))),
  inference(definition_folding,[],[f273,f439])).
fof(f439,plain,(
  ! [X0,X2,X1] : (sP10(X0,X2,X1) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f273,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  inference(rectify,[],[f262])).
fof(f262,axiom,(
  ! [X8,X21,X39] : (((p(s(bool,X21)) & p(s(bool,X39))) | p(s(bool,X8))) <=> ((p(s(bool,X21)) | p(s(bool,X8))) & (p(s(bool,X39)) | p(s(bool,X8)))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.bool.RIGHT_OR_OVER_AND')).
fof(f19396,plain,(
  $false | ($spl37 | $spl43)),
  inference(subsumption_resolution,[],[f19395,f854])).
fof(f19395,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl37 | $spl43)),
  inference(subsumption_resolution,[],[f19230,f1500])).
fof(f19230,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl43),
  inference(resolution,[],[f1509,f809])).
fof(f19394,plain,(
  $false | ($spl35 | $spl43)),
  inference(subsumption_resolution,[],[f19393,f3492])).
fof(f19393,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl43),
  inference(forward_demodulation,[],[f19229,f846])).
fof(f19229,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl43),
  inference(resolution,[],[f1509,f805])).
fof(f19392,plain,(
  $false | ($spl37 | $spl43)),
  inference(subsumption_resolution,[],[f19391,f819])).
fof(f19391,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl37 | $spl43)),
  inference(forward_demodulation,[],[f19390,f822])).
fof(f19390,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl37 | $spl43)),
  inference(subsumption_resolution,[],[f19227,f1500])).
fof(f19227,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl43),
  inference(resolution,[],[f1509,f786])).
fof(f19221,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(forward_demodulation,[],[f19220,f856])).
fof(f19220,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(forward_demodulation,[],[f19188,f844])).
fof(f19188,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(superposition,[],[f1484,f848])).
fof(f19187,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl30),
  inference(superposition,[],[f1484,f838])).
fof(f19186,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl30),
  inference(superposition,[],[f1484,f837])).
fof(f19185,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl30),
  inference(superposition,[],[f1484,f836])).
fof(f19184,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl30),
  inference(superposition,[],[f1484,f835])).
fof(f19219,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(forward_demodulation,[],[f19183,f845])).
fof(f19183,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl30),
  inference(superposition,[],[f1484,f833])).
fof(f1510,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl42),
  inference(cnf_transformation,[],[f1510_D])).
fof(f1510_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f19215,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(subsumption_resolution,[],[f19214,f819])).
fof(f19214,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl30),
  inference(forward_demodulation,[],[f19213,f822])).
fof(f19213,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl30),
  inference(forward_demodulation,[],[f19212,f821])).
fof(f19212,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl30),
  inference(forward_demodulation,[],[f19177,f822])).
fof(f19177,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl30),
  inference(superposition,[],[f1484,f786])).
fof(f19049,plain,(
  $false | ($spl0 | $spl31)),
  inference(subsumption_resolution,[],[f18904,f876])).
fof(f18904,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl31),
  inference(resolution,[],[f1485,f810])).
fof(f18747,plain,(
  $false | ($spl0 | $spl27)),
  inference(subsumption_resolution,[],[f18612,f876])).
fof(f18612,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl27),
  inference(resolution,[],[f1481,f805])).
fof(f17160,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl35),
  inference(forward_demodulation,[],[f17002,f846])).
fof(f17002,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl35),
  inference(resolution,[],[f3490,f803])).
fof(f14434,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | $spl101),
  inference(resolution,[],[f14144,f800])).
fof(f14423,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl5),
  inference(forward_demodulation,[],[f14289,f856])).
fof(f14289,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int',sK13)))) | $spl5),
  inference(resolution,[],[f11282,f800])).
fof(f11282,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl5),
  inference(resolution,[],[f881,f806])).
fof(f14288,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl5),
  inference(resolution,[],[f11282,f803])).
fof(f14286,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl5),
  inference(resolution,[],[f11282,f809])).
fof(f14149,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl101),
  inference(resolution,[],[f9004,f800])).
fof(f13967,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl29),
  inference(forward_demodulation,[],[f13838,f856])).
fof(f13838,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK14)))) | $spl29),
  inference(resolution,[],[f1482,f800])).
fof(f13706,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f808])).
fof(f13535,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X1))) ) | $spl2),
  inference(subsumption_resolution,[],[f13526,f758])).
fof(f13526,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X1))) ) | $spl2),
  inference(resolution,[],[f3408,f704])).
fof(f13534,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f13525,f758])).
fof(f13525,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X0))) ) | $spl2),
  inference(resolution,[],[f3408,f708])).
fof(f13504,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f13495,f758])).
fof(f13495,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(resolution,[],[f3407,f704])).
fof(f13503,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f13494,f758])).
fof(f13494,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(resolution,[],[f3407,f708])).
fof(f12745,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f1019,f808])).
fof(f12538,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl7),
  inference(resolution,[],[f1011,f809])).
fof(f12078,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f12077,f856])).
fof(f12077,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f12041,f844])).
fof(f12041,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(superposition,[],[f876,f848])).
fof(f12040,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(superposition,[],[f876,f838])).
fof(f12039,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(superposition,[],[f876,f837])).
fof(f12038,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(superposition,[],[f876,f836])).
fof(f12037,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(superposition,[],[f876,f835])).
fof(f12076,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f12036,f845])).
fof(f12036,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl0),
  inference(superposition,[],[f876,f833])).
fof(f12072,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(subsumption_resolution,[],[f12071,f819])).
fof(f12071,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f12070,f822])).
fof(f12070,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(forward_demodulation,[],[f12069,f821])).
fof(f12069,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(forward_demodulation,[],[f12030,f822])).
fof(f12030,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(superposition,[],[f876,f786])).
fof(f12066,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(forward_demodulation,[],[f11923,f845])).
fof(f11923,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK14)))) | $spl0),
  inference(resolution,[],[f876,f808])).
fof(f1483,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl28),
  inference(cnf_transformation,[],[f1483_D])).
fof(f1483_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f11635,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl12),
  inference(resolution,[],[f1164,f808])).
fof(f1313,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl14),
  inference(cnf_transformation,[],[f1313_D])).
fof(f1313_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f11431,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | p(s(bool,X1))) ) | $spl5),
  inference(subsumption_resolution,[],[f11425,f758])).
fof(f11425,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | p(s(bool,X1))) ) | $spl5),
  inference(resolution,[],[f5575,f704])).
fof(f5575,plain,(
  ( ! [X145] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X145)) ) | $spl5),
  inference(resolution,[],[f5216,f706])).
fof(f11430,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | p(s(bool,X0))) ) | $spl5),
  inference(subsumption_resolution,[],[f11424,f758])).
fof(f11424,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | p(s(bool,X0))) ) | $spl5),
  inference(resolution,[],[f5575,f708])).
fof(f11423,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f11417,f758])).
fof(f11417,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) ) | $spl5),
  inference(resolution,[],[f5574,f704])).
fof(f5574,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f5216,f705])).
fof(f11422,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f11416,f758])).
fof(f11416,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) ) | $spl5),
  inference(resolution,[],[f5574,f708])).
fof(f1010,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl6),
  inference(cnf_transformation,[],[f1010_D])).
fof(f1010_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f11284,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl5),
  inference(resolution,[],[f881,f800])).
fof(f1502,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl39),
  inference(cnf_transformation,[],[f1502_D])).
fof(f1502_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f878,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl2),
  inference(cnf_transformation,[],[f878_D])).
fof(f878_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f11283,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl5),
  inference(resolution,[],[f881,f803])).
fof(f11277,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13))))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f11273,f758])).
fof(f11273,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13))))) ) | ($spl12 | $spl16)),
  inference(resolution,[],[f11257,f704])).
fof(f11257,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f5912,f1316])).
fof(f5912,plain,(
  ( ! [X144] : (sP12(X144,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) ) | $spl12),
  inference(resolution,[],[f1164,f705])).
fof(f11276,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13))))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f11272,f758])).
fof(f11272,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13))))) ) | ($spl12 | $spl16)),
  inference(resolution,[],[f11257,f708])).
fof(f11140,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ($spl16 | $spl35)),
  inference(subsumption_resolution,[],[f11139,f3491])).
fof(f3491,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl35),
  inference(resolution,[],[f1498,f811])).
fof(f11139,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl16),
  inference(forward_demodulation,[],[f11138,f1316])).
fof(f11138,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl16),
  inference(forward_demodulation,[],[f11137,f1316])).
fof(f11137,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl16),
  inference(forward_demodulation,[],[f10964,f822])).
fof(f10964,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl16),
  inference(superposition,[],[f786,f1316])).
fof(f11136,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl16 | $spl35)),
  inference(subsumption_resolution,[],[f11135,f3491])).
fof(f11135,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl16),
  inference(forward_demodulation,[],[f11134,f1316])).
fof(f11134,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl16),
  inference(forward_demodulation,[],[f11133,f1316])).
fof(f11133,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl16),
  inference(forward_demodulation,[],[f10963,f822])).
fof(f10963,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl16),
  inference(superposition,[],[f786,f1316])).
fof(f8735,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl29),
  inference(resolution,[],[f1482,f808])).
fof(f10939,plain,(
  $false | ($spl2 | $spl64)),
  inference(subsumption_resolution,[],[f10938,f868])).
fof(f10938,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ($spl2 | $spl64)),
  inference(forward_demodulation,[],[f1024,f2795])).
fof(f1024,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f878,f797])).
fof(f10932,plain,(
  $false | ($spl2 | $spl64)),
  inference(subsumption_resolution,[],[f2810,f2795])).
fof(f2810,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl2),
  inference(resolution,[],[f1024,f811])).
fof(f10888,plain,(
  $false | ($spl2 | $spl64)),
  inference(subsumption_resolution,[],[f10887,f868])).
fof(f10887,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ($spl2 | $spl64)),
  inference(forward_demodulation,[],[f5354,f2795])).
fof(f5354,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f878,f797])).
fof(f10886,plain,(
  $false | ($spl2 | $spl64)),
  inference(subsumption_resolution,[],[f10885,f868])).
fof(f10885,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ($spl2 | $spl64)),
  inference(forward_demodulation,[],[f5355,f2795])).
fof(f5355,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f878,f799])).
fof(f10860,plain,(
  $false | ($spl2 | $spl64)),
  inference(subsumption_resolution,[],[f8053,f2795])).
fof(f8053,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl2),
  inference(resolution,[],[f1024,f811])).
fof(f10854,plain,(
  $false | ($spl2 | $spl64)),
  inference(subsumption_resolution,[],[f10853,f868])).
fof(f10853,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ($spl2 | $spl64)),
  inference(forward_demodulation,[],[f7310,f2795])).
fof(f7310,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f878,f797])).
fof(f10852,plain,(
  $false | ($spl2 | $spl64)),
  inference(subsumption_resolution,[],[f10851,f868])).
fof(f10851,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ($spl2 | $spl64)),
  inference(forward_demodulation,[],[f7311,f2795])).
fof(f7311,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f878,f799])).
fof(f10800,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | ($spl29 | $spl64)),
  inference(forward_demodulation,[],[f8735,f2795])).
fof(f10799,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl3),
  inference(forward_demodulation,[],[f9169,f844])).
fof(f9169,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | $spl3),
  inference(forward_demodulation,[],[f9036,f817])).
fof(f9036,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | $spl3),
  inference(resolution,[],[f879,f803])).
fof(f10118,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | ($spl7 | $spl64)),
  inference(forward_demodulation,[],[f9994,f2795])).
fof(f9994,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | ($spl7 | $spl64)),
  inference(resolution,[],[f9611,f800])).
fof(f10790,plain,(
  $false | ($spl4 | $spl12)),
  inference(subsumption_resolution,[],[f10662,f1164])).
fof(f10662,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f880,f734])).
fof(f10789,plain,(
  $false | ($spl4 | $spl12)),
  inference(subsumption_resolution,[],[f10661,f1164])).
fof(f10661,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f880,f734])).
fof(f10788,plain,(
  $false | ($spl4 | $spl12)),
  inference(subsumption_resolution,[],[f10660,f1164])).
fof(f10660,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f880,f679])).
fof(f10787,plain,(
  $false | ($spl4 | $spl12)),
  inference(subsumption_resolution,[],[f10659,f1164])).
fof(f10659,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f880,f679])).
fof(f1315,plain,(
  s('type.integer.int',sK13) != s('type.integer.int',sK14) | $spl17),
  inference(cnf_transformation,[],[f1315_D])).
fof(f1315_D,plain,(
  s('type.integer.int',sK13) != s('type.integer.int',sK14) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f10524,plain,(
  s('type.integer.int',sK13) != s('type.integer.int',sK14) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | ($spl35 | $spl64)),
  inference(superposition,[],[f5204,f853])).
fof(f5204,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ($spl35 | $spl64)),
  inference(forward_demodulation,[],[f5203,f2795])).
fof(f5203,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ($spl35 | $spl64)),
  inference(forward_demodulation,[],[f5202,f844])).
fof(f5202,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | ($spl35 | $spl64)),
  inference(subsumption_resolution,[],[f5029,f1498])).
fof(f5029,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | $spl64),
  inference(superposition,[],[f852,f2795])).
fof(f9033,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl3),
  inference(resolution,[],[f879,f803])).
fof(f10651,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | $spl3),
  inference(forward_demodulation,[],[f9170,f817])).
fof(f9170,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl3),
  inference(forward_demodulation,[],[f9037,f845])).
fof(f9037,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl3),
  inference(resolution,[],[f879,f800])).
fof(f10116,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl7 | $spl64)),
  inference(forward_demodulation,[],[f9993,f2795])).
fof(f9993,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl7 | $spl64)),
  inference(resolution,[],[f9611,f803])).
fof(f10643,plain,(
  $false | ($spl0 | $spl39)),
  inference(subsumption_resolution,[],[f5788,f1502])).
fof(f5788,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f5787,f856])).
fof(f5787,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(forward_demodulation,[],[f5645,f844])).
fof(f5645,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | $spl0),
  inference(resolution,[],[f876,f804])).
fof(f1321,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl22),
  inference(cnf_transformation,[],[f1321_D])).
fof(f1321_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f10139,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl39),
  inference(resolution,[],[f1502,f800])).
fof(f10268,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl115),
  inference(cnf_transformation,[],[f10268_D])).
fof(f10268_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f10136,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl39),
  inference(resolution,[],[f1502,f809])).
fof(f9988,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl64),
  inference(forward_demodulation,[],[f9815,f4877])).
fof(f9815,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK29('const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl64),
  inference(superposition,[],[f864,f4877])).
fof(f9987,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl64),
  inference(forward_demodulation,[],[f9810,f841])).
fof(f9810,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl64),
  inference(superposition,[],[f853,f4877])).
fof(f9986,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl64),
  inference(forward_demodulation,[],[f9985,f4877])).
fof(f9985,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl64),
  inference(forward_demodulation,[],[f9809,f844])).
fof(f9809,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | $spl64),
  inference(superposition,[],[f853,f4877])).
fof(f9982,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl64),
  inference(forward_demodulation,[],[f9808,f841])).
fof(f9808,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl64),
  inference(superposition,[],[f852,f4877])).
fof(f4038,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl81),
  inference(cnf_transformation,[],[f4038_D])).
fof(f4038_D,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f9981,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl64),
  inference(forward_demodulation,[],[f9980,f4877])).
fof(f9980,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl64),
  inference(forward_demodulation,[],[f9807,f844])).
fof(f9807,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | $spl64),
  inference(superposition,[],[f852,f4877])).
fof(f9975,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl64),
  inference(forward_demodulation,[],[f9974,f4877])).
fof(f9974,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl64),
  inference(forward_demodulation,[],[f9794,f841])).
fof(f9794,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl64),
  inference(superposition,[],[f839,f4877])).
fof(f9973,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl64),
  inference(forward_demodulation,[],[f9972,f841])).
fof(f9972,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl64),
  inference(forward_demodulation,[],[f9793,f841])).
fof(f9793,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl64),
  inference(superposition,[],[f839,f4877])).
fof(f9855,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | ($spl11 | $spl64)),
  inference(subsumption_resolution,[],[f9854,f1162])).
fof(f9854,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl64),
  inference(forward_demodulation,[],[f9853,f4877])).
fof(f9853,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f9852,f821])).
fof(f9852,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f9665,f822])).
fof(f9665,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(superposition,[],[f786,f4877])).
fof(f9851,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl11 | $spl64)),
  inference(subsumption_resolution,[],[f9850,f1162])).
fof(f9850,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl64),
  inference(forward_demodulation,[],[f9849,f4877])).
fof(f9849,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f9848,f822])).
fof(f9848,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f9664,f822])).
fof(f9664,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(superposition,[],[f786,f4877])).
fof(f9836,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ($spl35 | $spl64)),
  inference(subsumption_resolution,[],[f9621,f3491])).
fof(f9621,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(superposition,[],[f786,f4877])).
fof(f9598,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl64),
  inference(forward_demodulation,[],[f9597,f841])).
fof(f9597,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl64),
  inference(forward_demodulation,[],[f9424,f841])).
fof(f9424,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl64),
  inference(superposition,[],[f839,f2795])).
fof(f9488,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ($spl35 | $spl64)),
  inference(subsumption_resolution,[],[f9487,f3491])).
fof(f9487,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl64),
  inference(forward_demodulation,[],[f9486,f2795])).
fof(f9486,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f9485,f821])).
fof(f9485,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f9296,f822])).
fof(f9296,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(superposition,[],[f786,f2795])).
fof(f9482,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl35 | $spl64)),
  inference(subsumption_resolution,[],[f9481,f3491])).
fof(f9481,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl64),
  inference(forward_demodulation,[],[f9480,f2795])).
fof(f9480,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f9479,f822])).
fof(f9479,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f9295,f822])).
fof(f9295,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(superposition,[],[f786,f2795])).
fof(f9467,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | ($spl11 | $spl64)),
  inference(subsumption_resolution,[],[f9252,f1162])).
fof(f9252,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(superposition,[],[f786,f2795])).
fof(f9034,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl3),
  inference(resolution,[],[f879,f809])).
fof(f8881,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | ($spl3 | $spl64)),
  inference(resolution,[],[f4833,f800])).
fof(f4833,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | ($spl3 | $spl64)),
  inference(backward_demodulation,[],[f2795,f879])).
fof(f8880,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl3 | $spl64)),
  inference(resolution,[],[f4833,f803])).
fof(f8875,plain,(
  $false | ($spl2 | $spl39)),
  inference(subsumption_resolution,[],[f1502,f1019])).
fof(f8867,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl29),
  inference(forward_demodulation,[],[f8739,f856])).
fof(f8739,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK14)))) | $spl29),
  inference(resolution,[],[f1482,f800])).
fof(f8736,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl29),
  inference(resolution,[],[f1482,f809])).
fof(f8607,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f808])).
fof(f8402,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X1))) ) | $spl2),
  inference(subsumption_resolution,[],[f8393,f758])).
fof(f8393,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X1))) ) | $spl2),
  inference(resolution,[],[f3408,f704])).
fof(f8401,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f8392,f758])).
fof(f8392,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,X0))) ) | $spl2),
  inference(resolution,[],[f3408,f708])).
fof(f8372,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f8363,f758])).
fof(f8363,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(resolution,[],[f3407,f704])).
fof(f8371,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f8362,f758])).
fof(f8362,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(resolution,[],[f3407,f708])).
fof(f7772,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f1019,f808])).
fof(f7037,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl12),
  inference(resolution,[],[f1164,f808])).
fof(f6888,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl5),
  inference(resolution,[],[f881,f803])).
fof(f6877,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | p(s(bool,X1))) ) | $spl5),
  inference(subsumption_resolution,[],[f6871,f758])).
fof(f6871,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | p(s(bool,X1))) ) | $spl5),
  inference(resolution,[],[f5575,f704])).
fof(f6876,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | p(s(bool,X0))) ) | $spl5),
  inference(subsumption_resolution,[],[f6870,f758])).
fof(f6870,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | p(s(bool,X0))) ) | $spl5),
  inference(resolution,[],[f5575,f708])).
fof(f6869,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f6863,f758])).
fof(f6863,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) ) | $spl5),
  inference(resolution,[],[f5574,f704])).
fof(f6868,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f6862,f758])).
fof(f6862,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) ) | $spl5),
  inference(resolution,[],[f5574,f708])).
fof(f6857,plain,(
  $false | ($spl4 | $spl95)),
  inference(subsumption_resolution,[],[f891,f5340])).
fof(f6851,plain,(
  $false | ($spl16 | $spl95)),
  inference(subsumption_resolution,[],[f1316,f6069])).
fof(f6069,plain,(
  s('type.integer.int',sK13) != s('type.integer.int',sK14) | $spl95),
  inference(resolution,[],[f5340,f811])).
fof(f6847,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f6833,f758])).
fof(f6833,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(resolution,[],[f3038,f704])).
fof(f6845,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f6832,f758])).
fof(f6832,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))))) ) | $spl2),
  inference(resolution,[],[f3038,f708])).
fof(f6489,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f1019,f808])).
fof(f6219,plain,(
  $false | ($spl5 | $spl15)),
  inference(subsumption_resolution,[],[f1314,f5216])).
fof(f5834,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl12),
  inference(resolution,[],[f1164,f808])).
fof(f5825,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(subsumption_resolution,[],[f5824,f819])).
fof(f5824,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f5823,f822])).
fof(f5823,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(forward_demodulation,[],[f5822,f821])).
fof(f5822,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(forward_demodulation,[],[f5753,f822])).
fof(f5753,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(superposition,[],[f876,f786])).
fof(f5820,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f5759,f845])).
fof(f5759,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl0),
  inference(superposition,[],[f876,f833])).
fof(f5760,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(superposition,[],[f876,f835])).
fof(f5761,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(superposition,[],[f876,f836])).
fof(f5762,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(superposition,[],[f876,f837])).
fof(f5763,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(superposition,[],[f876,f838])).
fof(f5819,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f5818,f856])).
fof(f5818,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f5764,f844])).
fof(f5764,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(superposition,[],[f876,f848])).
fof(f5808,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl0 | $spl41)),
  inference(subsumption_resolution,[],[f5763,f1504])).
fof(f5807,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl0 | $spl41)),
  inference(subsumption_resolution,[],[f5762,f1504])).
fof(f5806,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl0 | $spl41)),
  inference(subsumption_resolution,[],[f5761,f1504])).
fof(f5805,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl0 | $spl41)),
  inference(subsumption_resolution,[],[f5760,f1504])).
fof(f5789,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl0),
  inference(forward_demodulation,[],[f5646,f845])).
fof(f5646,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK14)))) | $spl0),
  inference(resolution,[],[f876,f808])).
fof(f5495,plain,(
  s('type.integer.int',sK13) = s('type.integer.int',sK14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl5),
  inference(resolution,[],[f5216,f809])).
fof(f5214,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl5),
  inference(resolution,[],[f881,f803])).
fof(f5212,plain,(
  s('type.integer.int',sK13) = s('type.integer.int',sK14) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | $spl5),
  inference(resolution,[],[f881,f809])).
fof(f2502,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl7),
  inference(resolution,[],[f1011,f800])).
fof(f2642,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl7),
  inference(forward_demodulation,[],[f2505,f845])).
fof(f2505,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl7),
  inference(resolution,[],[f1011,f803])).
fof(f5209,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl7),
  inference(forward_demodulation,[],[f2643,f844])).
fof(f2643,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK14)))) | $spl7),
  inference(forward_demodulation,[],[f2506,f817])).
fof(f2506,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK14)))) | $spl7),
  inference(resolution,[],[f1011,f800])).
fof(f5190,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl64),
  inference(forward_demodulation,[],[f5189,f841])).
fof(f5189,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl64),
  inference(forward_demodulation,[],[f5015,f841])).
fof(f5015,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl64),
  inference(superposition,[],[f839,f2795])).
fof(f5080,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ($spl35 | $spl64)),
  inference(subsumption_resolution,[],[f5079,f3491])).
fof(f5079,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl64),
  inference(forward_demodulation,[],[f5078,f2795])).
fof(f5078,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f5077,f821])).
fof(f5077,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(forward_demodulation,[],[f4887,f822])).
fof(f4887,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(superposition,[],[f786,f2795])).
fof(f5058,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | ($spl11 | $spl64)),
  inference(subsumption_resolution,[],[f4843,f1162])).
fof(f4843,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl64),
  inference(superposition,[],[f786,f2795])).
fof(f4697,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl3),
  inference(resolution,[],[f4551,f809])).
fof(f4547,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl3),
  inference(resolution,[],[f879,f809])).
fof(f4402,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl4),
  inference(subsumption_resolution,[],[f4151,f758])).
fof(f4151,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl4),
  inference(superposition,[],[f800,f1279])).
fof(f4405,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl4),
  inference(subsumption_resolution,[],[f4153,f758])).
fof(f4153,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl4),
  inference(superposition,[],[f803,f1279])).
fof(f3908,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl4),
  inference(resolution,[],[f1187,f803])).
fof(f1187,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f882,f806])).
fof(f4544,plain,(
  $false | ($spl2 | $spl19)),
  inference(subsumption_resolution,[],[f1318,f1156])).
fof(f4042,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl4),
  inference(forward_demodulation,[],[f3909,f856])).
fof(f3909,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f1187,f800])).
fof(f3906,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f1187,f809])).
fof(f3890,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl41),
  inference(forward_demodulation,[],[f3889,f821])).
fof(f3889,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl41),
  inference(forward_demodulation,[],[f3731,f846])).
fof(f3731,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl41),
  inference(resolution,[],[f1504,f800])).
fof(f3487,plain,(
  ( ! [X0] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0)) ) | $spl72),
  inference(cnf_transformation,[],[f3487_D])).
fof(f3487_D,plain,(
  ( ! [X0] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3482,plain,(
  ( ! [X0] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14))))) ) | $spl2),
  inference(superposition,[],[f3408,f853])).
fof(f3475,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) ) | $spl70),
  inference(cnf_transformation,[],[f3475_D])).
fof(f3475_D,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3468,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14))))) ) | $spl2),
  inference(superposition,[],[f3407,f853])).
fof(f3329,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1156,f808])).
fof(f2793,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl63),
  inference(cnf_transformation,[],[f2793_D])).
fof(f2793_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f2657,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1020,f809])).
fof(f2503,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl7),
  inference(resolution,[],[f1011,f809])).
fof(f2163,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f889,f808])).
fof(f889,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl4),
  inference(resolution,[],[f880,f804])).
fof(f1504,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl41),
  inference(cnf_transformation,[],[f1504_D])).
fof(f1504_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f1530,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f1529,f856])).
fof(f1529,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f1458,f844])).
fof(f1458,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl1),
  inference(superposition,[],[f877,f850])).
fof(f1526,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f1525,f856])).
fof(f1525,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f1457,f844])).
fof(f1457,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(superposition,[],[f877,f848])).
fof(f1456,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(superposition,[],[f877,f838])).
fof(f1455,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(superposition,[],[f877,f837])).
fof(f1454,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(superposition,[],[f877,f836])).
fof(f1453,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK26('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(superposition,[],[f877,f835])).
fof(f1514,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f1452,f845])).
fof(f1452,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl1),
  inference(superposition,[],[f877,f833])).
fof(f1513,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f1451,f845])).
fof(f1451,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl1),
  inference(superposition,[],[f877,f831])).
fof(f1508,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(subsumption_resolution,[],[f1507,f819])).
fof(f1507,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f1506,f822])).
fof(f1506,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(forward_demodulation,[],[f1505,f821])).
fof(f1505,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(forward_demodulation,[],[f1446,f822])).
fof(f1446,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(superposition,[],[f877,f786])).
fof(f1443,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f877,f832])).
fof(f1442,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f877,f834])).
fof(f1319,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl21),
  inference(cnf_transformation,[],[f1319_D])).
fof(f1319_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1439,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl1),
  inference(superposition,[],[f877,f853])).
fof(f1492,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl1),
  inference(forward_demodulation,[],[f1491,f817])).
fof(f1491,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(forward_demodulation,[],[f1336,f845])).
fof(f1336,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(resolution,[],[f877,f800])).
fof(f1333,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(resolution,[],[f877,f809])).
fof(f1332,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(resolution,[],[f877,f803])).
fof(f1189,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f882,f800])).
fof(f1188,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl4),
  inference(resolution,[],[f882,f803])).
fof(f1314,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl15),
  inference(cnf_transformation,[],[f1314_D])).
fof(f1314_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1186,plain,(
  s('type.integer.int',sK13) = s('type.integer.int',sK14) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f882,f809])).
fof(f1128,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(superposition,[],[f878,f834])).
fof(f890,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f880,f808])).
fof(f572,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))),
  inference(cnf_transformation,[],[f447])).
fof(f447,plain,(
  (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14])],[f446])).
fof(f446,plain,(
  ? [X0,X1] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))),
  inference(flattening,[],[f445])).
fof(f445,plain,(
  ? [X0,X1] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))))),
  inference(nnf_transformation,[],[f398])).
fof(f398,plain,(
  ? [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) <~> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LT|split|3','thm.integer.INT_ABS_LT|split|3')).
fof(f573,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))),
  inference(cnf_transformation,[],[f447])).
fof(f877,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl1),
  inference(cnf_transformation,[],[f877_D])).
fof(f877_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f879,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl3),
  inference(cnf_transformation,[],[f879_D])).
fof(f879_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f881,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | $spl5),
  inference(cnf_transformation,[],[f881_D])).
fof(f881_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f574,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))))),
  inference(cnf_transformation,[],[f447])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1504
% Passive clauses: 37553
% Generated clauses: 72895
% Final active clauses: 580
% Final passive clauses: 14533
% Input formulas: 261
% Initial clauses: 297
% 
% Selected by SInE selection: 104
% SInE iterations: 1
% 
% Duplicate literals: 1368
% Trivial inequalities: 274
% Fw subsumption resolutions: 4683
% Fw demodulations: 6786
% Bw demodulations: 279
% 
% Simple tautologies: 1825
% Equational tautologies: 12
% Forward subsumptions: 18201
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 25785
% Factoring: 2
% Forward superposition: 18512
% Backward superposition: 10651
% Self superposition: 42
% Equality resolution: 5
% 
% Split clauses: 558
% Split components: 1181
% SAT solver clauses: 721
% SAT solver unit clauses: 28
% SAT solver binary clauses: 196
% 
% Sat splits: 558
% Sat splitting refutations: 162
% 
% Memory used [KB]: 24562
% Time elapsed: 1.185 s
% ------------------------------
% ------------------------------
% Success in time 12.951 s
