Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.combin
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.combin
% SZS output start Proof for thm.combin
fof(f4640,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4629,f104,f111])).
fof(f111,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f111_D])).
fof(f111_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f104,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f89])).
fof(f89,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X15] : ! [X10] : (s(X15,X10) = s(X15,X10) <=> p(s(bool,'T_0')))),
  file('thm.combin.UPDATE_APPLY|split|1','thm.bool.REFL_CLAUSE')).
fof(f4629,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f4628,f67])).
fof(f67,plain,(
  s(sK1,sK2) != s(sK1,sK3)),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  s(sK1,sK2) != s(sK1,sK3) & s(sK0,happ(s(fun(sK1,sK0),sK5),s(sK1,sK3))) != s(sK0,'const.combin.UPDATE_4'(s(sK1,sK2),s(sK0,sK4),s(fun(sK1,sK0),sK5),s(sK1,sK3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f44])).
fof(f44,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s(X1,X2) != s(X1,X3) & s(X0,happ(s(fun(X1,X0),X5),s(X1,X3))) != s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X4),s(fun(X1,X0),X5),s(X1,X3))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s(X1,X2) != s(X1,X3) => s(X0,happ(s(fun(X1,X0),X5),s(X1,X3))) = s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X4),s(fun(X1,X0),X5),s(X1,X3))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (~s(X1,X2) = s(X1,X3) => s(X0,happ(s(fun(X1,X0),X5),s(X1,X3))) = s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X4),s(fun(X1,X0),X5),s(X1,X3))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (~s(X7,X8) = s(X7,X9) => s(X6,happ(s(fun(X7,X6),X11),s(X7,X9))) = s(X6,'const.combin.UPDATE_4'(s(X7,X8),s(X6,X10),s(fun(X7,X6),X11),s(X7,X9))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (~s(X7,X8) = s(X7,X9) => s(X6,happ(s(fun(X7,X6),X11),s(X7,X9))) = s(X6,'const.combin.UPDATE_4'(s(X7,X8),s(X6,X10),s(fun(X7,X6),X11),s(X7,X9))))),
  file('thm.combin.UPDATE_APPLY|split|1','thm.combin.UPDATE_APPLY|split|1')).
fof(f4628,plain,(
  s(sK1,sK2) = s(sK1,sK3) | $spl6),
  inference(subsumption_resolution,[],[f4566,f111])).
fof(f4566,plain,(
  ~p(s(bool,'T_0')) | s(sK1,sK2) = s(sK1,sK3)),
  inference(superposition,[],[f101,f4542])).
fof(f4542,plain,(
  s(bool,'T_0') = s(bool,sK7(sK3,sK5,sK4,sK2,sK1,sK0))),
  inference(trivial_inequality_removal,[],[f4541])).
fof(f4541,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK5),s(sK1,sK3))) != s(sK0,happ(s(fun(sK1,sK0),sK5),s(sK1,sK3))) | s(bool,'T_0') = s(bool,sK7(sK3,sK5,sK4,sK2,sK1,sK0))),
  inference(superposition,[],[f68,f1705])).
fof(f1705,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(X22,happ(s(fun(X21,X22),X18),s(X21,X17))) = s(X22,'const.combin.UPDATE_4'(s(X21,X20),s(X22,X19),s(fun(X21,X22),X18),s(X21,X17))) | s(bool,'T_0') = s(bool,sK7(X17,X18,X19,X20,X21,X22))) )),
  inference(forward_demodulation,[],[f1579,f93])).
fof(f93,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X15] : ! [X12,X13] : s(X15,X13) = s(X15,'const.bool.COND_3'(s(bool,'F_0'),s(X15,X12),s(X15,X13)))),
  file('thm.combin.UPDATE_APPLY|split|1','thm.bool.COND_CLAUSES|split|1')).
fof(f1579,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(X22,'const.combin.UPDATE_4'(s(X21,X20),s(X22,X19),s(fun(X21,X22),X18),s(X21,X17))) = s(X22,'const.bool.COND_3'(s(bool,'F_0'),s(X22,X19),s(X22,happ(s(fun(X21,X22),X18),s(X21,X17))))) | s(bool,'T_0') = s(bool,sK7(X17,X18,X19,X20,X21,X22))) )),
  inference(superposition,[],[f103,f72])).
fof(f72,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.combin.UPDATE_APPLY|split|1',aHL_BOOL_CASES)).
fof(f103,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,sK7(X5,X4,X3,X2,X1,X0)),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,sK7(X5,X4,X3,X2,X1,X0))) | s(X1,X2) = s(X1,X5)) & (s(X1,X2) != s(X1,X5) | p(s(bool,sK7(X5,X4,X3,X2,X1,X0)))) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,sK7(X5,X4,X3,X2,X1,X0)),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((~p(s(bool,X6)) | s(X1,X2) = s(X1,X5)) & (s(X1,X2) != s(X1,X5) | p(s(bool,X6))) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : (((~p(s(bool,X6)) | s(X1,X2) = s(X1,X5)) & (s(X1,X2) != s(X1,X5) | p(s(bool,X6)))) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((p(s(bool,X6)) <=> s(X1,X2) = s(X1,X5)) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : ? [X6] : ((p(s(bool,X6)) <=> s(X1,X2) = s(X1,X5)) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X16,X15] : ! [X8,X9,X5,X17] : ? [X18] : ((p(s(bool,X18)) <=> s(X15,X8) = s(X15,X17)) & s(X16,'const.combin.UPDATE_4'(s(X15,X8),s(X16,X9),s(fun(X15,X16),X5),s(X15,X17))) = s(X16,'const.bool.COND_3'(s(bool,X18),s(X16,X9),s(X16,happ(s(fun(X15,X16),X5),s(X15,X17))))))),
  file('thm.combin.UPDATE_APPLY|split|1','thm.combin.UPDATE_def')).
fof(f68,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK5),s(sK1,sK3))) != s(sK0,'const.combin.UPDATE_4'(s(sK1,sK2),s(sK0,sK4),s(fun(sK1,sK0),sK5),s(sK1,sK3)))),
  inference(cnf_transformation,[],[f49])).
fof(f101,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK7(X5,X4,X3,X2,X1,X0))) | s(X1,X2) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f66])).
% SZS output end Proof for thm.combin
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 63
% Passive clauses: 1236
% Generated clauses: 4442
% Final active clauses: 46
% Final passive clauses: 1058
% Input formulas: 19
% Initial clauses: 37
% 
% Duplicate literals: 256
% Trivial inequalities: 35
% Fw subsumption resolutions: 360
% Fw demodulations: 278
% Bw demodulations: 14
% 
% Simple tautologies: 66
% Equational tautologies: 34
% Forward subsumptions: 2114
% 
% Binary resolution: 80
% Factoring: 1
% Forward superposition: 1686
% Backward superposition: 1500
% Self superposition: 99
% Equality factoring: 58
% Equality resolution: 2
% 
% Split clauses: 49
% Split components: 107
% SAT solver clauses: 110
% SAT solver unit clauses: 5
% SAT solver binary clauses: 34
% SAT solver learnt clauses: 1
% 
% Sat splits: 49
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 107
% TWLsolver calls for satisfiability: 36
% 
% Memory used [KB]: 3709
% Time elapsed: 0.241 s
% ------------------------------
% ------------------------------
% Success in time 0.249 s
