Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integerRing
% (73674)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 154
% Passive clauses: 3859
% Generated clauses: 19214
% Final active clauses: 146
% Final passive clauses: 3573
% Input formulas: 123
% Initial clauses: 128
% 
% Duplicate literals: 530
% Trivial inequalities: 304
% Fw subsumption resolutions: 828
% Fw demodulations: 1149
% Bw demodulations: 4
% 
% Simple tautologies: 265
% Equational tautologies: 22
% Forward subsumptions: 11949
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 31
% Forward superposition: 8483
% Backward superposition: 7480
% Self superposition: 211
% Equality factoring: 54
% Equality resolution: 22
% 
% Split clauses: 213
% Split components: 426
% SAT solver clauses: 410
% SAT solver unit clauses: 5
% SAT solver binary clauses: 99
% 
% Sat splits: 213
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 410
% TWLsolver calls for satisfiability: 50
% 
% Memory used [KB]: 20596
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integerRing
% (73698)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 126
% Passive clauses: 1070
% Generated clauses: 16560
% Final active clauses: 125
% Final passive clauses: 931
% Input formulas: 123
% Initial clauses: 128
% 
% Duplicate literals: 516
% Trivial inequalities: 296
% Fw subsumption resolutions: 533
% Fw demodulations: 1072
% Bw demodulations: 13
% Condensations: 672
% 
% Simple tautologies: 163
% Equational tautologies: 22
% Forward subsumptions: 9791
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 24
% Forward superposition: 6200
% Backward superposition: 5499
% Self superposition: 186
% Equality factoring: 52
% Equality resolution: 16
% 
% Split clauses: 146
% Split components: 292
% SAT solver clauses: 2263
% SAT solver unit clauses: 245
% SAT solver binary clauses: 40
% 
% Sat splits: 146
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2263
% TWLsolver calls for satisfiability: 261
% 
% Memory used [KB]: 23411
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integerRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integerRing
% SZS output start Proof for thm.integerRing
fof(f127290,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f127259])).
fof(f127259,plain,(
  s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pconst_1'(s('type.integer.int',sK0))))) != s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pconst_1'(s('type.integer.int',sK0)))))),
  inference(superposition,[],[f321,f115426])).
fof(f115426,plain,(
  ( ! [X1] : (s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pconst_1'(s('type.integer.int',X1))))) = s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_monom_3'(s('type.integer.int',X1),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))) )),
  inference(superposition,[],[f25752,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X2)))))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X2)))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X2)))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X7] : ! [X19,X6] : s('type.canonical.canonical_sum'(X7),'const.canonical.Cons_monom_3'(s(X7,X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X7),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X7),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X7),X19),s('type.ringNorm.polynom'(X7),'const.ringNorm.Pconst_1'(s(X7,X6)))))),
  file('thm.integerRing.int_ring_thms|split|52','thm.ringNorm.polynom_normalize_def|split|1')).
fof(f25752,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),X0))) = s('type.canonical.canonical_sum'('type.integer.int'),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'('type.integer.int'),'const.ring.recordtype|46|ring_5'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_add_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_mul_0'),s(fun('type.integer.int','type.integer.int'),'const.integer.int_neg_0'))),s('type.ringNorm.polynom'('type.integer.int'),X0)))) )),
  inference(forward_demodulation,[],[f25719,f343])).
fof(f343,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),X0))) = s('type.canonical.canonical_sum'('type.integer.int'),happ(s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.integerRing.int_polynom_normalize_0'),s('type.ringNorm.polynom'('type.integer.int'),X0)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),X0))) = s('type.canonical.canonical_sum'('type.integer.int'),happ(s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.integerRing.int_polynom_normalize_0'),s('type.ringNorm.polynom'('type.integer.int'),X0)))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X47] : s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),X47))) = s('type.canonical.canonical_sum'('type.integer.int'),happ(s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.integerRing.int_polynom_normalize_0'),s('type.ringNorm.polynom'('type.integer.int'),X47)))),
  file('thm.integerRing.int_ring_thms|split|52','ARITY_0_1_const.integerRing.int_polynom_normalize')).
fof(f25719,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.integer.int'),happ(s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.integerRing.int_polynom_normalize_0'),s('type.ringNorm.polynom'('type.integer.int'),X0))) = s('type.canonical.canonical_sum'('type.integer.int'),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'('type.integer.int'),'const.ring.recordtype|46|ring_5'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_add_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_mul_0'),s(fun('type.integer.int','type.integer.int'),'const.integer.int_neg_0'))),s('type.ringNorm.polynom'('type.integer.int'),X0)))) )),
  inference(superposition,[],[f360,f325])).
fof(f325,plain,(
  s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.integerRing.int_polynom_normalize_0') = s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.ringNorm.polynom_normalize_1'(s('type.ring.ring'('type.integer.int'),'const.ring.recordtype|46|ring_5'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_add_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_mul_0'),s(fun('type.integer.int','type.integer.int'),'const.integer.int_neg_0')))))),
  inference(cnf_transformation,[],[f11])).
fof(f11,axiom,(
  s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.integerRing.int_polynom_normalize_0') = s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.ringNorm.polynom_normalize_1'(s('type.ring.ring'('type.integer.int'),'const.ring.recordtype|46|ring_5'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_add_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_mul_0'),s(fun('type.integer.int','type.integer.int'),'const.integer.int_neg_0')))))),
  file('thm.integerRing.int_ring_thms|split|52','thm.integerRing.int_polynom_normalize_def')).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.polynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.polynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.polynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X7] : ! [X47,X48] : s('type.canonical.canonical_sum'(X7),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X7),X47),s('type.ringNorm.polynom'(X7),X48))) = s('type.canonical.canonical_sum'(X7),happ(s(fun('type.ringNorm.polynom'(X7),'type.canonical.canonical_sum'(X7)),'const.ringNorm.polynom_normalize_1'(s('type.ring.ring'(X7),X47))),s('type.ringNorm.polynom'(X7),X48)))),
  file('thm.integerRing.int_ring_thms|split|52','ARITY_1_2_const.ringNorm.polynom_normalize')).
fof(f321,plain,(
  s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pconst_1'(s('type.integer.int',sK0))))) != s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_monom_3'(s('type.integer.int',sK0),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pconst_1'(s('type.integer.int',sK0))))) != s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_monom_3'(s('type.integer.int',sK0),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f312])).
fof(f312,plain,(
  ? [X0] : s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pconst_1'(s('type.integer.int',X0))))) != s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_monom_3'(s('type.integer.int',X0),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ~! [X0] : s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pconst_1'(s('type.integer.int',X0))))) = s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_monom_3'(s('type.integer.int',X0),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pconst_1'(s('type.integer.int',X6))))) = s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_monom_3'(s('type.integer.int',X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pconst_1'(s('type.integer.int',X6))))) = s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_monom_3'(s('type.integer.int',X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))),
  file('thm.integerRing.int_ring_thms|split|52','thm.integerRing.int_ring_thms|split|52')).
% SZS output end Proof for thm.integerRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 473
% Passive clauses: 7618
% Generated clauses: 125830
% Final active clauses: 293
% Final passive clauses: 1851
% Input formulas: 123
% Initial clauses: 128
% Discarded non-redundant clauses: 15411
% 
% Duplicate literals: 2323
% Trivial inequalities: 698
% Fw subsumption resolutions: 1310
% Fw demodulations: 22482
% Bw demodulations: 607
% Condensations: 2207
% Global subsumptions: 139
% 
% Simple tautologies: 434
% Equational tautologies: 449
% Forward subsumptions: 87568
% Fw demodulations to eq. taut.: 8
% Bw demodulations to eq. taut.: 49
% 
% Unit resulting resolution: 15937
% Forward superposition: 42219
% Backward superposition: 36367
% Self superposition: 487
% Equality factoring: 126
% Equality resolution: 33
% 
% Split clauses: 185
% Split components: 375
% SAT solver clauses: 8317
% SAT solver unit clauses: 661
% SAT solver binary clauses: 1135
% 
% Sat splits: 185
% Sat splitting refutations: 188
% 
% Memory used [KB]: 160167
% Time elapsed: 6.591 s
% ------------------------------
% ------------------------------
% Success in time 7.939 s
