Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.combin
% (112630)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 120
% Passive clauses: 3370
% Generated clauses: 20609
% Final active clauses: 116
% Final passive clauses: 3094
% Input formulas: 18
% Initial clauses: 31
% 
% Duplicate literals: 983
% Trivial inequalities: 240
% Fw subsumption resolutions: 2462
% Fw demodulations: 3521
% Bw demodulations: 2
% 
% Simple tautologies: 322
% Equational tautologies: 33
% Forward subsumptions: 8879
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 6
% Forward superposition: 7653
% Backward superposition: 5327
% Self superposition: 192
% Equality factoring: 72
% Equality resolution: 37
% 
% Split clauses: 638
% Split components: 1304
% SAT solver clauses: 924
% SAT solver unit clauses: 13
% SAT solver binary clauses: 89
% 
% Sat splits: 638
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 910
% TWLsolver calls for satisfiability: 107
% 
% Memory used [KB]: 14200
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.combin
% (112647)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 99
% Passive clauses: 1400
% Generated clauses: 14488
% Final active clauses: 95
% Final passive clauses: 1274
% Input formulas: 18
% Initial clauses: 31
% 
% Duplicate literals: 757
% Trivial inequalities: 135
% Fw subsumption resolutions: 763
% Fw demodulations: 2303
% Bw demodulations: 27
% Condensations: 902
% Global subsumptions: 4
% 
% Simple tautologies: 175
% Equational tautologies: 31
% Forward subsumptions: 5932
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 7
% Forward superposition: 5384
% Backward superposition: 2569
% Self superposition: 132
% Equality factoring: 49
% Equality resolution: 18
% 
% Split clauses: 58
% Split components: 124
% SAT solver clauses: 2113
% SAT solver unit clauses: 38
% SAT solver binary clauses: 34
% SAT solver learnt clauses: 8
% 
% Sat splits: 58
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2097
% TWLsolver calls for satisfiability: 141
% 
% Memory used [KB]: 17782
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.combin
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.combin
% SZS output start Proof for thm.combin
fof(f2103,plain,(
  $false),
  inference(subsumption_resolution,[],[f2102,f463])).
fof(f463,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X1))) = s(X2,'const.combin.o_3'(s(fun(X0,X2),X3),s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(superposition,[],[f92,f96])).
fof(f96,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f84,f83])).
fof(f83,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X10] : s(X6,X10) = s(X6,'const.combin.I_1'(s(X6,X10)))),
  file('thm.combin.I_o_ID|split|1','thm.combin.I_THM')).
fof(f84,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X13] : s(X6,'const.combin.I_1'(s(X6,X13))) = s(X6,happ(s(fun(X6,X6),'const.combin.I_0'),s(X6,X13)))),
  file('thm.combin.I_o_ID|split|1','ARITY_0_1_const.combin.I')).
fof(f92,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7,X6,X12] : ! [X8,X11,X10] : s(X7,'const.combin.o_3'(s(fun(X6,X7),X8),s(fun(X12,X6),X11),s(X12,X10))) = s(X7,happ(s(fun(X6,X7),X8),s(X6,happ(s(fun(X12,X6),X11),s(X12,X10)))))),
  file('thm.combin.I_o_ID|split|1','thm.combin.o_THM')).
fof(f2102,plain,(
  s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK1),sK2),s(fun(sK0,sK0),'const.combin.I_0')),sK2,sK0,sK1)))) != s(sK1,'const.combin.o_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK1),sK2),s(fun(sK0,sK0),'const.combin.I_0')),sK2,sK0,sK1))))),
  inference(superposition,[],[f1119,f93])).
fof(f93,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X7,X12,X6] : ! [X13,X14,X15] : s(X7,'const.combin.o_3'(s(fun(X12,X7),X13),s(fun(X6,X12),X14),s(X6,X15))) = s(X7,happ(s(fun(X6,X7),'const.combin.o_2'(s(fun(X12,X7),X13),s(fun(X6,X12),X14))),s(X6,X15)))),
  file('thm.combin.I_o_ID|split|1','ARITY_2_3_const.combin.o')).
fof(f1119,plain,(
  s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK1),sK2),s(fun(sK0,sK0),'const.combin.I_0')),sK2,sK0,sK1)))) != s(sK1,happ(s(fun(sK0,sK1),'const.combin.o_2'(s(fun(sK0,sK1),sK2),s(fun(sK0,sK0),'const.combin.I_0'))),s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK1),sK2),s(fun(sK0,sK0),'const.combin.I_0')),sK2,sK0,sK1))))),
  inference(unit_resulting_resolution,[],[f63,f91])).
fof(f91,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK4(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK4(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK4(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK4(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X7,X6] : ! [X8,X11] : (s(fun(X6,X7),X8) = s(fun(X6,X7),X11) <=> ! [X10] : s(X7,happ(s(fun(X6,X7),X8),s(X6,X10))) = s(X7,happ(s(fun(X6,X7),X11),s(X6,X10))))),
  file('thm.combin.I_o_ID|split|1','thm.bool.FUN_EQ_THM')).
fof(f63,plain,(
  s(fun(sK0,sK1),sK2) != s(fun(sK0,sK1),'const.combin.o_2'(s(fun(sK0,sK1),sK2),s(fun(sK0,sK0),'const.combin.I_0')))),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  s(fun(sK0,sK1),sK2) != s(fun(sK0,sK1),'const.combin.o_2'(s(fun(sK0,sK1),sK2),s(fun(sK0,sK0),'const.combin.I_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f43])).
fof(f43,plain,(
  ? [X0,X1,X2] : s(fun(X0,X1),X2) != s(fun(X0,X1),'const.combin.o_2'(s(fun(X0,X1),X2),s(fun(X0,X0),'const.combin.I_0')))),
  inference(ennf_transformation,[],[f21])).
fof(f21,plain,(
  ~! [X0,X1,X2] : s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.o_2'(s(fun(X0,X1),X2),s(fun(X0,X0),'const.combin.I_0')))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ~! [X0,X1] : ! [X2] : s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.o_2'(s(fun(X0,X1),X2),s(fun(X0,X0),'const.combin.I_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8] : s(fun(X6,X7),X8) = s(fun(X6,X7),'const.combin.o_2'(s(fun(X6,X7),X8),s(fun(X6,X6),'const.combin.I_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8] : s(fun(X6,X7),X8) = s(fun(X6,X7),'const.combin.o_2'(s(fun(X6,X7),X8),s(fun(X6,X6),'const.combin.I_0')))),
  file('thm.combin.I_o_ID|split|1','thm.combin.I_o_ID|split|1')).
% SZS output end Proof for thm.combin
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 32
% Passive clauses: 157
% Generated clauses: 1979
% Final active clauses: 32
% Final passive clauses: 125
% Input formulas: 18
% Initial clauses: 31
% 
% Duplicate literals: 187
% Trivial inequalities: 30
% Fw subsumption resolutions: 44
% Fw demodulations: 183
% Condensations: 256
% 
% Simple tautologies: 58
% Equational tautologies: 42
% Forward subsumptions: 1008
% Fw demodulations to eq. taut.: 12
% 
% Unit resulting resolution: 78
% Forward superposition: 595
% Backward superposition: 457
% Self superposition: 68
% Equality factoring: 44
% Equality resolution: 6
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 165
% SAT solver unit clauses: 16
% SAT solver binary clauses: 59
% 
% Sat splits: 3
% 
% Memory used [KB]: 2046
% Time elapsed: 0.090 s
% ------------------------------
% ------------------------------
% Success in time 1.408 s
