Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (56404)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 127
% Passive clauses: 1434
% Generated clauses: 24816
% Final active clauses: 123
% Final passive clauses: 1223
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 1664
% Trivial inequalities: 35
% Fw subsumption resolutions: 2589
% Fw demodulations: 4417
% Bw demodulations: 2
% 
% Simple tautologies: 212
% Equational tautologies: 18
% Forward subsumptions: 13281
% 
% Binary resolution: 12
% Forward superposition: 9673
% Backward superposition: 6272
% Self superposition: 170
% Equality factoring: 34
% Equality resolution: 4
% 
% Split clauses: 670
% Split components: 1345
% SAT solver clauses: 1247
% SAT solver unit clauses: 1
% SAT solver binary clauses: 23
% 
% Sat splits: 670
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1247
% TWLsolver calls for satisfiability: 92
% 
% Memory used [KB]: 16375
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (56432)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 107
% Passive clauses: 887
% Generated clauses: 15642
% Final active clauses: 103
% Final passive clauses: 771
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 775
% Trivial inequalities: 27
% Fw subsumption resolutions: 670
% Fw demodulations: 1925
% Bw demodulations: 9
% Condensations: 1139
% 
% Simple tautologies: 132
% Equational tautologies: 17
% Forward subsumptions: 8901
% 
% Binary resolution: 9
% Forward superposition: 6187
% Backward superposition: 4157
% Self superposition: 148
% Equality factoring: 32
% Equality resolution: 3
% 
% Split clauses: 12
% Split components: 29
% SAT solver clauses: 1188
% SAT solver unit clauses: 38
% SAT solver binary clauses: 16
% 
% Sat splits: 12
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1188
% TWLsolver calls for satisfiability: 121
% 
% Memory used [KB]: 27121
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f58484,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f58433])).
fof(f58433,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK3),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) != s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK3),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(superposition,[],[f65,f11529])).
fof(f11529,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.canonical.canonical_sum'(X5),'const.canonical.Cons_monom_3'(s(X5,X7),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X5),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X5),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(X5),X6),s(X5,X7),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X5),'const.canonical.Nil_monom_0')))) )),
  inference(superposition,[],[f2091,f73])).
fof(f73,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.canonical.monom_insert_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.canonical.monom_insert_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.canonical.monom_insert_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X13,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X9),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X6),'const.canonical.monom_insert_4'(s('type.semi_ring.semi_ring'(X6),X13),s(X6,X9),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.monom_insert_def|split|2','thm.canonical.monom_insert_def|split|2')).
fof(f2091,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),'const.canonical.monom_insert_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4)))) )),
  inference(forward_demodulation,[],[f2024,f83])).
fof(f83,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_monom_insert_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_monom_insert_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_monom_insert_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X20,X21,X22,X23] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(X6),X20),s(X6,X21),s('type.list.list'('type.quote.index'),X22),s('type.canonical.canonical_sum'(X6),X23))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),happ(s(fun(X6,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)))),'const.ringNorm.r_monom_insert_1'(s('type.ring.ring'(X6),X20))),s(X6,X21))),s('type.list.list'('type.quote.index'),X22))),s('type.canonical.canonical_sum'(X6),X23)))),
  file('thm.ringNorm.monom_insert_def|split|2','ARITY_1_4_const.ringNorm.r_monom_insert')).
fof(f2024,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.monom_insert_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_monom_insert_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))) )),
  inference(superposition,[],[f82,f69])).
fof(f69,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_monom_insert_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.monom_insert_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_monom_insert_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.monom_insert_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.ringNorm.r_monom_insert_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.monom_insert_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)))),'const.ringNorm.r_monom_insert_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)))),'const.canonical.monom_insert_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.monom_insert_def|split|2','thm.ringNorm.r_monom_insert_def')).
fof(f82,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.monom_insert_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.monom_insert_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.canonical.monom_insert_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.monom_insert_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.canonical.monom_insert_4'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)))),'const.canonical.monom_insert_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.list.list'('type.quote.index'),X3))),s('type.canonical.canonical_sum'(X0),X4)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X20,X21,X22,X23] : s('type.canonical.canonical_sum'(X6),'const.canonical.monom_insert_4'(s('type.semi_ring.semi_ring'(X6),X20),s(X6,X21),s('type.list.list'('type.quote.index'),X22),s('type.canonical.canonical_sum'(X6),X23))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),happ(s(fun(X6,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)))),'const.canonical.monom_insert_1'(s('type.semi_ring.semi_ring'(X6),X20))),s(X6,X21))),s('type.list.list'('type.quote.index'),X22))),s('type.canonical.canonical_sum'(X6),X23)))),
  file('thm.ringNorm.monom_insert_def|split|2','ARITY_1_4_const.canonical.monom_insert')).
fof(f65,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK3),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(sK0),sK1),s(sK0,sK3),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK3),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(sK0),sK1),s(sK0,sK3),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f60])).
fof(f60,plain,(
  ? [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(X0),X1),s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(X0),X1),s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(X0),X1),s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X9),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(X6),X7),s(X6,X9),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X9),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_monom_insert_4'(s('type.ring.ring'(X6),X7),s(X6,X9),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.monom_insert_def|split|2','thm.ringNorm.monom_insert_def|split|2')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 167
% Passive clauses: 1655
% Generated clauses: 58597
% Final active clauses: 117
% Final passive clauses: 951
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 3978
% Trivial inequalities: 414
% Fw subsumption resolutions: 3828
% Fw demodulations: 6492
% Bw demodulations: 350
% Condensations: 8682
% Global subsumptions: 6
% 
% Simple tautologies: 484
% Equational tautologies: 191
% Forward subsumptions: 32874
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 123
% 
% Unit resulting resolution: 388
% Forward superposition: 15401
% Backward superposition: 17656
% Self superposition: 539
% Equality factoring: 261
% Equality resolution: 20
% 
% Split clauses: 117
% Split components: 243
% SAT solver clauses: 1917
% SAT solver unit clauses: 50
% SAT solver binary clauses: 153
% 
% Sat splits: 117
% Sat splitting refutations: 20
% 
% Memory used [KB]: 51427
% Time elapsed: 3.881 s
% ------------------------------
% ------------------------------
% Success in time 5.197 s
