Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 44
dis+11_7_41 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f143527,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f143437,f1799,f2821,f2819,f16019,f1800,f2822,f2820,f2058,f2889,f2887,f1844,f2886])).
fof(f2886,plain,(
  ~p(s(bool,'F_0')) | $spl69),
  inference(cnf_transformation,[],[f2886_D])).
fof(f2886_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f1844,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f537])).
fof(f537,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_ABS_LE|48|',aHL_FALSITY)).
fof(f2887,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(cnf_transformation,[],[f2887_D])).
fof(f2887_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2889,plain,(
  p(s(bool,'T_0')) | $spl72),
  inference(cnf_transformation,[],[f2889_D])).
fof(f2889_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2058,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f1488])).
fof(f1488,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f1487])).
fof(f1487,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f1215])).
fof(f1215,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f717])).
fof(f717,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_ABS_LE|48|','thm.bool.IMP_CLAUSES|split|2')).
fof(f2820,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(cnf_transformation,[],[f2820_D])).
fof(f2820_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f2822,plain,(
  s('type.integer.int',sK27) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(cnf_transformation,[],[f2822_D])).
fof(f2822_D,plain,(
  s('type.integer.int',sK27) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1800,plain,(
  s('type.integer.int',sK27) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(cnf_transformation,[],[f1418])).
fof(f1418,plain,(
  (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK27) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK27) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27])],[f1417])).
fof(f1417,plain,(
  ? [X0] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(nnf_transformation,[],[f1168])).
fof(f1168,plain,(
  ? [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <~> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LE|48|','thm.integer.INT_ABS_LE|48|')).
fof(f16019,plain,(
  $false | ($spl1 | $spl2 | $spl72)),
  inference(subsumption_resolution,[],[f16018,f2889])).
fof(f16018,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f15939,f2987])).
fof(f2987,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) )),
  inference(resolution,[],[f1996,f1850])).
fof(f1850,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f543])).
fof(f543,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X8] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X8),s('type.integer.int',X8))))),
  file('thm.integer.INT_ABS_LE|48|','thm.integer.INT_LE_REFL')).
fof(f1996,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f1450])).
fof(f1450,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f687])).
fof(f687,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.integer.INT_ABS_LE|48|','thm.bool.EQ_CLAUSES|split|0')).
fof(f15939,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f15855,f12834])).
fof(f12834,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f12833,f2906])).
fof(f2906,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK27) | $spl2),
  inference(backward_demodulation,[],[f1804,f2821])).
fof(f1804,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f195])).
fof(f195,axiom,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_ABS_LE|48|','thm.integer.INT_|48|')).
fof(f12833,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl1),
  inference(forward_demodulation,[],[f2820,f1804])).
fof(f15855,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_|48|_0')))),
  inference(superposition,[],[f1880,f1804])).
fof(f1880,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f573])).
fof(f573,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X13] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))),
  file('thm.integer.INT_ABS_LE|48|','thm.integer.INT_ABS_NUM')).
fof(f2819,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(cnf_transformation,[],[f2819_D])).
fof(f2819_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2821,plain,(
  s('type.integer.int',sK27) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(cnf_transformation,[],[f2821_D])).
fof(f2821_D,plain,(
  s('type.integer.int',sK27) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1799,plain,(
  s('type.integer.int',sK27) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(cnf_transformation,[],[f1418])).
fof(f143437,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f143433,f16028])).
fof(f16028,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK27) | $spl3),
  inference(forward_demodulation,[],[f2822,f1804])).
fof(f143433,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK27) | $spl0),
  inference(trivial_inequality_removal,[],[f142928])).
fof(f142928,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_|48|_0') | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK27) | $spl0),
  inference(superposition,[],[f92108,f142713])).
fof(f142713,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))) | $spl0),
  inference(subsumption_resolution,[],[f142558,f18571])).
fof(f18571,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_|48|_0'))))) )),
  inference(forward_demodulation,[],[f2081,f1804])).
fof(f2081,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f734])).
fof(f734,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(flattening,[],[f733])).
fof(f733,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  file('thm.integer.INT_ABS_LE|48|','thm.integer.INT_ABS_LT|48|')).
fof(f142558,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl0),
  inference(resolution,[],[f2330,f16029])).
fof(f16029,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK27))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl0),
  inference(forward_demodulation,[],[f2819,f1804])).
fof(f2330,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f1588])).
fof(f1588,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f1587])).
fof(f1587,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f940])).
fof(f940,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)))),
  inference(flattening,[],[f939])).
fof(f939,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & ~s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f196])).
fof(f196,axiom,(
  ! [X8,X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',X9)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X8),s('type.integer.int',X9)))) & ~s('type.integer.int',X8) = s('type.integer.int',X9)))),
  file('thm.integer.INT_ABS_LE|48|','thm.integer.INT_LT_LE')).
fof(f92108,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',X0)) )),
  inference(forward_demodulation,[],[f92107,f1804])).
fof(f92107,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',X0) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(forward_demodulation,[],[f2017,f1804])).
fof(f2017,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f1461])).
fof(f1461,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f701])).
fof(f701,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LE|48|','thm.integer.INT_ABS_EQ|48|')).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 3089
% Passive clauses: 35391
% Generated clauses: 141817
% Final active clauses: 2557
% Final passive clauses: 30888
% Input formulas: 517
% Initial clauses: 1008
% 
% Duplicate literals: 339
% Trivial inequalities: 273
% Fw subsumption resolutions: 5920
% Fw demodulations: 19804
% Bw demodulations: 425
% 
% Simple tautologies: 562
% Equational tautologies: 42
% Forward subsumptions: 78385
% Fw demodulations to eq. taut.: 23
% 
% Binary resolution: 19160
% Factoring: 920
% Forward superposition: 37152
% Backward superposition: 55795
% Self superposition: 108
% Equality factoring: 2
% Equality resolution: 39
% 
% Split clauses: 1395
% Split components: 2792
% SAT solver clauses: 1969
% SAT solver unit clauses: 45
% SAT solver binary clauses: 244
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 1395
% Sat splitting refutations: 22
% 
% TWLsolver clauses: 1958
% TWLsolver calls for satisfiability: 967
% 
% Memory used [KB]: 70489
% Time elapsed: 2.161 s
% ------------------------------
% ------------------------------
% Success in time 2.21 s
