Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (123590)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 926
% Passive clauses: 5857
% Generated clauses: 33871
% Final active clauses: 779
% Final passive clauses: 4303
% Input formulas: 133
% Initial clauses: 274
% 
% Duplicate literals: 218
% Trivial inequalities: 41
% Fw subsumption resolutions: 1562
% Fw demodulations: 3446
% Bw demodulations: 144
% 
% Simple tautologies: 189
% Equational tautologies: 27
% Forward subsumptions: 22724
% Fw demodulations to eq. taut.: 23
% 
% Binary resolution: 5783
% Factoring: 1232
% Forward superposition: 9868
% Backward superposition: 11117
% Self superposition: 85
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 220
% Split components: 440
% SAT solver clauses: 410
% SAT solver unit clauses: 12
% SAT solver binary clauses: 51
% 
% Sat splits: 220
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 410
% TWLsolver calls for satisfiability: 382
% 
% Memory used [KB]: 24690
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (123618)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 494
% Passive clauses: 2588
% Generated clauses: 21073
% Final active clauses: 420
% Final passive clauses: 1627
% Input formulas: 133
% Initial clauses: 274
% 
% Duplicate literals: 85
% Trivial inequalities: 31
% Fw subsumption resolutions: 672
% Fw demodulations: 2997
% Bw demodulations: 504
% Condensations: 69
% Global subsumptions: 15
% 
% Simple tautologies: 186
% Equational tautologies: 7
% Forward subsumptions: 11345
% Fw demodulations to eq. taut.: 21
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 1958
% Factoring: 211
% Forward superposition: 5639
% Backward superposition: 5972
% Self superposition: 93
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 11
% Split components: 22
% SAT solver clauses: 3318
% SAT solver unit clauses: 343
% SAT solver binary clauses: 81
% SAT solver learnt clauses: 30
% SAT solver learnt literals: 12
% 
% Sat splits: 11
% Sat splitting refutations: 15
% 
% TWLsolver clauses: 3258
% TWLsolver calls for satisfiability: 1699
% 
% Memory used [KB]: 24562
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (123659)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1872
% Passive clauses: 18741
% Generated clauses: 185334
% Final active clauses: 752
% Final passive clauses: 3865
% Input formulas: 133
% Initial clauses: 274
% Discarded non-redundant clauses: 24664
% 
% Duplicate literals: 161
% Trivial inequalities: 104
% Fw subsumption resolutions: 2714
% Fw demodulations: 35162
% Bw demodulations: 8120
% Condensations: 185
% Global subsumptions: 5
% 
% Simple tautologies: 119
% Equational tautologies: 76
% Forward subsumptions: 124238
% Fw demodulations to eq. taut.: 377
% Bw demodulations to eq. taut.: 44
% 
% Unit resulting resolution: 13732
% Factoring: 5
% Forward superposition: 68687
% Backward superposition: 49270
% Self superposition: 159
% Equality factoring: 28
% Equality resolution: 19
% 
% Split clauses: 123
% Split components: 246
% SAT solver clauses: 19313
% SAT solver unit clauses: 3041
% SAT solver binary clauses: 6474
% 
% Sat splits: 123
% Sat splitting refutations: 318
% 
% Memory used [KB]: 118462
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f78357,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f664,f665,f666,f667,f668,f670,f671,f672,f673,f742,f743,f835,f1118,f1120,f1209,f1206,f1198,f1275,f1276,f1281,f1284,f1285,f1286,f1776,f1801,f1804,f2104,f2106,f2109,f2111,f2407,f2408,f2311,f2310,f2305,f2304,f2303,f2302,f2510,f3007,f3014,f3974,f3947,f3946,f4047,f4048,f4054,f4056,f4058,f4059,f4060,f4061,f4064,f2306,f2308,f2309,f4542,f4311,f4310,f4309,f4308,f4306,f4305,f4304,f4303,f4302,f4745,f4751,f4746,f5480,f7138,f7147,f7139,f7314,f7331,f7333,f4920,f4919,f4918,f4917,f4915,f4914,f4913,f4912,f4911,f7367,f7368,f8133,f8130,f8122,f8204,f8205,f8210,f8213,f8214,f8215,f7477,f7476,f7475,f7474,f7472,f7471,f7470,f7469,f7468,f8326,f8311,f8298,f8297,f8399,f8400,f8406,f8408,f8410,f8411,f8412,f8413,f8416,f8419,f8420,f8423,f8424,f8443,f8444,f8451,f8452,f8453,f8454,f8455,f8457,f8458,f8459,f8460,f8610,f8611,f9134,f9872,f9382,f9381,f9380,f9379,f9377,f9376,f9375,f9374,f9373,f9888,f9889,f9892,f9893,f9912,f9913,f9920,f9921,f9922,f9923,f9924,f9926,f9927,f9928,f9929,f10087,f10088,f10602,f10619,f10637,f10655,f10667,f10684,f10685,f7332,f10688,f7146,f832,f10705,f10707,f10706,f10795,f10801,f15711,f15712,f15746,f15744,f15779,f15782,f15781,f15883,f15884,f15638,f15637,f15899,f15636,f15635,f15900,f15627,f15901,f15623,f15904,f15655,f15653,f15909,f15959,f15960,f15961,f15962,f15963,f16070,f16071,f16100,f16101,f16243,f16244,f16245,f16246,f16247,f16249,f16250,f16251,f16252,f16526,f16543,f16601,f16624,f16602,f16623,f18269,f18312,f18348,f18382,f9873,f18420,f18423,f18421,f18448,f18450,f18449,f18507,f18508,f18509,f18510,f18511,f18512,f18717,f18718,f18787,f18800,f18788,f18799,f19375,f19376,f19444,f19445,f19925,f19926,f19974,f19993,f19994,f20000,f19972,f20025,f20026,f20034,f20178,f20183,f20182,f20180,f20189,f20190,f20206,f20204,f20197,f20198,f20036,f20210,f20327,f20366,f20328,f20845,f22617,f24312,f24311,f24330,f24328,f24705,f24737,f24706,f24736,f25027,f25053,f27950,f27951,f19973,f28000,f28738,f28988,f29037,f28989,f30466,f30469,f30467,f30472,f30473,f39525,f20878,f39528,f28780,f41553,f41554,f41723,f42436,f42499,f42857,f42858,f20205,f41721,f46899,f46901,f46904,f46922,f46924,f32057,f32058,f32060,f32099,f32097,f46957,f32074,f32072,f32070,f15903,f32068,f32066,f46958,f15908,f15907,f1780,f15647,f15646,f15898,f46968,f46969,f46976,f46977,f46996,f46997,f47004,f47005,f676,f47006,f47007,f677,f47008,f47010,f47011,f679,f674,f47012,f47013,f681,f675,f47104,f833,f47105,f47297,f47298,f1803,f47883,f48050,f48306,f48429,f9135,f48663,f48674,f48768,f48779,f48874,f48885,f48985,f49004,f48986,f49003,f49137,f49148,f49247,f19457,f49248,f19456,f49343,f19388,f49344,f19387,f50057,f50061,f50059,f50686,f51508,f51577,f52394,f55299,f55300,f55375,f55386,f55389,f55392,f55972,f56913,f56948,f56914,f56947,f57423,f57450,f60329,f62408,f62463,f69109,f29036,f69111,f20365,f69114,f50761,f69115,f60390,f69243,f69339,f71714,f71715,f71719,f71721,f71807,f71795,f71794,f71805,f72556,f72557,f72561,f72563,f72648,f71806,f72637,f72636,f72647,f77300,f77302,f77301,f77327,f77328,f77329,f77279,f77334,f836,f77335,f837,f77342,f77355,f77357,f77362,f77368,f77381,f77383,f77388,f77363,f77531,f77534,f77535,f77536,f77537,f77538,f77539,f77540,f77542,f77543,f77545,f77546,f77547,f77548,f77549,f77564,f77567,f77568,f77569,f77570,f77571,f77572,f77573,f77574,f77575,f77576,f77577,f77578,f77593,f77598,f77599,f77600,f77601,f77602,f77603,f77604,f77605,f77606,f77607,f77608,f77609,f77622,f77626,f77627,f77629,f77630,f77631,f77632,f77633,f77634,f77635,f77636,f77656,f77657,f77658,f77661,f77662,f77663,f77664,f77665,f77666,f77672,f77673,f77674,f77675,f77700,f77702,f18422,f77701,f77717,f77720,f77721,f77722,f77723,f77724,f77725,f77726,f77727,f77728,f77729,f77730,f77745,f77748,f77749,f77750,f77751,f77752,f77753,f77754,f77755,f77756,f77757,f77758,f77774,f77776,f77779,f77780,f77777,f77781,f77782,f77783,f77784,f77785,f77786,f77787,f77788,f77789,f77790,f2108,f77958,f77976,f2107,f77965,f1119,f5486,f78130,f18730,f78131,f78135,f18729,f1778,f78350,f78351])).
fof(f78351,plain,(
  $false | ($spl8 | $spl15)),
  inference(subsumption_resolution,[],[f78258,f401])).
fof(f401,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X0))))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0] : ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X0))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0] : ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X0))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X6),s('type.pair.prod'('type.num.num','type.num.num'),X6))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.integer.TINT_LT_REFL')).
fof(f78258,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | ($spl8 | $spl15)),
  inference(resolution,[],[f77507,f47630])).
fof(f47630,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X2),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X2),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(subsumption_resolution,[],[f47447,f514])).
fof(f514,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f335])).
fof(f335,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0] : ((p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X19] : ((p(s(bool,'T_0')) | p(s(bool,X19))) <=> p(s(bool,'T_0')))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.OR_CLAUSES|split|0')).
fof(f47447,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X2),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X2),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl8),
  inference(superposition,[],[f396,f680])).
fof(f680,plain,(
  s(bool,'T_0') = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | $spl8),
  inference(cnf_transformation,[],[f680_D])).
fof(f680_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X2)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X2)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X2)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X2)))))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X2))))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X2)))))),
  inference(ennf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) & p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X2))))) => p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X2)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X11,X18] : ((p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X6),s('type.pair.prod'('type.num.num','type.num.num'),X11)))) & p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X11),s('type.pair.prod'('type.num.num','type.num.num'),X18))))) => p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X6),s('type.pair.prod'('type.num.num','type.num.num'),X18)))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.integer.TINT_LT_TRANS')).
fof(f77507,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ($spl8 | $spl15)),
  inference(forward_demodulation,[],[f77506,f568])).
fof(f568,plain,(
  ( ! [X0] : (s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X0))) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X0)))) )),
  inference(resolution,[],[f350,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X2))) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X2))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) | s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X2))) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X2))))),
  inference(ennf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) => s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X2))) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X2))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9,X10,X11] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X9),s('type.pair.prod'('type.num.num','type.num.num'),X10)))) => s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X9),s('type.pair.prod'('type.num.num','type.num.num'),X11))) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X10),s('type.pair.prod'('type.num.num','type.num.num'),X11))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.integer.TINT_LT_WELLDEFR')).
fof(f350,plain,(
  p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) & s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12])],[f226])).
fof(f226,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X2)))) & s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X2))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X2)))) => s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X7),s('type.pair.prod'('type.num.num','type.num.num'),X8)))) => s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X6),s('type.pair.prod'('type.num.num','type.num.num'),X7))) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X6),s('type.pair.prod'('type.num.num','type.num.num'),X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X7),s('type.pair.prod'('type.num.num','type.num.num'),X8)))) => s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X6),s('type.pair.prod'('type.num.num','type.num.num'),X7))) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X6),s('type.pair.prod'('type.num.num','type.num.num'),X8))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.integer.TINT_LT_WELLDEFL')).
fof(f77506,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ($spl8 | $spl15)),
  inference(subsumption_resolution,[],[f77505,f519])).
fof(f519,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f339])).
fof(f339,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X19] : ((p(s(bool,'F_0')) | p(s(bool,X19))) <=> p(s(bool,X19)))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.OR_CLAUSES|split|2')).
fof(f77505,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ($spl8 | $spl15)),
  inference(forward_demodulation,[],[f77419,f4706])).
fof(f4706,plain,(
  s(bool,'F_0') = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | $spl8),
  inference(resolution,[],[f2862,f507])).
fof(f507,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f331])).
fof(f331,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X19] : ((p(s(bool,X19)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X19))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.IMP_F_EQ_F')).
fof(f2862,plain,(
  ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl8),
  inference(trivial_inequality_removal,[],[f2861])).
fof(f2861,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl8),
  inference(superposition,[],[f2540,f532])).
fof(f532,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X19] : (s(bool,'T_0') = s(bool,X19) <=> p(s(bool,X19)))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.EQ_CLAUSES|split|0')).
fof(f2540,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | $spl8),
  inference(backward_demodulation,[],[f680,f351])).
fof(f351,plain,(
  s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))),
  inference(cnf_transformation,[],[f265])).
fof(f77419,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl15),
  inference(resolution,[],[f836,f400])).
fof(f400,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X0)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X11] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X6),s('type.pair.prod'('type.num.num','type.num.num'),X11)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X6),s('type.pair.prod'('type.num.num','type.num.num'),X11)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X11),s('type.pair.prod'('type.num.num','type.num.num'),X6)))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.integer.TINT_LT_TOTAL')).
fof(f78350,plain,(
  $false | ($spl8 | $spl15)),
  inference(subsumption_resolution,[],[f78255,f401])).
fof(f78255,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ($spl8 | $spl15)),
  inference(resolution,[],[f77507,f47629])).
fof(f47629,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X1))))) ) | $spl8),
  inference(subsumption_resolution,[],[f47446,f514])).
fof(f47446,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X1))))) ) | $spl8),
  inference(superposition,[],[f396,f680])).
fof(f1778,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) | $spl20),
  inference(cnf_transformation,[],[f1778_D])).
fof(f1778_D,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f18729,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl85),
  inference(cnf_transformation,[],[f18729_D])).
fof(f18729_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f78135,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~p(s(bool,X13))) ) | ($spl8 | $spl15)),
  inference(resolution,[],[f77424,f48436])).
fof(f48436,plain,(
  ( ! [X8,X9] : (~sP2(X8,X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,X9) | ~p(s(bool,X8))) ) | $spl8),
  inference(subsumption_resolution,[],[f48435,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f48435,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | ~sP2(X8,X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~sP0(X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X8)) ) | $spl8),
  inference(subsumption_resolution,[],[f48434,f46994])).
fof(f46994,plain,(
  ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl8),
  inference(trivial_inequality_removal,[],[f46993])).
fof(f46993,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl8),
  inference(superposition,[],[f2540,f532])).
fof(f48434,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~sP2(X8,X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~sP0(X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X8)) ) | $spl8),
  inference(subsumption_resolution,[],[f48408,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f273])).
fof(f273,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f48408,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~sP2(X8,X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~sP1('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X9,X8) | ~sP0(X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X8)) ) | $spl8),
  inference(resolution,[],[f46983,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f267])).
fof(f267,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f46983,plain,(
  ( ! [X8,X9] : (~sP3(X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X8) | ~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X8)) ) | $spl8),
  inference(superposition,[],[f2540,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP3(X2,X1,X0)) & (~sP3(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP3(X2,X1,X0))),
  inference(definition_folding,[],[f137,f253,f252,f251,f250])).
fof(f252,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f137,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X26,X27,X22] : ((p(s(bool,X22)) <=> s(bool,X26) = s(bool,X27)) <=> ((p(s(bool,X22)) | p(s(bool,X27)) | p(s(bool,X26))) & (p(s(bool,X22)) | ~p(s(bool,X26)) | ~p(s(bool,X27))) & (p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,X22))) & (p(s(bool,X26)) | ~p(s(bool,X27)) | ~p(s(bool,X22)))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.sat.sat_2|49|_prove_ltb')).
fof(f77424,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X9)) ) | $spl15),
  inference(resolution,[],[f836,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f270])).
fof(f270,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f252])).
fof(f78131,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~p(s(bool,X9))) ) | $spl15),
  inference(resolution,[],[f77424,f1771])).
fof(f1771,plain,(
  ( ! [X14,X15] : (~sP2(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f1770,f539])).
fof(f539,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.TINT_LT_WELLDEFL',aHL_FALSITY)).
fof(f1770,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP2(X14,X15,'F_0')) )),
  inference(subsumption_resolution,[],[f1769,f1359])).
fof(f1359,plain,(
  ( ! [X4,X5] : (sP0(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1354,f368])).
fof(f1354,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP0(X4,'F_0',X5)) )),
  inference(resolution,[],[f1095,f352])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (sP0(X2,X1,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f1095,plain,(
  ( ! [X6,X7] : (sP3(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f1088,f371])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f277])).
fof(f1088,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f1080,f350])).
fof(f1080,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f632,f530])).
fof(f530,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X19] : (s(bool,'T_0') = s(bool,X19) <=> p(s(bool,X19)))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.EQ_CLAUSES|split|1')).
fof(f632,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))),
  inference(resolution,[],[f350,f525])).
fof(f525,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X19] : (s(bool,'F_0') = s(bool,X19) <=> ~p(s(bool,X19)))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.EQ_CLAUSES|split|2')).
fof(f1769,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP2(X14,X15,'F_0') | ~sP0(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f1762,f366])).
fof(f1762,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP2(X14,X15,'F_0') | ~sP1('F_0',X15,X14) | ~sP0(X15,'F_0',X14)) )),
  inference(resolution,[],[f1097,f357])).
fof(f1097,plain,(
  ( ! [X10,X11] : (~sP3(X11,'F_0',X10) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) )),
  inference(superposition,[],[f1088,f373])).
fof(f18730,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl86),
  inference(cnf_transformation,[],[f18730_D])).
fof(f18730_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f78130,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl15),
  inference(resolution,[],[f77424,f1367])).
fof(f1367,plain,(
  ( ! [X4,X5] : (~sP2(X5,X4,'F_0') | sP3(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1366,f539])).
fof(f1366,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP3(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP2(X5,X4,'F_0')) )),
  inference(subsumption_resolution,[],[f1363,f366])).
fof(f1363,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP3(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP2(X5,X4,'F_0') | ~sP1('F_0',X4,X5)) )),
  inference(resolution,[],[f1359,f357])).
fof(f5486,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl37),
  inference(cnf_transformation,[],[f5486_D])).
fof(f5486_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f1119,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl16),
  inference(cnf_transformation,[],[f1119_D])).
fof(f1119_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f77965,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl28),
  inference(resolution,[],[f4940,f1700])).
fof(f1700,plain,(
  ( ! [X4,X5] : (~sP0(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(duplicate_literal_removal,[],[f1695])).
fof(f1695,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | ~sP0(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(resolution,[],[f1669,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f277])).
fof(f1669,plain,(
  ( ! [X6,X7] : (sP3(X6,'F_0',X7) | p(s(bool,X6)) | ~sP0(X7,'F_0',X6)) )),
  inference(subsumption_resolution,[],[f1667,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f271])).
fof(f1667,plain,(
  ( ! [X6,X7] : (sP3(X6,'F_0',X7) | p(s(bool,X6)) | ~sP2(X6,X7,'F_0') | ~sP0(X7,'F_0',X6)) )),
  inference(resolution,[],[f1663,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f1663,plain,(
  ( ! [X4,X5] : (sP1('F_0',X4,X5)) )),
  inference(subsumption_resolution,[],[f1659,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f274])).
fof(f1659,plain,(
  ( ! [X4,X5] : (sP1('F_0',X4,X5) | p(s(bool,X5))) )),
  inference(resolution,[],[f1656,f362])).
fof(f1656,plain,(
  ( ! [X6,X7] : (~sP2(X7,X6,'F_0') | sP1('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f1652,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f274])).
fof(f1652,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP2(X7,X6,'F_0') | sP1('F_0',X6,X7)) )),
  inference(resolution,[],[f1374,f353])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,X2,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f1374,plain,(
  ( ! [X2,X3] : (sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f1373,f530])).
fof(f1373,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f1369,f1359])).
fof(f1369,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0') | ~sP0(X2,'F_0',X3)) )),
  inference(resolution,[],[f1360,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f1360,plain,(
  ( ! [X6,X7] : (sP1('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1355,f364])).
fof(f1355,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP1('F_0',X6,X7)) )),
  inference(resolution,[],[f1095,f353])).
fof(f4940,plain,(
  ( ! [X26,X25] : (sP0('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X25,X26)) ) | $spl28),
  inference(resolution,[],[f2108,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f276])).
fof(f2107,plain,(
  ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl29),
  inference(cnf_transformation,[],[f2107_D])).
fof(f2107_D,plain,(
  ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f77976,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~sP1(X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X2) | ~sP2(X2,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X3)) ) | $spl230),
  inference(cnf_transformation,[],[f77976_D])).
fof(f77976_D,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~sP1(X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X2) | ~sP2(X2,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X3)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f77958,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~sP2(X2,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X3) | ~sP1(X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X2)) ) | $spl28),
  inference(resolution,[],[f4940,f356])).
fof(f2108,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl28),
  inference(cnf_transformation,[],[f2108_D])).
fof(f2108_D,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f77790,plain,(
  ( ! [X19] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X19))) ) | ($spl12 | $spl54 | $spl109)),
  inference(forward_demodulation,[],[f77773,f568])).
fof(f77773,plain,(
  ( ! [X19] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X19))) ) | ($spl12 | $spl54 | $spl109)),
  inference(resolution,[],[f935,f43447])).
fof(f43447,plain,(
  ( ! [X173,X174] : (~sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X173,X174) | p(s(bool,X174)) | ~p(s(bool,X173))) ) | ($spl54 | $spl109)),
  inference(subsumption_resolution,[],[f43322,f514])).
fof(f43322,plain,(
  ( ! [X173,X174] : (~p(s(bool,'T_0')) | ~p(s(bool,X173)) | p(s(bool,X174)) | ~sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X173,X174)) ) | ($spl54 | $spl109)),
  inference(superposition,[],[f426,f42131])).
fof(f42131,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ($spl54 | $spl109)),
  inference(forward_demodulation,[],[f41987,f15745])).
fof(f15745,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl54),
  inference(cnf_transformation,[],[f15745_D])).
fof(f15745_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f41987,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl109),
  inference(superposition,[],[f561,f41769])).
fof(f41769,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl109),
  inference(subsumption_resolution,[],[f41602,f24592])).
fof(f24592,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')))))) | $spl109),
  inference(forward_demodulation,[],[f24591,f392])).
fof(f392,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X21] : s('type.num.num',X21) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X21)))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.arithmetic.ADD|split|0')).
fof(f24591,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')))))))) | $spl109),
  inference(forward_demodulation,[],[f24590,f384])).
fof(f384,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X20,X21] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num',X21))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X20)))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.arithmetic.ADD_SYM')).
fof(f24590,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')),s('type.num.num','const.num.|48|_0'))))))) | $spl109),
  inference(forward_demodulation,[],[f24589,f550])).
fof(f550,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X21,X20] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num',X21)))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f24589,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl109),
  inference(forward_demodulation,[],[f24485,f555])).
fof(f555,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f50])).
fof(f50,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.numeral.numeral_distrib|split|13')).
fof(f24485,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl109),
  inference(resolution,[],[f24342,f502])).
fof(f502,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13(X1,X0)),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13(X1,X0)),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f244])).
fof(f244,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ? [X2] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(ennf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) => ? [X2] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(rectify,[],[f134])).
fof(f134,axiom,(
  ! [X20,X21] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num',X20)))) => ? [X22] : s('type.num.num',X20) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.arithmetic.LESS_ADD_|49|')).
fof(f24342,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl109),
  inference(resolution,[],[f24310,f463])).
fof(f463,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X20,X21] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num',X21)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X21),s('type.num.num',X20)))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.arithmetic.NOT_LESS')).
fof(f24310,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) | $spl109),
  inference(equality_resolution,[],[f20221])).
fof(f20221,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X4) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0'))))) ) | $spl109),
  inference(subsumption_resolution,[],[f20217,f538])).
fof(f538,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X21] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X21))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.arithmetic.ZERO_LESS_EQ')).
fof(f20217,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X4) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4))))) ) | $spl109),
  inference(superposition,[],[f20036,f501])).
fof(f501,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(rectify,[],[f108])).
fof(f108,axiom,(
  ! [X20,X21] : (s('type.num.num',X20) = s('type.num.num',X21) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X20),s('type.num.num',X21)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X21),s('type.num.num',X20))))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.arithmetic.EQ_LESS_EQ')).
fof(f41602,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')))))) | $spl109),
  inference(superposition,[],[f388,f24592])).
fof(f388,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X21] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X21))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.numeral.numeral_distrib|split|18')).
fof(f561,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X21] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X21),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.numeral.numeral_distrib|split|27')).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f935,plain,(
  ( ! [X80,X79] : (sP6(X79,X80,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(resolution,[],[f835,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f296])).
fof(f77789,plain,(
  ( ! [X18] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X18))) ) | ($spl12 | $spl54)),
  inference(forward_demodulation,[],[f77772,f568])).
fof(f77772,plain,(
  ( ! [X18] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X18))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f935,f17001])).
fof(f17001,plain,(
  ( ! [X173,X174] : (~sP6('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X173,X174) | p(s(bool,X174)) | ~p(s(bool,X173))) ) | $spl54),
  inference(subsumption_resolution,[],[f16872,f514])).
fof(f16872,plain,(
  ( ! [X173,X174] : (~p(s(bool,'T_0')) | ~p(s(bool,X173)) | p(s(bool,X174)) | ~sP6('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X173,X174)) ) | $spl54),
  inference(superposition,[],[f426,f15745])).
fof(f77788,plain,(
  ( ! [X17] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X17))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77771,f568])).
fof(f77771,plain,(
  ( ! [X17] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X17))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f935,f24525])).
fof(f24525,plain,(
  ( ! [X72,X73] : (~sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X72,X73) | p(s(bool,X73)) | ~p(s(bool,X72))) ) | $spl109),
  inference(resolution,[],[f24342,f426])).
fof(f77787,plain,(
  ( ! [X16] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X16))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77770,f568])).
fof(f77770,plain,(
  ( ! [X16] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X16))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f935,f47680])).
fof(f47680,plain,(
  ( ! [X177,X176] : (~sP6('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X176,X177) | p(s(bool,X177)) | ~p(s(bool,X176))) ) | $spl8),
  inference(subsumption_resolution,[],[f47546,f514])).
fof(f47546,plain,(
  ( ! [X177,X176] : (~p(s(bool,'T_0')) | ~p(s(bool,X176)) | p(s(bool,X177)) | ~sP6('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X176,X177)) ) | $spl8),
  inference(superposition,[],[f426,f680])).
fof(f77786,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X15))) ) | $spl12),
  inference(forward_demodulation,[],[f77769,f568])).
fof(f77769,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X15))) ) | $spl12),
  inference(resolution,[],[f935,f607])).
fof(f607,plain,(
  ( ! [X74,X73] : (~sP6('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X73,X74) | p(s(bool,X74)) | ~p(s(bool,X73))) )),
  inference(resolution,[],[f350,f426])).
fof(f77785,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X14))) ) | ($spl12 | $spl54 | $spl109)),
  inference(forward_demodulation,[],[f77768,f568])).
fof(f77768,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X14))) ) | ($spl12 | $spl54 | $spl109)),
  inference(resolution,[],[f935,f43448])).
fof(f43448,plain,(
  ( ! [X175,X176] : (~sP6(X175,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X176) | p(s(bool,X176)) | ~p(s(bool,X175))) ) | ($spl54 | $spl109)),
  inference(subsumption_resolution,[],[f43323,f514])).
fof(f43323,plain,(
  ( ! [X175,X176] : (~p(s(bool,'T_0')) | ~p(s(bool,X175)) | p(s(bool,X176)) | ~sP6(X175,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X176)) ) | ($spl54 | $spl109)),
  inference(superposition,[],[f426,f42131])).
fof(f77784,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X13))) ) | ($spl12 | $spl54)),
  inference(forward_demodulation,[],[f77767,f568])).
fof(f77767,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X13))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f935,f17002])).
fof(f17002,plain,(
  ( ! [X175,X176] : (~sP6(X175,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X176) | p(s(bool,X176)) | ~p(s(bool,X175))) ) | $spl54),
  inference(subsumption_resolution,[],[f16873,f514])).
fof(f16873,plain,(
  ( ! [X175,X176] : (~p(s(bool,'T_0')) | ~p(s(bool,X175)) | p(s(bool,X176)) | ~sP6(X175,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X176)) ) | $spl54),
  inference(superposition,[],[f426,f15745])).
fof(f77783,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X12))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77766,f568])).
fof(f77766,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X12))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f935,f24526])).
fof(f24526,plain,(
  ( ! [X74,X75] : (~sP6(X74,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X75) | p(s(bool,X75)) | ~p(s(bool,X74))) ) | $spl109),
  inference(resolution,[],[f24342,f426])).
fof(f77782,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X11))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77765,f568])).
fof(f77765,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X11))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f935,f47681])).
fof(f47681,plain,(
  ( ! [X178,X179] : (~sP6(X178,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X179) | p(s(bool,X179)) | ~p(s(bool,X178))) ) | $spl8),
  inference(subsumption_resolution,[],[f47547,f514])).
fof(f47547,plain,(
  ( ! [X178,X179] : (~p(s(bool,'T_0')) | ~p(s(bool,X178)) | p(s(bool,X179)) | ~sP6(X178,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X179)) ) | $spl8),
  inference(superposition,[],[f426,f680])).
fof(f77781,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X10))) ) | $spl12),
  inference(forward_demodulation,[],[f77764,f568])).
fof(f77764,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X10))) ) | $spl12),
  inference(resolution,[],[f935,f608])).
fof(f608,plain,(
  ( ! [X76,X75] : (~sP6(X75,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X76) | p(s(bool,X76)) | ~p(s(bool,X75))) )),
  inference(resolution,[],[f350,f426])).
fof(f77777,plain,(
  ( ! [X4,X5] : (sP7('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X4,X5)) ) | $spl220),
  inference(cnf_transformation,[],[f77777_D])).
fof(f77777_D,plain,(
  ( ! [X4,X5] : (sP7('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X4,X5)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f77780,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | sP7('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X8,X9)) ) | $spl12),
  inference(forward_demodulation,[],[f77763,f568])).
fof(f77763,plain,(
  ( ! [X8,X9] : (sP7('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X8,X9) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f935,f565])).
fof(f565,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f293])).
fof(f293,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f77779,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | sP7('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7)) ) | $spl12),
  inference(forward_demodulation,[],[f77778,f568])).
fof(f77778,plain,(
  ( ! [X6,X7] : (sP7('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(subsumption_resolution,[],[f77762,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f162,f259,f258])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X26,X27,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X27)) & p(s(bool,X26)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X27)) | ~p(s(bool,X26))) & (p(s(bool,X27)) | ~p(s(bool,X22))) & (p(s(bool,X26)) | ~p(s(bool,X22)))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.sat.sat_23_prove_ltb')).
fof(f77762,plain,(
  ( ! [X6,X7] : (sP7('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X7))) ) | $spl12),
  inference(resolution,[],[f935,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f294])).
fof(f77776,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | sP7('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X4,X5)) ) | $spl12),
  inference(forward_demodulation,[],[f77775,f568])).
fof(f77775,plain,(
  ( ! [X4,X5] : (sP7('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X4,X5) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(subsumption_resolution,[],[f77761,f431])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f298])).
fof(f77761,plain,(
  ( ! [X4,X5] : (sP7('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f935,f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f294])).
fof(f77774,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl12),
  inference(forward_demodulation,[],[f77759,f568])).
fof(f77759,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f935,f426])).
fof(f77758,plain,(
  ( ! [X18] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X18))) ) | ($spl12 | $spl54 | $spl109)),
  inference(forward_demodulation,[],[f77744,f568])).
fof(f77744,plain,(
  ( ! [X18] : (p(s(bool,X18)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl12 | $spl54 | $spl109)),
  inference(resolution,[],[f919,f43433])).
fof(f43433,plain,(
  ( ! [X116,X115] : (~sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X115,X116) | p(s(bool,X116)) | p(s(bool,X115))) ) | ($spl54 | $spl109)),
  inference(subsumption_resolution,[],[f43293,f544])).
fof(f544,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.TINT_LT_WELLDEFL',aHL_TRUTH)).
fof(f43293,plain,(
  ( ! [X116,X115] : (~p(s(bool,'T_0')) | p(s(bool,X115)) | p(s(bool,X116)) | ~sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X115,X116)) ) | ($spl54 | $spl109)),
  inference(superposition,[],[f409,f42131])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f288])).
fof(f288,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f919,plain,(
  ( ! [X47,X48] : (sP4(X47,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X48)) ) | $spl12),
  inference(resolution,[],[f835,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f289])).
fof(f77757,plain,(
  ( ! [X17] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X17))) ) | ($spl12 | $spl54)),
  inference(forward_demodulation,[],[f77743,f568])).
fof(f77743,plain,(
  ( ! [X17] : (p(s(bool,X17)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f919,f16989])).
fof(f16989,plain,(
  ( ! [X116,X115] : (~sP4('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X115,X116) | p(s(bool,X116)) | p(s(bool,X115))) ) | $spl54),
  inference(subsumption_resolution,[],[f16843,f544])).
fof(f16843,plain,(
  ( ! [X116,X115] : (~p(s(bool,'T_0')) | p(s(bool,X115)) | p(s(bool,X116)) | ~sP4('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X115,X116)) ) | $spl54),
  inference(superposition,[],[f409,f15745])).
fof(f77756,plain,(
  ( ! [X16] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X16))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77742,f568])).
fof(f77742,plain,(
  ( ! [X16] : (p(s(bool,X16)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f919,f24509])).
fof(f24509,plain,(
  ( ! [X41,X40] : (~sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X40,X41) | p(s(bool,X41)) | p(s(bool,X40))) ) | $spl109),
  inference(resolution,[],[f24342,f409])).
fof(f77755,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X15))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77741,f568])).
fof(f77741,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f919,f47666])).
fof(f47666,plain,(
  ( ! [X118,X119] : (~sP4('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X118,X119) | p(s(bool,X119)) | p(s(bool,X118))) ) | $spl8),
  inference(subsumption_resolution,[],[f47517,f544])).
fof(f47517,plain,(
  ( ! [X118,X119] : (~p(s(bool,'T_0')) | p(s(bool,X118)) | p(s(bool,X119)) | ~sP4('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X118,X119)) ) | $spl8),
  inference(superposition,[],[f409,f680])).
fof(f77754,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X14))) ) | $spl12),
  inference(forward_demodulation,[],[f77740,f568])).
fof(f77740,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f919,f591])).
fof(f591,plain,(
  ( ! [X41,X42] : (~sP4('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X41,X42) | p(s(bool,X42)) | p(s(bool,X41))) )),
  inference(resolution,[],[f350,f409])).
fof(f77753,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X13))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77739,f568])).
fof(f77739,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X13))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f919,f24368])).
fof(f24368,plain,(
  ( ! [X45,X46] : (~sP4(X45,X46,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X46)) | ~p(s(bool,X45))) ) | $spl109),
  inference(resolution,[],[f24310,f409])).
fof(f77752,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X12))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77738,f568])).
fof(f77738,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X12))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f919,f42320])).
fof(f42320,plain,(
  ( ! [X70,X71,X69] : (~sP4(X69,X70,'const.prim_rec.<_2'(s('type.num.num',X71),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X70)) | ~p(s(bool,X69))) ) | $spl109),
  inference(resolution,[],[f42160,f409])).
fof(f42160,plain,(
  ( ! [X53] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X53),s('type.num.num','const.num.|48|_0'))))) ) | $spl109),
  inference(forward_demodulation,[],[f42159,f41769])).
fof(f42159,plain,(
  ( ! [X53] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X53),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl109),
  inference(subsumption_resolution,[],[f42030,f538])).
fof(f42030,plain,(
  ( ! [X53] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X53)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X53),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl109),
  inference(superposition,[],[f464,f41769])).
fof(f464,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f309])).
fof(f77751,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X11))) ) | ($spl12 | $spl56)),
  inference(forward_demodulation,[],[f77737,f568])).
fof(f77737,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X11))) ) | ($spl12 | $spl56)),
  inference(resolution,[],[f919,f16278])).
fof(f16278,plain,(
  ( ! [X45,X44] : (~sP4(X44,X45,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | p(s(bool,X45)) | ~p(s(bool,X44))) ) | $spl56),
  inference(resolution,[],[f16105,f409])).
fof(f16105,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl56),
  inference(resolution,[],[f15780,f464])).
fof(f15780,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl56),
  inference(cnf_transformation,[],[f15780_D])).
fof(f15780_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f77750,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X10))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77736,f568])).
fof(f77736,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X10))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f919,f47041])).
fof(f47041,plain,(
  ( ! [X45,X46] : (~sP4(X45,X46,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | p(s(bool,X46)) | ~p(s(bool,X45))) ) | $spl8),
  inference(resolution,[],[f46994,f409])).
fof(f77749,plain,(
  ( ! [X8,X9] : (sP5(X8,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X9) | p(s(bool,X9))) ) | $spl218),
  inference(cnf_transformation,[],[f77749_D])).
fof(f77749_D,plain,(
  ( ! [X8,X9] : (sP5(X8,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X9) | p(s(bool,X9))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f77748,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | sP5(X8,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X9) | p(s(bool,X9))) ) | $spl12),
  inference(forward_demodulation,[],[f77735,f568])).
fof(f77735,plain,(
  ( ! [X8,X9] : (sP5(X8,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X9) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X9))) ) | $spl12),
  inference(resolution,[],[f919,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP4(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP4(X0,X1,X2) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f285])).
fof(f285,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1)) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f77745,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl12),
  inference(forward_demodulation,[],[f77731,f568])).
fof(f77731,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X1))) ) | $spl12),
  inference(resolution,[],[f919,f409])).
fof(f77730,plain,(
  ( ! [X18] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X18))) ) | ($spl12 | $spl54 | $spl109)),
  inference(forward_demodulation,[],[f77716,f568])).
fof(f77716,plain,(
  ( ! [X18] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X18))) ) | ($spl12 | $spl54 | $spl109)),
  inference(resolution,[],[f918,f43433])).
fof(f918,plain,(
  ( ! [X45,X46] : (sP4(X45,X46,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(resolution,[],[f835,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f289])).
fof(f77729,plain,(
  ( ! [X17] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X17))) ) | ($spl12 | $spl54)),
  inference(forward_demodulation,[],[f77715,f568])).
fof(f77715,plain,(
  ( ! [X17] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X17))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f918,f16989])).
fof(f77728,plain,(
  ( ! [X16] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X16))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77714,f568])).
fof(f77714,plain,(
  ( ! [X16] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X16))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f918,f24509])).
fof(f77727,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X15))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77713,f568])).
fof(f77713,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X15))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f918,f47666])).
fof(f77726,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X14))) ) | $spl12),
  inference(forward_demodulation,[],[f77712,f568])).
fof(f77712,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X14))) ) | $spl12),
  inference(resolution,[],[f918,f591])).
fof(f77725,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X13))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77711,f568])).
fof(f77711,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X13))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f918,f24369])).
fof(f24369,plain,(
  ( ! [X47,X48] : (~sP4(X47,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')),X48) | p(s(bool,X48)) | ~p(s(bool,X47))) ) | $spl109),
  inference(resolution,[],[f24310,f409])).
fof(f77724,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X12))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77710,f568])).
fof(f77710,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X12))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f918,f42321])).
fof(f42321,plain,(
  ( ! [X74,X72,X73] : (~sP4(X72,'const.prim_rec.<_2'(s('type.num.num',X74),s('type.num.num','const.num.|48|_0')),X73) | p(s(bool,X73)) | ~p(s(bool,X72))) ) | $spl109),
  inference(resolution,[],[f42160,f409])).
fof(f77723,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X11))) ) | ($spl12 | $spl56)),
  inference(forward_demodulation,[],[f77709,f568])).
fof(f77709,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X11))) ) | ($spl12 | $spl56)),
  inference(resolution,[],[f918,f16279])).
fof(f16279,plain,(
  ( ! [X47,X46] : (~sP4(X46,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X47) | p(s(bool,X47)) | ~p(s(bool,X46))) ) | $spl56),
  inference(resolution,[],[f16105,f409])).
fof(f77722,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X10))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77708,f568])).
fof(f77708,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X10))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f918,f47042])).
fof(f47042,plain,(
  ( ! [X47,X48] : (~sP4(X47,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X48) | p(s(bool,X48)) | ~p(s(bool,X47))) ) | $spl8),
  inference(resolution,[],[f46994,f409])).
fof(f77721,plain,(
  ( ! [X8,X9] : (sP5(X8,X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | p(s(bool,X9))) ) | $spl216),
  inference(cnf_transformation,[],[f77721_D])).
fof(f77721_D,plain,(
  ( ! [X8,X9] : (sP5(X8,X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | p(s(bool,X9))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f77720,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | sP5(X8,X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | p(s(bool,X9))) ) | $spl12),
  inference(forward_demodulation,[],[f77707,f568])).
fof(f77707,plain,(
  ( ! [X8,X9] : (sP5(X8,X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | p(s(bool,X9)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f918,f408])).
fof(f77717,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl12),
  inference(forward_demodulation,[],[f77703,f568])).
fof(f77703,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f918,f409])).
fof(f77701,plain,(
  ( ! [X20] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X20)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl214),
  inference(cnf_transformation,[],[f77701_D])).
fof(f77701_D,plain,(
  ( ! [X20] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X20)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f18422,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,X21)) ) | $spl80),
  inference(cnf_transformation,[],[f18422_D])).
fof(f18422_D,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,X21)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f77702,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X22)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X22),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(subsumption_resolution,[],[f77687,f525])).
fof(f77687,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X22)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X22),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X23))) ) | $spl7),
  inference(superposition,[],[f47278,f383])).
fof(f383,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1))))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(rectify,[],[f109])).
fof(f109,axiom,(
  ! [X23,X24] : (s(bool,X23) = s(bool,X24) <=> ((p(s(bool,X23)) & p(s(bool,X24))) | (~p(s(bool,X23)) & ~p(s(bool,X24)))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.EQ_EXPAND')).
fof(f47278,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X1))))) ) | $spl7),
  inference(superposition,[],[f679,f398])).
fof(f77700,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X21) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X20)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(subsumption_resolution,[],[f77686,f525])).
fof(f77686,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X21) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X20)))) | p(s(bool,X21)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(superposition,[],[f47278,f383])).
fof(f77675,plain,(
  ( ! [X17] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X17))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77655,f568])).
fof(f77655,plain,(
  ( ! [X17] : (p(s(bool,X17)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f906,f24354])).
fof(f24354,plain,(
  ( ! [X21,X22] : (~sP0(X21,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')),X22) | p(s(bool,X22)) | p(s(bool,X21))) ) | $spl109),
  inference(resolution,[],[f24310,f367])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f276])).
fof(f906,plain,(
  ( ! [X26,X25] : (sP0('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X25,X26)) ) | $spl12),
  inference(resolution,[],[f835,f370])).
fof(f77674,plain,(
  ( ! [X16] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X16))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77654,f568])).
fof(f77654,plain,(
  ( ! [X16] : (p(s(bool,X16)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f906,f42306])).
fof(f42306,plain,(
  ( ! [X33,X31,X32] : (~sP0(X31,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num','const.num.|48|_0')),X32) | p(s(bool,X32)) | p(s(bool,X31))) ) | $spl109),
  inference(resolution,[],[f42160,f367])).
fof(f77673,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X15))) ) | ($spl12 | $spl56)),
  inference(forward_demodulation,[],[f77653,f568])).
fof(f77653,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl12 | $spl56)),
  inference(resolution,[],[f906,f16264])).
fof(f16264,plain,(
  ( ! [X21,X20] : (~sP0(X20,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X21) | p(s(bool,X21)) | p(s(bool,X20))) ) | $spl56),
  inference(resolution,[],[f16105,f367])).
fof(f77672,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X14))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77652,f568])).
fof(f77652,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f906,f47027])).
fof(f47027,plain,(
  ( ! [X21,X22] : (~sP0(X21,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X22) | p(s(bool,X22)) | p(s(bool,X21))) ) | $spl8),
  inference(resolution,[],[f46994,f367])).
fof(f77666,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X13))) ) | $spl12),
  inference(forward_demodulation,[],[f77646,f568])).
fof(f77646,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X13))) ) | $spl12),
  inference(resolution,[],[f906,f1710])).
fof(f1710,plain,(
  ( ! [X4,X5] : (~sP0(X4,'F_0',X5) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(subsumption_resolution,[],[f1704,f507])).
fof(f1704,plain,(
  ( ! [X4,X5] : (~sP0(X4,'F_0',X5) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,'F_0') != s(bool,X4)) )),
  inference(resolution,[],[f1688,f374])).
fof(f1688,plain,(
  ( ! [X2,X3] : (sP3(X2,'F_0',X3) | ~sP0(X3,'F_0',X2) | p(s(bool,X3))) )),
  inference(resolution,[],[f1668,f361])).
fof(f1668,plain,(
  ( ! [X4,X5] : (~sP2(X4,X5,'F_0') | sP3(X4,'F_0',X5) | ~sP0(X5,'F_0',X4)) )),
  inference(subsumption_resolution,[],[f1666,f539])).
fof(f1666,plain,(
  ( ! [X4,X5] : (sP3(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP2(X4,X5,'F_0') | ~sP0(X5,'F_0',X4)) )),
  inference(resolution,[],[f1663,f357])).
fof(f77665,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | p(s(bool,X12))) ) | $spl12),
  inference(forward_demodulation,[],[f77645,f568])).
fof(f77645,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(resolution,[],[f906,f1700])).
fof(f77664,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X11))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77644,f568])).
fof(f77644,plain,(
  ( ! [X11] : (p(s(bool,X11)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f906,f24353])).
fof(f24353,plain,(
  ( ! [X19,X20] : (~sP0(X19,X20,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X20)) | p(s(bool,X19))) ) | $spl109),
  inference(resolution,[],[f24310,f367])).
fof(f77663,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X10))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77643,f568])).
fof(f77643,plain,(
  ( ! [X10] : (p(s(bool,X10)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f906,f42305])).
fof(f42305,plain,(
  ( ! [X30,X28,X29] : (~sP0(X28,X29,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl109),
  inference(resolution,[],[f42160,f367])).
fof(f77662,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X9))) ) | ($spl12 | $spl56)),
  inference(forward_demodulation,[],[f77642,f568])).
fof(f77642,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl12 | $spl56)),
  inference(resolution,[],[f906,f16263])).
fof(f16263,plain,(
  ( ! [X19,X18] : (~sP0(X18,X19,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | p(s(bool,X19)) | p(s(bool,X18))) ) | $spl56),
  inference(resolution,[],[f16105,f367])).
fof(f77661,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X8))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77641,f568])).
fof(f77641,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f906,f47026])).
fof(f47026,plain,(
  ( ! [X19,X20] : (~sP0(X19,X20,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | p(s(bool,X20)) | p(s(bool,X19))) ) | $spl8),
  inference(resolution,[],[f46994,f367])).
fof(f77658,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~sP1(X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X2) | ~sP2(X2,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X3)) ) | $spl212),
  inference(cnf_transformation,[],[f77658_D])).
fof(f77658_D,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~sP1(X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X2) | ~sP2(X2,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X3)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f77657,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | sP3(X2,X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~sP2(X2,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X3) | ~sP1(X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X2)) ) | $spl12),
  inference(forward_demodulation,[],[f77638,f568])).
fof(f77638,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~sP2(X2,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X3) | ~sP1(X3,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X2)) ) | $spl12),
  inference(resolution,[],[f906,f356])).
fof(f77656,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl12),
  inference(forward_demodulation,[],[f77637,f568])).
fof(f77637,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl12),
  inference(resolution,[],[f906,f367])).
fof(f77636,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X15))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77621,f568])).
fof(f77621,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f905,f24355])).
fof(f24355,plain,(
  ( ! [X24,X23] : (~sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | p(s(bool,X24)) | p(s(bool,X23))) ) | $spl109),
  inference(resolution,[],[f24310,f367])).
fof(f905,plain,(
  ( ! [X24,X23] : (sP0(X23,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X24)) ) | $spl12),
  inference(resolution,[],[f835,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f276])).
fof(f77635,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X14))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77620,f568])).
fof(f77620,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f905,f42307])).
fof(f42307,plain,(
  ( ! [X35,X36,X34] : (~sP0('const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num','const.num.|48|_0')),X34,X35) | p(s(bool,X35)) | p(s(bool,X34))) ) | $spl109),
  inference(resolution,[],[f42160,f367])).
fof(f77634,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X13))) ) | ($spl12 | $spl56)),
  inference(forward_demodulation,[],[f77619,f568])).
fof(f77619,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl12 | $spl56)),
  inference(resolution,[],[f905,f16265])).
fof(f16265,plain,(
  ( ! [X23,X22] : (~sP0('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X22,X23) | p(s(bool,X23)) | p(s(bool,X22))) ) | $spl56),
  inference(resolution,[],[f16105,f367])).
fof(f77633,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X12))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77618,f568])).
fof(f77618,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f905,f47028])).
fof(f47028,plain,(
  ( ! [X24,X23] : (~sP0('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X23,X24) | p(s(bool,X24)) | p(s(bool,X23))) ) | $spl8),
  inference(resolution,[],[f46994,f367])).
fof(f77632,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X11))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77617,f568])).
fof(f77617,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X11))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f905,f24353])).
fof(f77631,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X10))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77616,f568])).
fof(f77616,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X10))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f905,f42305])).
fof(f77630,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X9))) ) | ($spl12 | $spl56)),
  inference(forward_demodulation,[],[f77615,f568])).
fof(f77615,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X9))) ) | ($spl12 | $spl56)),
  inference(resolution,[],[f905,f16263])).
fof(f77629,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X8))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77614,f568])).
fof(f77614,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X8))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f905,f47026])).
fof(f77627,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5)) ) | $spl210),
  inference(cnf_transformation,[],[f77627_D])).
fof(f77627_D,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f77626,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | sP3(X4,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5)) ) | $spl12),
  inference(forward_demodulation,[],[f77625,f568])).
fof(f77625,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(subsumption_resolution,[],[f77624,f900])).
fof(f900,plain,(
  ( ! [X14,X13] : (sP2(X13,X14,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(resolution,[],[f835,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f271])).
fof(f77624,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~sP2(X4,X5,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(subsumption_resolution,[],[f77612,f366])).
fof(f77612,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~sP2(X4,X5,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~sP1('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5,X4)) ) | $spl12),
  inference(resolution,[],[f905,f357])).
fof(f77622,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl12),
  inference(forward_demodulation,[],[f77610,f568])).
fof(f77610,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X1))) ) | $spl12),
  inference(resolution,[],[f905,f367])).
fof(f77609,plain,(
  ( ! [X17] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X17))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77592,f568])).
fof(f77592,plain,(
  ( ! [X17] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X17))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f904,f24355])).
fof(f904,plain,(
  ( ! [X21,X22] : (sP0(X21,X22,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(resolution,[],[f835,f368])).
fof(f77608,plain,(
  ( ! [X16] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X16))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77591,f568])).
fof(f77591,plain,(
  ( ! [X16] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X16))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f904,f42307])).
fof(f77607,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X15))) ) | ($spl12 | $spl56)),
  inference(forward_demodulation,[],[f77590,f568])).
fof(f77590,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X15))) ) | ($spl12 | $spl56)),
  inference(resolution,[],[f904,f16265])).
fof(f77606,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X14))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77589,f568])).
fof(f77589,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X14))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f904,f47028])).
fof(f77605,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X13))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77588,f568])).
fof(f77588,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X13))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f904,f24354])).
fof(f77604,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X12))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77587,f568])).
fof(f77587,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X12))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f904,f42306])).
fof(f77603,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X11))) ) | ($spl12 | $spl56)),
  inference(forward_demodulation,[],[f77586,f568])).
fof(f77586,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X11))) ) | ($spl12 | $spl56)),
  inference(resolution,[],[f904,f16264])).
fof(f77602,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X10))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77585,f568])).
fof(f77585,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X10))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f904,f47027])).
fof(f77601,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X9))) ) | $spl12),
  inference(forward_demodulation,[],[f77584,f568])).
fof(f77584,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f904,f1710])).
fof(f77600,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | s(bool,'F_0') != s(bool,X8)) ) | $spl12),
  inference(forward_demodulation,[],[f77583,f568])).
fof(f77583,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | s(bool,'F_0') != s(bool,X8)) ) | $spl12),
  inference(resolution,[],[f904,f1700])).
fof(f77599,plain,(
  ( ! [X6,X7] : (sP3('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7)) ) | $spl208),
  inference(cnf_transformation,[],[f77599_D])).
fof(f77599_D,plain,(
  ( ! [X6,X7] : (sP3('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f77598,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | sP3('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7)) ) | $spl12),
  inference(forward_demodulation,[],[f77597,f568])).
fof(f77597,plain,(
  ( ! [X6,X7] : (sP3('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(subsumption_resolution,[],[f77596,f903])).
fof(f903,plain,(
  ( ! [X19,X20] : (sP1(X19,X20,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(resolution,[],[f835,f364])).
fof(f77596,plain,(
  ( ! [X6,X7] : (sP3('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~sP1(X6,X7,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(subsumption_resolution,[],[f77582,f362])).
fof(f77582,plain,(
  ( ! [X6,X7] : (sP3('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~sP2('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X7,X6) | ~sP1(X6,X7,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(resolution,[],[f904,f358])).
fof(f77593,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl12),
  inference(forward_demodulation,[],[f77579,f568])).
fof(f77579,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f904,f367])).
fof(f77578,plain,(
  ( ! [X17] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X17))) ) | ($spl12 | $spl54 | $spl109)),
  inference(forward_demodulation,[],[f77563,f568])).
fof(f77563,plain,(
  ( ! [X17] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X17))) ) | ($spl12 | $spl54 | $spl109)),
  inference(resolution,[],[f903,f43409])).
fof(f43409,plain,(
  ( ! [X30,X29] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X29,X30) | p(s(bool,X30)) | ~p(s(bool,X29))) ) | ($spl54 | $spl109)),
  inference(subsumption_resolution,[],[f43242,f514])).
fof(f43242,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X29,X30)) ) | ($spl54 | $spl109)),
  inference(superposition,[],[f363,f42131])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f274])).
fof(f77577,plain,(
  ( ! [X16] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X16))) ) | ($spl12 | $spl54)),
  inference(forward_demodulation,[],[f77562,f568])).
fof(f77562,plain,(
  ( ! [X16] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X16))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f903,f16965])).
fof(f16965,plain,(
  ( ! [X30,X29] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X29,X30) | p(s(bool,X30)) | ~p(s(bool,X29))) ) | $spl54),
  inference(subsumption_resolution,[],[f16792,f514])).
fof(f16792,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X29,X30)) ) | $spl54),
  inference(superposition,[],[f363,f15745])).
fof(f77576,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X15))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77561,f568])).
fof(f77561,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X15))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f903,f24493])).
fof(f24493,plain,(
  ( ! [X12,X13] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X12,X13) | p(s(bool,X13)) | ~p(s(bool,X12))) ) | $spl109),
  inference(resolution,[],[f24342,f363])).
fof(f77575,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X14))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77560,f568])).
fof(f77560,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X14))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f903,f47638])).
fof(f47638,plain,(
  ( ! [X33,X32] : (~sP1('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X32,X33) | p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl8),
  inference(subsumption_resolution,[],[f47466,f514])).
fof(f47466,plain,(
  ( ! [X33,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP1('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X32,X33)) ) | $spl8),
  inference(superposition,[],[f363,f680])).
fof(f77574,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X13))) ) | $spl12),
  inference(forward_demodulation,[],[f77559,f568])).
fof(f77559,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X13))) ) | $spl12),
  inference(resolution,[],[f903,f575])).
fof(f575,plain,(
  ( ! [X14,X13] : (~sP1('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X13,X14) | p(s(bool,X14)) | ~p(s(bool,X13))) )),
  inference(resolution,[],[f350,f363])).
fof(f77573,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X12))) ) | ($spl12 | $spl54 | $spl109)),
  inference(forward_demodulation,[],[f77558,f568])).
fof(f77558,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X12))) ) | ($spl12 | $spl54 | $spl109)),
  inference(resolution,[],[f903,f43410])).
fof(f43410,plain,(
  ( ! [X31,X32] : (~sP1(X31,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X32) | p(s(bool,X32)) | ~p(s(bool,X31))) ) | ($spl54 | $spl109)),
  inference(subsumption_resolution,[],[f43243,f514])).
fof(f43243,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | p(s(bool,X32)) | ~sP1(X31,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X32)) ) | ($spl54 | $spl109)),
  inference(superposition,[],[f363,f42131])).
fof(f77572,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X11))) ) | ($spl12 | $spl54)),
  inference(forward_demodulation,[],[f77557,f568])).
fof(f77557,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X11))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f903,f16966])).
fof(f16966,plain,(
  ( ! [X31,X32] : (~sP1(X31,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X32) | p(s(bool,X32)) | ~p(s(bool,X31))) ) | $spl54),
  inference(subsumption_resolution,[],[f16793,f514])).
fof(f16793,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | p(s(bool,X32)) | ~sP1(X31,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X32)) ) | $spl54),
  inference(superposition,[],[f363,f15745])).
fof(f77571,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X10))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77556,f568])).
fof(f77556,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X10))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f903,f24494])).
fof(f24494,plain,(
  ( ! [X14,X15] : (~sP1(X14,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X15) | p(s(bool,X15)) | ~p(s(bool,X14))) ) | $spl109),
  inference(resolution,[],[f24342,f363])).
fof(f77570,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X9))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77555,f568])).
fof(f77555,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X9))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f903,f47639])).
fof(f47639,plain,(
  ( ! [X35,X34] : (~sP1(X34,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X35) | p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl8),
  inference(subsumption_resolution,[],[f47467,f514])).
fof(f47467,plain,(
  ( ! [X35,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X34)) | p(s(bool,X35)) | ~sP1(X34,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X35)) ) | $spl8),
  inference(superposition,[],[f363,f680])).
fof(f77569,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X8))) ) | $spl12),
  inference(forward_demodulation,[],[f77554,f568])).
fof(f77554,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X8))) ) | $spl12),
  inference(resolution,[],[f903,f576])).
fof(f576,plain,(
  ( ! [X15,X16] : (~sP1(X15,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X16) | p(s(bool,X16)) | ~p(s(bool,X15))) )),
  inference(resolution,[],[f350,f363])).
fof(f77568,plain,(
  ( ! [X6,X7] : (sP3('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7) | ~sP0(X7,X6,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl206),
  inference(cnf_transformation,[],[f77568_D])).
fof(f77568_D,plain,(
  ( ! [X6,X7] : (sP3('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7) | ~sP0(X7,X6,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f77567,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | sP3('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7) | ~sP0(X7,X6,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(forward_demodulation,[],[f77566,f568])).
fof(f77566,plain,(
  ( ! [X6,X7] : (sP3('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~sP0(X7,X6,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(subsumption_resolution,[],[f77553,f362])).
fof(f77553,plain,(
  ( ! [X6,X7] : (sP3('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X6,X7) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~sP2('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X7,X6) | ~sP0(X7,X6,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(resolution,[],[f903,f358])).
fof(f77564,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl12),
  inference(forward_demodulation,[],[f77550,f568])).
fof(f77550,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f903,f363])).
fof(f77549,plain,(
  ( ! [X20] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X20))) ) | ($spl12 | $spl54 | $spl109)),
  inference(forward_demodulation,[],[f77530,f568])).
fof(f77530,plain,(
  ( ! [X20] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X20))) ) | ($spl12 | $spl54 | $spl109)),
  inference(resolution,[],[f900,f43406])).
fof(f43406,plain,(
  ( ! [X17,X18] : (~sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17,X18) | p(s(bool,X18)) | ~p(s(bool,X17))) ) | ($spl54 | $spl109)),
  inference(subsumption_resolution,[],[f43236,f514])).
fof(f43236,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17,X18)) ) | ($spl54 | $spl109)),
  inference(superposition,[],[f359,f42131])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f271])).
fof(f77548,plain,(
  ( ! [X19] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X19))) ) | ($spl12 | $spl54)),
  inference(forward_demodulation,[],[f77529,f568])).
fof(f77529,plain,(
  ( ! [X19] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X19))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f900,f16963])).
fof(f16963,plain,(
  ( ! [X17,X18] : (~sP2('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X17,X18) | p(s(bool,X18)) | ~p(s(bool,X17))) ) | $spl54),
  inference(subsumption_resolution,[],[f16786,f514])).
fof(f16786,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~sP2('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X17,X18)) ) | $spl54),
  inference(superposition,[],[f359,f15745])).
fof(f77547,plain,(
  ( ! [X18] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X18))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77528,f568])).
fof(f77528,plain,(
  ( ! [X18] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X18))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f900,f24490])).
fof(f24490,plain,(
  ( ! [X6,X7] : (~sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X6,X7) | p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl109),
  inference(resolution,[],[f24342,f359])).
fof(f77546,plain,(
  ( ! [X17] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X17))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77527,f568])).
fof(f77527,plain,(
  ( ! [X17] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X17))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f900,f47635])).
fof(f47635,plain,(
  ( ! [X21,X20] : (~sP2('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X20,X21) | p(s(bool,X21)) | ~p(s(bool,X20))) ) | $spl8),
  inference(subsumption_resolution,[],[f47460,f514])).
fof(f47460,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | p(s(bool,X21)) | ~sP2('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X20,X21)) ) | $spl8),
  inference(superposition,[],[f359,f680])).
fof(f77545,plain,(
  ( ! [X16] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X16))) ) | $spl12),
  inference(forward_demodulation,[],[f77526,f568])).
fof(f77526,plain,(
  ( ! [X16] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X16))) ) | $spl12),
  inference(resolution,[],[f900,f572])).
fof(f572,plain,(
  ( ! [X8,X7] : (~sP2('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X7,X8) | p(s(bool,X8)) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f350,f359])).
fof(f77543,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl204),
  inference(cnf_transformation,[],[f77543_D])).
fof(f77543_D,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f77542,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(forward_demodulation,[],[f77541,f568])).
fof(f77541,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | s(bool,X14) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(forward_demodulation,[],[f77524,f568])).
fof(f77524,plain,(
  ( ! [X14] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | s(bool,X14) = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(resolution,[],[f900,f3855])).
fof(f3855,plain,(
  ( ! [X8,X9] : (~sP2('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X8,X9) | p(s(bool,X9)) | s(bool,X8) = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f3854,f578])).
fof(f578,plain,(
  ( ! [X19,X20] : (sP0(X19,X20,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) )),
  inference(resolution,[],[f350,f368])).
fof(f3854,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | p(s(bool,X9)) | ~sP2('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X8,X9) | ~sP0(X8,X9,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) )),
  inference(subsumption_resolution,[],[f3843,f366])).
fof(f3843,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | p(s(bool,X9)) | ~sP2('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X8,X9) | ~sP1(X9,X8,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~sP0(X8,X9,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) )),
  inference(resolution,[],[f582,f357])).
fof(f582,plain,(
  ( ! [X28,X27] : (~sP3('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X28,X27) | s(bool,X27) = s(bool,X28)) )),
  inference(resolution,[],[f350,f373])).
fof(f77540,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X13))) ) | ($spl12 | $spl54 | $spl109)),
  inference(forward_demodulation,[],[f77523,f568])).
fof(f77523,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X13))) ) | ($spl12 | $spl54 | $spl109)),
  inference(resolution,[],[f900,f43407])).
fof(f43407,plain,(
  ( ! [X19,X20] : (~sP2(X19,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | p(s(bool,X20)) | ~p(s(bool,X19))) ) | ($spl54 | $spl109)),
  inference(subsumption_resolution,[],[f43237,f514])).
fof(f43237,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | p(s(bool,X20)) | ~sP2(X19,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20)) ) | ($spl54 | $spl109)),
  inference(superposition,[],[f359,f42131])).
fof(f77539,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X12))) ) | ($spl12 | $spl54)),
  inference(forward_demodulation,[],[f77522,f568])).
fof(f77522,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X12))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f900,f16964])).
fof(f16964,plain,(
  ( ! [X19,X20] : (~sP2(X19,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X20) | p(s(bool,X20)) | ~p(s(bool,X19))) ) | $spl54),
  inference(subsumption_resolution,[],[f16787,f514])).
fof(f16787,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | p(s(bool,X20)) | ~sP2(X19,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X20)) ) | $spl54),
  inference(superposition,[],[f359,f15745])).
fof(f77538,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X11))) ) | ($spl12 | $spl109)),
  inference(forward_demodulation,[],[f77521,f568])).
fof(f77521,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X11))) ) | ($spl12 | $spl109)),
  inference(resolution,[],[f900,f24491])).
fof(f24491,plain,(
  ( ! [X8,X9] : (~sP2(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X9) | p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl109),
  inference(resolution,[],[f24342,f359])).
fof(f77537,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X10))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f77520,f568])).
fof(f77520,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X10))) ) | ($spl8 | $spl12)),
  inference(resolution,[],[f900,f47636])).
fof(f47636,plain,(
  ( ! [X23,X22] : (~sP2(X22,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X23) | p(s(bool,X23)) | ~p(s(bool,X22))) ) | $spl8),
  inference(subsumption_resolution,[],[f47461,f514])).
fof(f47461,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X22)) | p(s(bool,X23)) | ~sP2(X22,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X23)) ) | $spl8),
  inference(superposition,[],[f359,f680])).
fof(f77536,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X9))) ) | $spl12),
  inference(forward_demodulation,[],[f77519,f568])).
fof(f77519,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X9))) ) | $spl12),
  inference(resolution,[],[f900,f573])).
fof(f573,plain,(
  ( ! [X10,X9] : (~sP2(X9,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X10) | p(s(bool,X10)) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f350,f359])).
fof(f77535,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5) | ~sP0(X5,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X4)) ) | $spl202),
  inference(cnf_transformation,[],[f77535_D])).
fof(f77535_D,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5) | ~sP0(X5,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X4)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f77534,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | sP3(X4,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5) | ~sP0(X5,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X4)) ) | $spl12),
  inference(forward_demodulation,[],[f77533,f568])).
fof(f77533,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~sP0(X5,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X4)) ) | $spl12),
  inference(subsumption_resolution,[],[f77516,f366])).
fof(f77516,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~sP1('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X5,X4) | ~sP0(X5,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X4)) ) | $spl12),
  inference(resolution,[],[f900,f357])).
fof(f77531,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl12),
  inference(forward_demodulation,[],[f77514,f568])).
fof(f77514,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f900,f359])).
fof(f77363,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X3))))) ) | $spl200),
  inference(cnf_transformation,[],[f77363_D])).
fof(f77363_D,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X3))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f77388,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X3))))) ) | $spl12),
  inference(forward_demodulation,[],[f77375,f568])).
fof(f77375,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X3))))) ) | $spl12),
  inference(resolution,[],[f952,f7006])).
fof(f7006,plain,(
  ( ! [X366,X367] : (~sP9(X367,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X366))) | p(s(bool,X367)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X366))))) )),
  inference(subsumption_resolution,[],[f6943,f511])).
fof(f511,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f333])).
fof(f333,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X19] : ((p(s(bool,X19)) | p(s(bool,'F_0'))) <=> p(s(bool,X19)))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.OR_CLAUSES|split|3')).
fof(f6943,plain,(
  ( ! [X366,X367] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X366)))) | p(s(bool,'F_0')) | p(s(bool,X367)) | ~sP9(X367,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X366)))) )),
  inference(superposition,[],[f484,f568])).
fof(f484,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP9(X0,X1)) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(nnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1] : (sP9(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f952,plain,(
  ( ! [X104] : (sP9('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X104)) ) | $spl12),
  inference(resolution,[],[f835,f486])).
fof(f486,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f319])).
fof(f77383,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X2))) ) | $spl12),
  inference(forward_demodulation,[],[f77382,f568])).
fof(f77382,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X2))) ) | $spl12),
  inference(subsumption_resolution,[],[f77370,f511])).
fof(f77370,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X2))) ) | $spl12),
  inference(resolution,[],[f952,f484])).
fof(f77381,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X1))) ) | $spl12),
  inference(forward_demodulation,[],[f77380,f568])).
fof(f77380,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X1))) ) | $spl12),
  inference(subsumption_resolution,[],[f77369,f511])).
fof(f77369,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X1))) ) | $spl12),
  inference(resolution,[],[f952,f488])).
fof(f488,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP9(X0,X1)) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(nnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1] : (sP9(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f242,f263])).
fof(f242,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X14,X13] : ((~(p(s(bool,X13)) | p(s(bool,X14))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X13)) => p(s(bool,'F_0'))) => (~p(s(bool,X14)) => p(s(bool,'F_0')))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.sat.sat_9_prove_ltb')).
fof(f77368,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f952,f7007])).
fof(f7007,plain,(
  ( ! [X368,X369] : (~sP9('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X368)),X369) | p(s(bool,X369)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X368))))) )),
  inference(subsumption_resolution,[],[f6944,f511])).
fof(f6944,plain,(
  ( ! [X368,X369] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X368)))) | p(s(bool,'F_0')) | p(s(bool,X369)) | ~sP9('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X368)),X369)) )),
  inference(superposition,[],[f484,f568])).
fof(f77362,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X3))))) ) | $spl12),
  inference(forward_demodulation,[],[f77349,f568])).
fof(f77349,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X3))))) ) | $spl12),
  inference(resolution,[],[f951,f7007])).
fof(f951,plain,(
  ( ! [X103] : (sP9(X103,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(resolution,[],[f835,f485])).
fof(f485,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f319])).
fof(f77357,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X2))) ) | $spl12),
  inference(forward_demodulation,[],[f77356,f568])).
fof(f77356,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(subsumption_resolution,[],[f77344,f511])).
fof(f77344,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f951,f484])).
fof(f77355,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X1))) ) | $spl12),
  inference(forward_demodulation,[],[f77354,f568])).
fof(f77354,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(subsumption_resolution,[],[f77343,f511])).
fof(f77343,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f951,f488])).
fof(f77342,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl12),
  inference(resolution,[],[f951,f7006])).
fof(f837,plain,(
  p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl14),
  inference(cnf_transformation,[],[f837_D])).
fof(f837_D,plain,(
  p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f77335,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl8),
  inference(forward_demodulation,[],[f52249,f568])).
fof(f52249,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl8),
  inference(subsumption_resolution,[],[f52069,f511])).
fof(f52069,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl8),
  inference(superposition,[],[f400,f4706])).
fof(f836,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl15),
  inference(cnf_transformation,[],[f836_D])).
fof(f836_D,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f77334,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | $spl199),
  inference(cnf_transformation,[],[f77334_D])).
fof(f77334_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f77279,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl7),
  inference(superposition,[],[f47277,f568])).
fof(f47277,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl7),
  inference(superposition,[],[f679,f398])).
fof(f77329,plain,(
  $false | ($spl7 | $spl14)),
  inference(subsumption_resolution,[],[f77319,f401])).
fof(f77319,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | ($spl7 | $spl14)),
  inference(trivial_inequality_removal,[],[f77316])).
fof(f77316,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | ($spl7 | $spl14)),
  inference(superposition,[],[f77299,f528])).
fof(f528,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f198])).
fof(f198,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X19] : (s(bool,'F_0') = s(bool,X19) <=> ~p(s(bool,X19)))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.EQ_CLAUSES|split|3')).
fof(f77299,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ($spl7 | $spl14)),
  inference(subsumption_resolution,[],[f77298,f544])).
fof(f77298,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ($spl7 | $spl14)),
  inference(forward_demodulation,[],[f77279,f47794])).
fof(f47794,plain,(
  s(bool,'T_0') = s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl14),
  inference(resolution,[],[f837,f530])).
fof(f77328,plain,(
  $false | ($spl7 | $spl14)),
  inference(subsumption_resolution,[],[f77320,f401])).
fof(f77320,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | ($spl7 | $spl14)),
  inference(trivial_inequality_removal,[],[f77315])).
fof(f77315,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | ($spl7 | $spl14)),
  inference(superposition,[],[f77299,f526])).
fof(f526,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f343])).
fof(f77327,plain,(
  $false | ($spl7 | $spl14)),
  inference(subsumption_resolution,[],[f77322,f401])).
fof(f77322,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | ($spl7 | $spl14)),
  inference(trivial_inequality_removal,[],[f77313])).
fof(f77313,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | ($spl7 | $spl14)),
  inference(superposition,[],[f77299,f507])).
fof(f77301,plain,(
  ( ! [X20] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl196),
  inference(cnf_transformation,[],[f77301_D])).
fof(f77301_D,plain,(
  ( ! [X20] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f77302,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X22),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X22),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(subsumption_resolution,[],[f77287,f525])).
fof(f77287,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X22),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X22),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X23))) ) | $spl7),
  inference(superposition,[],[f47277,f383])).
fof(f77300,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X21) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(subsumption_resolution,[],[f77286,f525])).
fof(f77286,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X21) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X21)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(superposition,[],[f47277,f383])).
fof(f72647,plain,(
  ( ! [X246] : (sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X246)),'T_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X246))))) ) | $spl192),
  inference(cnf_transformation,[],[f72647_D])).
fof(f72647_D,plain,(
  ( ! [X246] : (sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X246)),'T_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X246))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f72636,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X43)),'T_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X43))))) )),
  inference(duplicate_literal_removal,[],[f72600])).
fof(f72600,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X43)),'T_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X43)))) | ~p(s(bool,X42))) )),
  inference(superposition,[],[f6933,f530])).
fof(f6933,plain,(
  ( ! [X346,X347] : (~p(s(bool,X347)) | sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X346)),X347) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X346))))) )),
  inference(superposition,[],[f473,f568])).
fof(f473,plain,(
  ( ! [X0,X1] : (sP8(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1] : ((~sP8(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1)))),
  inference(rectify,[],[f313])).
fof(f313,plain,(
  ! [X1,X0] : ((~sP8(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP8(X1,X0)))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X1,X0] : ((~sP8(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP8(X1,X0)))),
  inference(nnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X1,X0] : (sP8(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f72637,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X41)),'T_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X41))))) )),
  inference(duplicate_literal_removal,[],[f72599])).
fof(f72599,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X41)),'T_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X41)))) | ~p(s(bool,X40))) )),
  inference(superposition,[],[f6933,f532])).
fof(f71806,plain,(
  ( ! [X255,X256] : (p(s(bool,X255)) | ~sP9(X255,X256) | p(s(bool,X256))) ) | $spl188),
  inference(cnf_transformation,[],[f71806_D])).
fof(f71806_D,plain,(
  ( ! [X255,X256] : (p(s(bool,X255)) | ~sP9(X255,X256) | p(s(bool,X256))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f72648,plain,(
  ( ! [X254] : (sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X254)),'F_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X254))))) ) | $spl194),
  inference(cnf_transformation,[],[f72648_D])).
fof(f72648_D,plain,(
  ( ! [X254] : (sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X254)),'F_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X254))))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f72563,plain,(
  ( ! [X261,X260,X262] : (sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X260)),'F_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X260)))) | p(s(bool,X261)) | p(s(bool,X262)) | ~sP9(X261,X262)) )),
  inference(resolution,[],[f6933,f484])).
fof(f72561,plain,(
  ( ! [X255,X254,X256] : (sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X254)),'F_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X254)))) | p(s(bool,X255)) | p(s(bool,X256)) | ~sP9(X255,X256)) )),
  inference(resolution,[],[f6933,f488])).
fof(f72557,plain,(
  ( ! [X249,X248] : (sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X248)),'T_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X248)))) | ~p(s(bool,X249))) )),
  inference(resolution,[],[f6933,f514])).
fof(f72556,plain,(
  ( ! [X246,X247] : (sP8('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X246)),'T_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X246)))) | ~p(s(bool,X247))) )),
  inference(resolution,[],[f6933,f522])).
fof(f522,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f341])).
fof(f341,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X19] : ((p(s(bool,X19)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.OR_CLAUSES|split|1')).
fof(f71805,plain,(
  ( ! [X246] : (sP8('T_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X246))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X246))))) ) | $spl186),
  inference(cnf_transformation,[],[f71805_D])).
fof(f71805_D,plain,(
  ( ! [X246] : (sP8('T_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X246))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X246))))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f71794,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP8('T_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X43))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X43))))) )),
  inference(duplicate_literal_removal,[],[f71758])).
fof(f71758,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP8('T_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X43))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X43)))) | ~p(s(bool,X42))) )),
  inference(superposition,[],[f6931,f530])).
fof(f6931,plain,(
  ( ! [X343,X342] : (~p(s(bool,X343)) | sP8(X343,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X342))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X342))))) )),
  inference(superposition,[],[f470,f568])).
fof(f470,plain,(
  ( ! [X0,X1] : (sP8(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f71795,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP8('T_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X41))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X41))))) )),
  inference(duplicate_literal_removal,[],[f71757])).
fof(f71757,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP8('T_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X41))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X41)))) | ~p(s(bool,X40))) )),
  inference(superposition,[],[f6931,f532])).
fof(f71807,plain,(
  ( ! [X254] : (sP8('F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X254))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X254))))) ) | $spl190),
  inference(cnf_transformation,[],[f71807_D])).
fof(f71807_D,plain,(
  ( ! [X254] : (sP8('F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X254))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X254))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f71721,plain,(
  ( ! [X261,X260,X262] : (sP8('F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X260))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X260)))) | p(s(bool,X261)) | p(s(bool,X262)) | ~sP9(X261,X262)) )),
  inference(resolution,[],[f6931,f484])).
fof(f71719,plain,(
  ( ! [X255,X254,X256] : (sP8('F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X254))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X254)))) | p(s(bool,X255)) | p(s(bool,X256)) | ~sP9(X255,X256)) )),
  inference(resolution,[],[f6931,f488])).
fof(f71715,plain,(
  ( ! [X249,X248] : (sP8('T_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X248))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X248)))) | ~p(s(bool,X249))) )),
  inference(resolution,[],[f6931,f514])).
fof(f71714,plain,(
  ( ! [X246,X247] : (sP8('T_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),X246))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),X246)))) | ~p(s(bool,X247))) )),
  inference(resolution,[],[f6931,f522])).
fof(f69339,plain,(
  $false | ($spl10 | $spl48 | $spl82)),
  inference(subsumption_resolution,[],[f69316,f545])).
fof(f545,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f41])).
fof(f41,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.TRUTH')).
fof(f69316,plain,(
  ~p(s(bool,'T_0')) | ($spl10 | $spl48 | $spl82)),
  inference(superposition,[],[f69246,f47397])).
fof(f47397,plain,(
  s(bool,'T_0') = s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | $spl10),
  inference(resolution,[],[f833,f530])).
fof(f69246,plain,(
  ( ! [X20] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X20))))) ) | ($spl48 | $spl82)),
  inference(subsumption_resolution,[],[f18449,f10706])).
fof(f69243,plain,(
  $false | ($spl14 | $spl42 | $spl78)),
  inference(subsumption_resolution,[],[f69221,f545])).
fof(f69221,plain,(
  ~p(s(bool,'T_0')) | ($spl14 | $spl42 | $spl78)),
  inference(superposition,[],[f69150,f47794])).
fof(f69150,plain,(
  ( ! [X20] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | ($spl42 | $spl78)),
  inference(subsumption_resolution,[],[f18421,f7332])).
fof(f60390,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) ) | $spl182),
  inference(cnf_transformation,[],[f60390_D])).
fof(f60390_D,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f69115,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X26)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25)) ) | ($spl8 | $spl109)),
  inference(resolution,[],[f48436,f56454])).
fof(f56454,plain,(
  ( ! [X21,X22,X20] : (sP2(X21,'const.arithmetic.<=_2'(s('type.num.num',X20),s('type.num.num','const.num.|48|_0')),X22) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) ) | $spl109),
  inference(resolution,[],[f42130,f361])).
fof(f42130,plain,(
  ( ! [X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) ) | $spl109),
  inference(forward_demodulation,[],[f42129,f41769])).
fof(f42129,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl109),
  inference(subsumption_resolution,[],[f42128,f538])).
fof(f42128,plain,(
  ( ! [X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl109),
  inference(forward_demodulation,[],[f41977,f41769])).
fof(f41977,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X7)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl109),
  inference(superposition,[],[f41769,f501])).
fof(f50761,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X14))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl172),
  inference(cnf_transformation,[],[f50761_D])).
fof(f50761_D,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X14))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f69114,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X23))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X24))) ) | ($spl8 | $spl109)),
  inference(resolution,[],[f48436,f50196])).
fof(f50196,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X20))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X21)) ) | $spl109),
  inference(resolution,[],[f42185,f361])).
fof(f42185,plain,(
  ( ! [X71] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X71))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl109),
  inference(forward_demodulation,[],[f42184,f41769])).
fof(f42184,plain,(
  ( ! [X71] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X71))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl109),
  inference(subsumption_resolution,[],[f42048,f538])).
fof(f42048,plain,(
  ( ! [X71] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X71)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X71))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl109),
  inference(superposition,[],[f497,f41769])).
fof(f497,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X21,X20] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X21))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X20),s('type.num.num',X21)))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.numeral.numeral_lte|split|5')).
fof(f20365,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl121),
  inference(cnf_transformation,[],[f20365_D])).
fof(f20365_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f69111,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,X20))) ) | ($spl8 | $spl54)),
  inference(resolution,[],[f48436,f19859])).
fof(f19859,plain,(
  ( ! [X10,X9] : (sP2(X9,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X10)) ) | $spl54),
  inference(resolution,[],[f16958,f361])).
fof(f16958,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl54),
  inference(subsumption_resolution,[],[f16773,f522])).
fof(f16773,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl54),
  inference(superposition,[],[f497,f15745])).
fof(f29036,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | $spl141),
  inference(cnf_transformation,[],[f29036_D])).
fof(f29036_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f69109,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X17))) ) | ($spl8 | $spl109)),
  inference(resolution,[],[f48436,f28384])).
fof(f28384,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),X9)) ) | $spl109),
  inference(resolution,[],[f27869,f361])).
fof(f27869,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) | $spl109),
  inference(resolution,[],[f24341,f464])).
fof(f24341,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl109),
  inference(resolution,[],[f24310,f498])).
fof(f498,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f324])).
fof(f62463,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)))) ) | $spl184),
  inference(cnf_transformation,[],[f62463_D])).
fof(f62463_D,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f62408,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | p(s(bool,X19)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)))) ) | $spl109),
  inference(resolution,[],[f59943,f1700])).
fof(f59943,plain,(
  ( ! [X37,X35,X36] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X35)),X36,X37) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X35)) ) | $spl109),
  inference(resolution,[],[f56446,f370])).
fof(f56446,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3)) ) | $spl109),
  inference(resolution,[],[f42130,f463])).
fof(f60329,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X15))) ) | $spl109),
  inference(resolution,[],[f56454,f1771])).
fof(f57450,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl181),
  inference(cnf_transformation,[],[f57450_D])).
fof(f57450_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f57423,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl109),
  inference(resolution,[],[f56747,f1700])).
fof(f56747,plain,(
  ( ! [X23,X22] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X22,X23)) ) | $spl109),
  inference(resolution,[],[f56623,f370])).
fof(f56623,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl109),
  inference(resolution,[],[f56606,f463])).
fof(f56606,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl109),
  inference(subsumption_resolution,[],[f56546,f28004])).
fof(f28004,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl109),
  inference(subsumption_resolution,[],[f28003,f539])).
fof(f28003,plain,(
  p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl109),
  inference(forward_demodulation,[],[f27968,f562])).
fof(f562,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X21] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X21))),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.numeral.numeral_lte|split|2')).
fof(f27968,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl109),
  inference(superposition,[],[f24341,f387])).
fof(f387,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f282])).
fof(f56546,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl109),
  inference(superposition,[],[f42130,f555])).
fof(f56947,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | $spl177),
  inference(cnf_transformation,[],[f56947_D])).
fof(f56947_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f56914,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | $spl109),
  inference(resolution,[],[f56629,f1771])).
fof(f56629,plain,(
  ( ! [X10,X9] : (sP2(X9,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X10)) ) | $spl109),
  inference(resolution,[],[f56606,f361])).
fof(f56948,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) ) | $spl178),
  inference(cnf_transformation,[],[f56948_D])).
fof(f56948_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f56913,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) ) | $spl109),
  inference(resolution,[],[f56629,f1367])).
fof(f55972,plain,(
  $false | ($spl29 | $spl37)),
  inference(subsumption_resolution,[],[f5486,f48133])).
fof(f48133,plain,(
  s(bool,'F_0') = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl29),
  inference(resolution,[],[f2107,f507])).
fof(f55392,plain,(
  $false | ($spl29 | $spl33)),
  inference(subsumption_resolution,[],[f55391,f539])).
fof(f55391,plain,(
  p(s(bool,'F_0')) | ($spl29 | $spl33)),
  inference(forward_demodulation,[],[f55390,f48133])).
fof(f55390,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl33),
  inference(forward_demodulation,[],[f55371,f568])).
fof(f55371,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl33),
  inference(trivial_inequality_removal,[],[f55368])).
fof(f55368,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl33),
  inference(superposition,[],[f3014,f528])).
fof(f55389,plain,(
  $false | ($spl29 | $spl33)),
  inference(subsumption_resolution,[],[f55388,f539])).
fof(f55388,plain,(
  p(s(bool,'F_0')) | ($spl29 | $spl33)),
  inference(forward_demodulation,[],[f55387,f48133])).
fof(f55387,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl33),
  inference(forward_demodulation,[],[f55372,f568])).
fof(f55372,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl33),
  inference(trivial_inequality_removal,[],[f55367])).
fof(f55367,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl33),
  inference(superposition,[],[f3014,f526])).
fof(f55386,plain,(
  $false | ($spl29 | $spl33)),
  inference(subsumption_resolution,[],[f55385,f539])).
fof(f55385,plain,(
  p(s(bool,'F_0')) | ($spl29 | $spl33)),
  inference(forward_demodulation,[],[f55384,f48133])).
fof(f55384,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl33),
  inference(forward_demodulation,[],[f55374,f568])).
fof(f55374,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl33),
  inference(trivial_inequality_removal,[],[f55365])).
fof(f55365,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl33),
  inference(superposition,[],[f3014,f507])).
fof(f55375,plain,(
  $false | ($spl29 | $spl33)),
  inference(subsumption_resolution,[],[f55354,f48133])).
fof(f55354,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl33),
  inference(superposition,[],[f3014,f568])).
fof(f55300,plain,(
  $false | ($spl10 | $spl89)),
  inference(subsumption_resolution,[],[f18799,f47397])).
fof(f55299,plain,(
  $false | ($spl14 | $spl85)),
  inference(subsumption_resolution,[],[f18729,f47794])).
fof(f52394,plain,(
  $false | ($spl8 | $spl67)),
  inference(subsumption_resolution,[],[f15903,f4706])).
fof(f51577,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17)))))) ) | $spl174),
  inference(cnf_transformation,[],[f51577_D])).
fof(f51577_D,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17)))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f51508,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17)))))) ) | $spl109),
  inference(resolution,[],[f50401,f1700])).
fof(f50401,plain,(
  ( ! [X37,X35,X36] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X35)))),X36,X37)) ) | $spl109),
  inference(resolution,[],[f50190,f370])).
fof(f50190,plain,(
  ( ! [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))))))) ) | $spl109),
  inference(resolution,[],[f42185,f463])).
fof(f50686,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X14))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X15))) ) | $spl109),
  inference(resolution,[],[f50196,f1771])).
fof(f50059,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl169),
  inference(cnf_transformation,[],[f50059_D])).
fof(f50059_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl169),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl169])])).
fof(f50061,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl171),
  inference(cnf_transformation,[],[f50061_D])).
fof(f50061_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl171])])).
fof(f50057,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl109),
  inference(forward_demodulation,[],[f49882,f561])).
fof(f49882,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.num.|48|_0')))) | $spl109),
  inference(superposition,[],[f20221,f41984])).
fof(f41984,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl109),
  inference(superposition,[],[f552,f41769])).
fof(f552,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X21] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21)))))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.numeral.numeral_distrib|split|14')).
fof(f19387,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl93),
  inference(cnf_transformation,[],[f19387_D])).
fof(f19387_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f49344,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~p(s(bool,X9))) ) | $spl31),
  inference(resolution,[],[f49154,f1771])).
fof(f49154,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X9)) ) | $spl31),
  inference(resolution,[],[f2110,f361])).
fof(f2110,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl31),
  inference(cnf_transformation,[],[f2110_D])).
fof(f2110_D,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f19388,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl94),
  inference(cnf_transformation,[],[f19388_D])).
fof(f19388_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f49343,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl31),
  inference(resolution,[],[f49154,f1367])).
fof(f19456,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | $spl97),
  inference(cnf_transformation,[],[f19456_D])).
fof(f19456_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f49248,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X9))) ) | $spl27),
  inference(resolution,[],[f48892,f1771])).
fof(f48892,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X9)) ) | $spl27),
  inference(resolution,[],[f2105,f361])).
fof(f2105,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl27),
  inference(cnf_transformation,[],[f2105_D])).
fof(f2105_D,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f19457,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) ) | $spl98),
  inference(cnf_transformation,[],[f19457_D])).
fof(f19457_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f49247,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) ) | $spl27),
  inference(resolution,[],[f48892,f1367])).
fof(f49148,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl167),
  inference(cnf_transformation,[],[f49148_D])).
fof(f49148_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) <=> ~$spl167),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl167])])).
fof(f49137,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl30),
  inference(resolution,[],[f48180,f1700])).
fof(f48180,plain,(
  ( ! [X24,X23] : (sP0('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X23,X24)) ) | $spl30),
  inference(resolution,[],[f2111,f370])).
fof(f49003,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl163),
  inference(cnf_transformation,[],[f49003_D])).
fof(f49003_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f48986,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~p(s(bool,X9))) ) | $spl29),
  inference(resolution,[],[f48077,f1771])).
fof(f48077,plain,(
  ( ! [X10,X9] : (sP2(X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X10)) ) | $spl29),
  inference(resolution,[],[f2107,f361])).
fof(f49004,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl164),
  inference(cnf_transformation,[],[f49004_D])).
fof(f49004_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f48985,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl29),
  inference(resolution,[],[f48077,f1367])).
fof(f48885,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | $spl161),
  inference(cnf_transformation,[],[f48885_D])).
fof(f48885_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f48874,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) ) | $spl26),
  inference(resolution,[],[f47919,f1700])).
fof(f47919,plain,(
  ( ! [X24,X23] : (sP0('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X23,X24)) ) | $spl26),
  inference(resolution,[],[f2106,f370])).
fof(f48779,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl159),
  inference(cnf_transformation,[],[f48779_D])).
fof(f48779_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f48768,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl14),
  inference(resolution,[],[f47740,f1700])).
fof(f47740,plain,(
  ( ! [X24,X23] : (sP0('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X23,X24)) ) | $spl14),
  inference(resolution,[],[f837,f370])).
fof(f48674,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | $spl157),
  inference(cnf_transformation,[],[f48674_D])).
fof(f48674_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f48663,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl10),
  inference(resolution,[],[f47343,f1700])).
fof(f47343,plain,(
  ( ! [X24,X23] : (sP0('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X23,X24)) ) | $spl10),
  inference(resolution,[],[f833,f370])).
fof(f9135,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,X2)) ) | $spl44),
  inference(cnf_transformation,[],[f9135_D])).
fof(f9135_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f48429,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f48406])).
fof(f48406,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X4))) ) | $spl8),
  inference(resolution,[],[f46983,f372])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f277])).
fof(f48306,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f48277])).
fof(f48277,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X4))) ) | $spl8),
  inference(resolution,[],[f46982,f372])).
fof(f46982,plain,(
  ( ! [X6,X7] : (~sP3(X7,X6,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | $spl8),
  inference(superposition,[],[f2540,f373])).
fof(f48050,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f48015])).
fof(f48015,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(resolution,[],[f46981,f374])).
fof(f46981,plain,(
  ( ! [X4,X5] : (sP3(X5,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X4) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl8),
  inference(superposition,[],[f2540,f371])).
fof(f47883,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f47838])).
fof(f47838,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(resolution,[],[f46980,f374])).
fof(f46980,plain,(
  ( ! [X2,X3] : (sP3(X3,X2,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl8),
  inference(superposition,[],[f2540,f371])).
fof(f1803,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl24),
  inference(cnf_transformation,[],[f1803_D])).
fof(f1803_D,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f47298,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X13))) ) | $spl7),
  inference(subsumption_resolution,[],[f47286,f507])).
fof(f47286,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X13))) ) | $spl7),
  inference(superposition,[],[f679,f383])).
fof(f47297,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(subsumption_resolution,[],[f47285,f507])).
fof(f47285,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(superposition,[],[f679,f383])).
fof(f47105,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl8),
  inference(forward_demodulation,[],[f47015,f568])).
fof(f47015,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl8),
  inference(resolution,[],[f46994,f400])).
fof(f833,plain,(
  p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl10),
  inference(cnf_transformation,[],[f833_D])).
fof(f833_D,plain,(
  p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f47104,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl8),
  inference(forward_demodulation,[],[f47014,f568])).
fof(f47014,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl8),
  inference(resolution,[],[f46994,f400])).
fof(f675,plain,(
  ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl1),
  inference(cnf_transformation,[],[f675_D])).
fof(f675_D,plain,(
  ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f681,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | $spl9),
  inference(cnf_transformation,[],[f681_D])).
fof(f681_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f47013,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f532])).
fof(f47012,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f530])).
fof(f674,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl0),
  inference(cnf_transformation,[],[f674_D])).
fof(f674_D,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f679,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | $spl7),
  inference(cnf_transformation,[],[f679_D])).
fof(f679_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f47011,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f528])).
fof(f47010,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f526])).
fof(f47008,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f507])).
fof(f677,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,X12))) ) | $spl4),
  inference(cnf_transformation,[],[f677_D])).
fof(f677_D,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,X12))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f47007,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) )),
  inference(superposition,[],[f351,f383])).
fof(f47006,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f383])).
fof(f676,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X10))) ) | $spl2),
  inference(cnf_transformation,[],[f676_D])).
fof(f676_D,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X10))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f47005,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f351,f382])).
fof(f382,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f280])).
fof(f47004,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f382])).
fof(f46997,plain,(
  ( ! [X11] : (~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) ) | $spl8),
  inference(subsumption_resolution,[],[f46985,f530])).
fof(f46985,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) ) | $spl8),
  inference(superposition,[],[f2540,f382])).
fof(f46996,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl8),
  inference(subsumption_resolution,[],[f46984,f530])).
fof(f46984,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl8),
  inference(superposition,[],[f2540,f382])).
fof(f46977,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(subsumption_resolution,[],[f46971,f511])).
fof(f46971,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(resolution,[],[f8224,f484])).
fof(f8224,plain,(
  ( ! [X232] : (sP9(X232,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) ) | $spl8),
  inference(subsumption_resolution,[],[f5233,f544])).
fof(f5233,plain,(
  ( ! [X232] : (~p(s(bool,'T_0')) | sP9(X232,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) ) | $spl8),
  inference(superposition,[],[f485,f680])).
fof(f46976,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(subsumption_resolution,[],[f46970,f511])).
fof(f46970,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(resolution,[],[f8224,f488])).
fof(f46969,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X1))) ) | $spl8),
  inference(subsumption_resolution,[],[f46963,f511])).
fof(f46963,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X1))) ) | $spl8),
  inference(resolution,[],[f8222,f484])).
fof(f8222,plain,(
  ( ! [X237] : (sP9('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X237)) ) | $spl8),
  inference(subsumption_resolution,[],[f5238,f544])).
fof(f5238,plain,(
  ( ! [X237] : (~p(s(bool,'T_0')) | sP9('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X237)) ) | $spl8),
  inference(superposition,[],[f491,f680])).
fof(f491,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f321])).
fof(f46968,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f46962,f511])).
fof(f46962,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X0))) ) | $spl8),
  inference(resolution,[],[f8222,f488])).
fof(f15898,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,X30)) ) | $spl58),
  inference(cnf_transformation,[],[f15898_D])).
fof(f15898_D,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,X30)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f15646,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,X30) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl44),
  inference(duplicate_literal_removal,[],[f15644])).
fof(f15644,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,X30) | s(bool,'T_0') != s(bool,X30) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl44),
  inference(superposition,[],[f9135,f532])).
fof(f15647,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,X29) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl44),
  inference(duplicate_literal_removal,[],[f15643])).
fof(f15643,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') != s(bool,X29) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl44),
  inference(superposition,[],[f9135,f530])).
fof(f1780,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | $spl23),
  inference(cnf_transformation,[],[f1780_D])).
fof(f1780_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f15907,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X25))) ) | $spl44),
  inference(subsumption_resolution,[],[f15611,f530])).
fof(f15611,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,X25) | ~p(s(bool,X25))) ) | $spl44),
  inference(superposition,[],[f9135,f532])).
fof(f15908,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X24))) ) | $spl44),
  inference(subsumption_resolution,[],[f15610,f530])).
fof(f15610,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,X24) | ~p(s(bool,X24))) ) | $spl44),
  inference(superposition,[],[f9135,f530])).
fof(f46958,plain,(
  ( ! [X19] : (s(bool,X19) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | p(s(bool,X19))) ) | $spl154),
  inference(cnf_transformation,[],[f46958_D])).
fof(f46958_D,plain,(
  ( ! [X19] : (s(bool,X19) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | p(s(bool,X19))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f32066,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl46),
  inference(subsumption_resolution,[],[f32015,f507])).
fof(f32015,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'F_0') != s(bool,X16) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl46),
  inference(superposition,[],[f9873,f383])).
fof(f32068,plain,(
  ( ! [X19,X18] : (s(bool,X19) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | p(s(bool,X18)) | p(s(bool,X19))) ) | $spl46),
  inference(subsumption_resolution,[],[f32016,f507])).
fof(f32016,plain,(
  ( ! [X19,X18] : (s(bool,X19) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'F_0') != s(bool,X18) | p(s(bool,X18)) | p(s(bool,X19))) ) | $spl46),
  inference(superposition,[],[f9873,f383])).
fof(f15903,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | $spl67),
  inference(cnf_transformation,[],[f15903_D])).
fof(f15903_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f32070,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | p(s(bool,X20))) ) | $spl46),
  inference(subsumption_resolution,[],[f32017,f507])).
fof(f32017,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'F_0') != s(bool,X20) | p(s(bool,X20))) ) | $spl46),
  inference(superposition,[],[f9873,f507])).
fof(f32072,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | p(s(bool,X22))) ) | $spl46),
  inference(subsumption_resolution,[],[f32019,f507])).
fof(f32019,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'F_0') != s(bool,X22) | p(s(bool,X22))) ) | $spl46),
  inference(superposition,[],[f9873,f526])).
fof(f32074,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | p(s(bool,X23))) ) | $spl46),
  inference(subsumption_resolution,[],[f32020,f507])).
fof(f32020,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'F_0') != s(bool,X23) | p(s(bool,X23))) ) | $spl46),
  inference(superposition,[],[f9873,f528])).
fof(f46957,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,X24) | s(bool,'F_0') != s(bool,X24)) ) | $spl152),
  inference(cnf_transformation,[],[f46957_D])).
fof(f46957_D,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,X24) | s(bool,'F_0') != s(bool,X24)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f32097,plain,(
  ( ! [X21,X22] : (s(bool,X21) != s(bool,X22) | s(bool,'F_0') != s(bool,X22) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | ($spl18 | $spl46)),
  inference(subsumption_resolution,[],[f32048,f1298])).
fof(f1298,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X0))) ) | $spl18),
  inference(superposition,[],[f1120,f532])).
fof(f32048,plain,(
  ( ! [X21,X22] : (s(bool,X21) != s(bool,X22) | s(bool,'F_0') != s(bool,X22) | p(s(bool,X21)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl46),
  inference(superposition,[],[f9873,f383])).
fof(f32099,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,X24) | s(bool,'F_0') != s(bool,X24) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | ($spl18 | $spl46)),
  inference(subsumption_resolution,[],[f32049,f1298])).
fof(f32049,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,X24) | s(bool,'F_0') != s(bool,X24) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X23))) ) | $spl46),
  inference(superposition,[],[f9873,f383])).
fof(f32060,plain,(
  ( ! [X25] : (s(bool,'F_0') != s(bool,X25) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl46),
  inference(duplicate_literal_removal,[],[f32050])).
fof(f32050,plain,(
  ( ! [X25] : (s(bool,'F_0') != s(bool,X25) | s(bool,'F_0') != s(bool,X25) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl46),
  inference(superposition,[],[f9873,f507])).
fof(f32058,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,X27) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl46),
  inference(duplicate_literal_removal,[],[f32052])).
fof(f32052,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,X27) | s(bool,'F_0') != s(bool,X27) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl46),
  inference(superposition,[],[f9873,f526])).
fof(f32057,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,X28) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl46),
  inference(duplicate_literal_removal,[],[f32053])).
fof(f32053,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,X28) | s(bool,'F_0') != s(bool,X28) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl46),
  inference(superposition,[],[f9873,f528])).
fof(f46924,plain,(
  $false | ($spl6 | $spl26)),
  inference(subsumption_resolution,[],[f46923,f539])).
fof(f46923,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl26)),
  inference(forward_demodulation,[],[f2106,f19008])).
fof(f19008,plain,(
  s(bool,'F_0') = s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | $spl6),
  inference(resolution,[],[f18446,f507])).
fof(f18446,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl6),
  inference(trivial_inequality_removal,[],[f18426])).
fof(f18426,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl6),
  inference(superposition,[],[f16081,f7096])).
fof(f7096,plain,(
  s(bool,'F_0') = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))),
  inference(resolution,[],[f6814,f507])).
fof(f6814,plain,(
  ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f401,f568])).
fof(f16081,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X1))))) ) | $spl6),
  inference(superposition,[],[f1907,f398])).
fof(f1907,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | $spl6),
  inference(backward_demodulation,[],[f678,f351])).
fof(f678,plain,(
  s(bool,'F_0') = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | $spl6),
  inference(cnf_transformation,[],[f678_D])).
fof(f678_D,plain,(
  s(bool,'F_0') = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f46922,plain,(
  $false | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f46921,f539])).
fof(f46921,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl30)),
  inference(forward_demodulation,[],[f2111,f18913])).
fof(f18913,plain,(
  s(bool,'F_0') = s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl6),
  inference(resolution,[],[f18418,f507])).
fof(f18418,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl6),
  inference(trivial_inequality_removal,[],[f18398])).
fof(f18398,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl6),
  inference(superposition,[],[f16080,f7096])).
fof(f16080,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl6),
  inference(superposition,[],[f1907,f398])).
fof(f46904,plain,(
  $false | ($spl0 | $spl28)),
  inference(subsumption_resolution,[],[f46903,f545])).
fof(f46903,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl28)),
  inference(forward_demodulation,[],[f46902,f4994])).
fof(f4994,plain,(
  s(bool,'T_0') = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl28),
  inference(resolution,[],[f2108,f530])).
fof(f46902,plain,(
  ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl0),
  inference(subsumption_resolution,[],[f46875,f401])).
fof(f46875,plain,(
  ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl0),
  inference(superposition,[],[f16430,f568])).
fof(f16430,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X2),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X2),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl0),
  inference(resolution,[],[f674,f396])).
fof(f46901,plain,(
  $false | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f46900,f401])).
fof(f46900,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f46874,f544])).
fof(f46874,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ($spl0 | $spl12)),
  inference(superposition,[],[f16430,f960])).
fof(f960,plain,(
  s(bool,'T_0') = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl12),
  inference(resolution,[],[f835,f530])).
fof(f46899,plain,(
  $false | ($spl0 | $spl28)),
  inference(subsumption_resolution,[],[f46898,f539])).
fof(f46898,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl28)),
  inference(forward_demodulation,[],[f46897,f7096])).
fof(f46897,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ($spl0 | $spl28)),
  inference(subsumption_resolution,[],[f46873,f544])).
fof(f46873,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ($spl0 | $spl28)),
  inference(superposition,[],[f16430,f4994])).
fof(f41721,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')))) | $spl147),
  inference(cnf_transformation,[],[f41721_D])).
fof(f41721_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')))) <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f20205,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl116),
  inference(cnf_transformation,[],[f20205_D])).
fof(f20205_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f42858,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')))) | $spl109),
  inference(superposition,[],[f41567,f389])).
fof(f389,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X21] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X21))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.numeral.numeral_distrib|split|17')).
fof(f41567,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')))) | $spl109),
  inference(superposition,[],[f24592,f394])).
fof(f394,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X20] : s('type.num.num',X20) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X20)))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.arithmetic.ADD_CLAUSES|split|0')).
fof(f42857,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')))) | $spl109),
  inference(superposition,[],[f41567,f387])).
fof(f42499,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) ) | $spl150),
  inference(cnf_transformation,[],[f42499_D])).
fof(f42499_D,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f42436,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X15))) ) | $spl109),
  inference(resolution,[],[f42300,f1771])).
fof(f42300,plain,(
  ( ! [X14,X15,X13] : (sP2(X13,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')),X15)) ) | $spl109),
  inference(resolution,[],[f42160,f361])).
fof(f41723,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl148),
  inference(cnf_transformation,[],[f41723_D])).
fof(f41723_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f41554,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')))) | $spl109),
  inference(superposition,[],[f24592,f389])).
fof(f41553,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13('const.num.|48|_0','const.arithmetic.ZERO_0')))) | $spl109),
  inference(superposition,[],[f24592,f387])).
fof(f28780,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl139),
  inference(cnf_transformation,[],[f28780_D])).
fof(f28780_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f39528,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,X16))) ) | ($spl0 | $spl6 | $spl109)),
  inference(resolution,[],[f18352,f27877])).
fof(f27877,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl109),
  inference(resolution,[],[f24341,f360])).
fof(f18352,plain,(
  ( ! [X6,X7] : (~sP2(X6,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X7) | s(bool,'F_0') != s(bool,X7) | ~p(s(bool,X6))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f18351,f368])).
fof(f18351,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,'F_0') != s(bool,X7) | ~sP2(X6,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X7) | ~sP0('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X7,X6)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f18350,f364])).
fof(f18350,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,'F_0') != s(bool,X7) | ~sP2(X6,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X7) | ~sP1(X7,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X6) | ~sP0('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X7,X6)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f18329,f674])).
fof(f18329,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,'F_0') != s(bool,X7) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~sP2(X6,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X7) | ~sP1(X7,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X6) | ~sP0('const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X7,X6)) ) | $spl6),
  inference(resolution,[],[f16084,f356])).
fof(f16084,plain,(
  ( ! [X6,X7] : (~sP3(X7,X6,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X6)) ) | $spl6),
  inference(superposition,[],[f1907,f373])).
fof(f20878,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl125),
  inference(cnf_transformation,[],[f20878_D])).
fof(f20878_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f39525,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,X13))) ) | ($spl0 | $spl6 | $spl54)),
  inference(resolution,[],[f18352,f20067])).
fof(f20067,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl54),
  inference(resolution,[],[f19853,f360])).
fof(f19853,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl54),
  inference(resolution,[],[f16958,f463])).
fof(f30473,plain,(
  ( ! [X43,X42] : (~p(s(bool,X43)) | ~p(s(bool,X42))) ) | $spl62),
  inference(subsumption_resolution,[],[f30457,f530])).
fof(f30457,plain,(
  ( ! [X43,X42] : (~p(s(bool,X43)) | s(bool,'T_0') != s(bool,X43) | ~p(s(bool,X42))) ) | $spl62),
  inference(trivial_inequality_removal,[],[f30417])).
fof(f30417,plain,(
  ( ! [X43,X42] : (s(bool,X42) != s(bool,X42) | ~p(s(bool,X43)) | s(bool,'T_0') != s(bool,X43) | ~p(s(bool,X42))) ) | $spl62),
  inference(superposition,[],[f15900,f530])).
fof(f30472,plain,(
  ( ! [X41,X40] : (~p(s(bool,X41)) | ~p(s(bool,X40))) ) | $spl62),
  inference(subsumption_resolution,[],[f30458,f530])).
fof(f30458,plain,(
  ( ! [X41,X40] : (~p(s(bool,X41)) | s(bool,'T_0') != s(bool,X41) | ~p(s(bool,X40))) ) | $spl62),
  inference(trivial_inequality_removal,[],[f30416])).
fof(f30416,plain,(
  ( ! [X41,X40] : (s(bool,X40) != s(bool,X40) | ~p(s(bool,X41)) | s(bool,'T_0') != s(bool,X41) | ~p(s(bool,X40))) ) | $spl62),
  inference(superposition,[],[f15900,f532])).
fof(f30467,plain,(
  ( ! [X18,X16] : (s(bool,X16) != s(bool,X18) | ~p(s(bool,X16))) ) | $spl144),
  inference(cnf_transformation,[],[f30467_D])).
fof(f30467_D,plain,(
  ( ! [X18,X16] : (s(bool,X16) != s(bool,X18) | ~p(s(bool,X16))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f30469,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X21) | ~p(s(bool,X19)) | ~p(s(bool,X20))) ) | $spl62),
  inference(subsumption_resolution,[],[f30468,f380])).
fof(f380,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f280])).
fof(f30468,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X21)) | s(bool,X19) != s(bool,X21) | ~p(s(bool,X19)) | ~p(s(bool,X20))) ) | $spl62),
  inference(subsumption_resolution,[],[f30407,f530])).
fof(f30407,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,X20) | ~p(s(bool,X21)) | s(bool,X19) != s(bool,X21) | ~p(s(bool,X19)) | ~p(s(bool,X20))) ) | $spl62),
  inference(superposition,[],[f15900,f382])).
fof(f30466,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,X18) | ~p(s(bool,X17)) | ~p(s(bool,X16))) ) | $spl62),
  inference(subsumption_resolution,[],[f30465,f380])).
fof(f30465,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X18)) | s(bool,X16) != s(bool,X18) | ~p(s(bool,X17)) | ~p(s(bool,X16))) ) | $spl62),
  inference(subsumption_resolution,[],[f30406,f530])).
fof(f30406,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,X18)) | s(bool,X16) != s(bool,X18) | ~p(s(bool,X17)) | ~p(s(bool,X16))) ) | $spl62),
  inference(superposition,[],[f15900,f382])).
fof(f28989,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X9))) ) | $spl109),
  inference(resolution,[],[f28384,f1771])).
fof(f29037,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl142),
  inference(cnf_transformation,[],[f29037_D])).
fof(f29037_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f28988,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl109),
  inference(resolution,[],[f28384,f1367])).
fof(f28738,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl109),
  inference(resolution,[],[f27883,f1700])).
fof(f27883,plain,(
  ( ! [X24,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X24,X25)) ) | $spl109),
  inference(resolution,[],[f24341,f370])).
fof(f28000,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | $spl137),
  inference(cnf_transformation,[],[f28000_D])).
fof(f28000_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f19973,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl102),
  inference(cnf_transformation,[],[f19973_D])).
fof(f19973_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f27951,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | $spl109),
  inference(superposition,[],[f24341,f389])).
fof(f27950,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | $spl109),
  inference(superposition,[],[f24341,f387])).
fof(f25053,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl135),
  inference(cnf_transformation,[],[f25053_D])).
fof(f25053_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f25027,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl109),
  inference(resolution,[],[f24498,f1700])).
fof(f24498,plain,(
  ( ! [X23,X22] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X22,X23)) ) | $spl109),
  inference(resolution,[],[f24342,f370])).
fof(f24736,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0'))) | $spl131),
  inference(cnf_transformation,[],[f24736_D])).
fof(f24736_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f24706,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | $spl109),
  inference(resolution,[],[f24348,f1771])).
fof(f24348,plain,(
  ( ! [X10,X9] : (sP2(X9,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')),X10)) ) | $spl109),
  inference(resolution,[],[f24310,f361])).
fof(f24737,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl132),
  inference(cnf_transformation,[],[f24737_D])).
fof(f24737_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f24705,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl109),
  inference(resolution,[],[f24348,f1367])).
fof(f24328,plain,(
  ( ! [X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))) ) | $spl126),
  inference(cnf_transformation,[],[f24328_D])).
fof(f24328_D,plain,(
  ( ! [X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f24330,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) | $spl129),
  inference(cnf_transformation,[],[f24330_D])).
fof(f24330_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f24311,plain,(
  ( ! [X41] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41)))) ) | $spl109),
  inference(trivial_inequality_removal,[],[f24300])).
fof(f24300,plain,(
  ( ! [X41] : (s('type.num.num',X41) != s('type.num.num',X41) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41)))) ) | $spl109),
  inference(superposition,[],[f20221,f387])).
fof(f24312,plain,(
  ( ! [X40] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))) ) | $spl109),
  inference(trivial_inequality_removal,[],[f24299])).
fof(f24299,plain,(
  ( ! [X40] : (s('type.num.num',X40) != s('type.num.num',X40) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))) ) | $spl109),
  inference(superposition,[],[f20221,f389])).
fof(f22617,plain,(
  $false | ($spl0 | $spl23)),
  inference(subsumption_resolution,[],[f1780,f16496])).
fof(f16496,plain,(
  s(bool,'T_0') = s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | $spl0),
  inference(resolution,[],[f674,f530])).
fof(f20845,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl54),
  inference(resolution,[],[f20073,f1700])).
fof(f20073,plain,(
  ( ! [X23,X22] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X22,X23)) ) | $spl54),
  inference(resolution,[],[f19853,f370])).
fof(f20328,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,X9))) ) | $spl54),
  inference(resolution,[],[f19859,f1771])).
fof(f20366,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl122),
  inference(cnf_transformation,[],[f20366_D])).
fof(f20366_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f20327,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl54),
  inference(resolution,[],[f19859,f1367])).
fof(f20210,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl119),
  inference(cnf_transformation,[],[f20210_D])).
fof(f20210_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f20036,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.|48|_0') | $spl109),
  inference(cnf_transformation,[],[f20036_D])).
fof(f20036_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.|48|_0') <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f20198,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl113),
  inference(superposition,[],[f20182,f389])).
fof(f20197,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl113),
  inference(superposition,[],[f20182,f387])).
fof(f20204,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl115),
  inference(cnf_transformation,[],[f20204_D])).
fof(f20204_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f20206,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl117),
  inference(cnf_transformation,[],[f20206_D])).
fof(f20206_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f20190,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl113),
  inference(superposition,[],[f20182,f389])).
fof(f20189,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl113),
  inference(superposition,[],[f20182,f387])).
fof(f20180,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl110),
  inference(cnf_transformation,[],[f20180_D])).
fof(f20180_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f20182,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl113),
  inference(cnf_transformation,[],[f20182_D])).
fof(f20182_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f20183,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl54),
  inference(forward_demodulation,[],[f20133,f555])).
fof(f20133,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl54),
  inference(superposition,[],[f19853,f389])).
fof(f20178,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl54),
  inference(forward_demodulation,[],[f20132,f555])).
fof(f20132,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl54),
  inference(superposition,[],[f19853,f387])).
fof(f20034,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl107),
  inference(cnf_transformation,[],[f20034_D])).
fof(f20034_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f20026,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl101),
  inference(superposition,[],[f19972,f389])).
fof(f20025,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl101),
  inference(superposition,[],[f19972,f387])).
fof(f19972,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl101),
  inference(cnf_transformation,[],[f19972_D])).
fof(f19972_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f20000,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl105),
  inference(cnf_transformation,[],[f20000_D])).
fof(f20000_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f19994,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl54),
  inference(superposition,[],[f19850,f389])).
fof(f19850,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl54),
  inference(resolution,[],[f16958,f499])).
fof(f499,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f19993,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl54),
  inference(superposition,[],[f19850,f387])).
fof(f19974,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl103),
  inference(cnf_transformation,[],[f19974_D])).
fof(f19974_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f19926,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl54),
  inference(superposition,[],[f16958,f389])).
fof(f19925,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl54),
  inference(superposition,[],[f16958,f387])).
fof(f19445,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f18952,f1771])).
fof(f18952,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X9)) ) | $spl6),
  inference(resolution,[],[f18446,f361])).
fof(f19444,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) ) | $spl6),
  inference(resolution,[],[f18952,f1367])).
fof(f19376,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f18857,f1771])).
fof(f18857,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X9)) ) | $spl6),
  inference(resolution,[],[f18418,f361])).
fof(f19375,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl6),
  inference(resolution,[],[f18857,f1367])).
fof(f18799,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | $spl89),
  inference(cnf_transformation,[],[f18799_D])).
fof(f18799_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f18788,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f18622,f1771])).
fof(f18622,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X9)) ) | $spl6),
  inference(resolution,[],[f18447,f361])).
fof(f18447,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl6),
  inference(subsumption_resolution,[],[f18427,f7096])).
fof(f18427,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl6),
  inference(superposition,[],[f16081,f568])).
fof(f18800,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl90),
  inference(cnf_transformation,[],[f18800_D])).
fof(f18800_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f18787,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl6),
  inference(resolution,[],[f18622,f1367])).
fof(f18718,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | ~p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f18530,f1771])).
fof(f18530,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)),X9)) ) | $spl6),
  inference(resolution,[],[f18419,f361])).
fof(f18419,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl6),
  inference(subsumption_resolution,[],[f18399,f7096])).
fof(f18399,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl6),
  inference(superposition,[],[f16080,f568])).
fof(f18717,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | s(bool,'T_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl6),
  inference(resolution,[],[f18530,f1367])).
fof(f18512,plain,(
  $false | ($spl6 | $spl80)),
  inference(trivial_inequality_removal,[],[f18492])).
fof(f18492,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl6 | $spl80)),
  inference(superposition,[],[f18422,f678])).
fof(f18511,plain,(
  $false | $spl80),
  inference(trivial_inequality_removal,[],[f18494])).
fof(f18494,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl80),
  inference(superposition,[],[f18422,f7096])).
fof(f18510,plain,(
  $false | ($spl56 | $spl80)),
  inference(trivial_inequality_removal,[],[f18497])).
fof(f18497,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl56 | $spl80)),
  inference(superposition,[],[f18422,f16314])).
fof(f16314,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl56),
  inference(resolution,[],[f16105,f507])).
fof(f18509,plain,(
  $false | $spl80),
  inference(trivial_inequality_removal,[],[f18502])).
fof(f18502,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl80),
  inference(superposition,[],[f18422,f563])).
fof(f563,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X21] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X21))),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.numeral.numeral_lte|split|1')).
fof(f18508,plain,(
  $false | $spl80),
  inference(trivial_inequality_removal,[],[f18506])).
fof(f18506,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl80),
  inference(superposition,[],[f18422,f562])).
fof(f18507,plain,(
  $false | $spl80),
  inference(equality_resolution,[],[f18422])).
fof(f18449,plain,(
  ( ! [X20] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X20)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl82),
  inference(cnf_transformation,[],[f18449_D])).
fof(f18449_D,plain,(
  ( ! [X20] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X20)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f18450,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X22)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X22),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl6),
  inference(subsumption_resolution,[],[f18435,f525])).
fof(f18435,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X22)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X22),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X23))) ) | $spl6),
  inference(superposition,[],[f16081,f383])).
fof(f18448,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X21) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X20)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl6),
  inference(subsumption_resolution,[],[f18434,f525])).
fof(f18434,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X21) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X20)))) | p(s(bool,X21)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl6),
  inference(superposition,[],[f16081,f383])).
fof(f18421,plain,(
  ( ! [X20] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl78),
  inference(cnf_transformation,[],[f18421_D])).
fof(f18421_D,plain,(
  ( ! [X20] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f18423,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X22),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X22),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl6),
  inference(subsumption_resolution,[],[f18407,f525])).
fof(f18407,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X22),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X22),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X23))) ) | $spl6),
  inference(superposition,[],[f16080,f383])).
fof(f18420,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X21) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl6),
  inference(subsumption_resolution,[],[f18406,f525])).
fof(f18406,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X21) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,X21)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X20),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl6),
  inference(superposition,[],[f16080,f383])).
fof(f9873,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'F_0') != s(bool,X2)) ) | $spl46),
  inference(cnf_transformation,[],[f9873_D])).
fof(f9873_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f18382,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f18369])).
fof(f18369,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X4))) ) | $spl6),
  inference(resolution,[],[f16085,f372])).
fof(f16085,plain,(
  ( ! [X8,X9] : (~sP3(X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X8) | ~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X8)) ) | $spl6),
  inference(superposition,[],[f1907,f373])).
fof(f18348,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f18328])).
fof(f18328,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X4))) ) | $spl6),
  inference(resolution,[],[f16084,f372])).
fof(f18312,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f18296])).
fof(f18296,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl6),
  inference(resolution,[],[f16083,f374])).
fof(f16083,plain,(
  ( ! [X4,X5] : (sP3(X5,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X4) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl6),
  inference(superposition,[],[f1907,f371])).
fof(f18269,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f18239])).
fof(f18239,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl6),
  inference(resolution,[],[f16082,f374])).
fof(f16082,plain,(
  ( ! [X2,X3] : (sP3(X3,X2,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'F_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f1907,f371])).
fof(f16623,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl75),
  inference(cnf_transformation,[],[f16623_D])).
fof(f16623_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f16602,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,X9))) ) | $spl56),
  inference(resolution,[],[f16258,f1771])).
fof(f16258,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X9)) ) | $spl56),
  inference(resolution,[],[f16105,f361])).
fof(f16624,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl76),
  inference(cnf_transformation,[],[f16624_D])).
fof(f16624_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f16601,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl56),
  inference(resolution,[],[f16258,f1367])).
fof(f16543,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl73),
  inference(cnf_transformation,[],[f16543_D])).
fof(f16543_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f16526,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl56),
  inference(resolution,[],[f16119,f1700])).
fof(f16119,plain,(
  ( ! [X24,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X24,X25)) ) | $spl56),
  inference(resolution,[],[f15780,f370])).
fof(f16252,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f532])).
fof(f16251,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f530])).
fof(f16250,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f528])).
fof(f16249,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f526])).
fof(f16247,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f507])).
fof(f16246,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) )),
  inference(superposition,[],[f351,f383])).
fof(f16245,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f383])).
fof(f16244,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f351,f382])).
fof(f16243,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f382])).
fof(f16101,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) ) | $spl6),
  inference(subsumption_resolution,[],[f16089,f507])).
fof(f16089,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) ) | $spl6),
  inference(superposition,[],[f1907,f383])).
fof(f16100,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl6),
  inference(subsumption_resolution,[],[f16088,f507])).
fof(f16088,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl6),
  inference(superposition,[],[f1907,f383])).
fof(f16071,plain,(
  $false | $spl57),
  inference(trivial_inequality_removal,[],[f15970])).
fof(f15970,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl57),
  inference(resolution,[],[f15781,f499])).
fof(f16070,plain,(
  $false | $spl57),
  inference(trivial_inequality_removal,[],[f15971])).
fof(f15971,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl57),
  inference(resolution,[],[f15781,f500])).
fof(f500,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f15963,plain,(
  $false | $spl58),
  inference(trivial_inequality_removal,[],[f15942])).
fof(f15942,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl58),
  inference(superposition,[],[f15898,f634])).
fof(f634,plain,(
  s(bool,'T_0') = s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))),
  inference(resolution,[],[f350,f530])).
fof(f15962,plain,(
  $false | $spl58),
  inference(trivial_inequality_removal,[],[f15943])).
fof(f15943,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl58),
  inference(superposition,[],[f15898,f7282])).
fof(f7282,plain,(
  s(bool,'T_0') = s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))),
  inference(resolution,[],[f7127,f530])).
fof(f7127,plain,(
  p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))),
  inference(subsumption_resolution,[],[f7126,f401])).
fof(f7126,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))),
  inference(forward_demodulation,[],[f7034,f568])).
fof(f7034,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))),
  inference(resolution,[],[f6814,f400])).
fof(f15961,plain,(
  $false | ($spl28 | $spl58)),
  inference(trivial_inequality_removal,[],[f15947])).
fof(f15947,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl28 | $spl58)),
  inference(superposition,[],[f15898,f4994])).
fof(f15960,plain,(
  $false | ($spl12 | $spl58)),
  inference(trivial_inequality_removal,[],[f15949])).
fof(f15949,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl58)),
  inference(superposition,[],[f15898,f960])).
fof(f15959,plain,(
  $false | $spl58),
  inference(equality_resolution,[],[f15898])).
fof(f15909,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X15))) ) | $spl70),
  inference(cnf_transformation,[],[f15909_D])).
fof(f15909_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X15))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f15653,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X13)) | ~p(s(bool,X12))) ) | $spl44),
  inference(subsumption_resolution,[],[f15602,f530])).
fof(f15602,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,X12) | ~p(s(bool,X13)) | ~p(s(bool,X12))) ) | $spl44),
  inference(superposition,[],[f9135,f382])).
fof(f15655,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X14)) | ~p(s(bool,X15))) ) | $spl44),
  inference(subsumption_resolution,[],[f15603,f530])).
fof(f15603,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,X14) | ~p(s(bool,X14)) | ~p(s(bool,X15))) ) | $spl44),
  inference(superposition,[],[f9135,f382])).
fof(f15904,plain,(
  ( ! [X50] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl68),
  inference(cnf_transformation,[],[f15904_D])).
fof(f15904_D,plain,(
  ( ! [X50] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.ZERO_0')))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f15623,plain,(
  ( ! [X50] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl44),
  inference(superposition,[],[f9135,f563])).
fof(f15901,plain,(
  ( ! [X56] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X56))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f15901_D])).
fof(f15901_D,plain,(
  ( ! [X56] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X56))),s('type.num.num','const.arithmetic.ZERO_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f15627,plain,(
  ( ! [X56] : (s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X56))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl44),
  inference(superposition,[],[f9135,f562])).
fof(f15900,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') != s(bool,X20) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,X20)) ) | $spl62),
  inference(cnf_transformation,[],[f15900_D])).
fof(f15900_D,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') != s(bool,X20) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,X20)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f15635,plain,(
  ( ! [X17,X18] : (s(bool,X17) != s(bool,X18) | s(bool,'T_0') != s(bool,X18) | ~p(s(bool,X17)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl44),
  inference(superposition,[],[f9135,f382])).
fof(f15636,plain,(
  ( ! [X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') != s(bool,X20) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X19))) ) | $spl44),
  inference(superposition,[],[f9135,f382])).
fof(f15899,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,X24) | p(s(bool,X23)) | s(bool,'T_0') != s(bool,X24)) ) | $spl60),
  inference(cnf_transformation,[],[f15899_D])).
fof(f15899_D,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,X24) | p(s(bool,X23)) | s(bool,'T_0') != s(bool,X24)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f15637,plain,(
  ( ! [X21,X22] : (s(bool,X21) != s(bool,X22) | s(bool,'T_0') != s(bool,X22) | p(s(bool,X21)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl44),
  inference(superposition,[],[f9135,f383])).
fof(f15638,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,X24) | s(bool,'T_0') != s(bool,X24) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X23))) ) | $spl44),
  inference(superposition,[],[f9135,f383])).
fof(f15884,plain,(
  $false | $spl55),
  inference(trivial_inequality_removal,[],[f15783])).
fof(f15783,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl55),
  inference(resolution,[],[f15777,f499])).
fof(f15777,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl55),
  inference(trivial_inequality_removal,[],[f15776])).
fof(f15776,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl55),
  inference(superposition,[],[f15746,f532])).
fof(f15883,plain,(
  $false | $spl55),
  inference(trivial_inequality_removal,[],[f15784])).
fof(f15784,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl55),
  inference(resolution,[],[f15777,f500])).
fof(f15781,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl57),
  inference(cnf_transformation,[],[f15781_D])).
fof(f15781_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f15782,plain,(
  ( ! [X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | ~p(s(bool,X9))) ) | $spl55),
  inference(subsumption_resolution,[],[f15768,f530])).
fof(f15768,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | ~p(s(bool,X9))) ) | $spl55),
  inference(superposition,[],[f15746,f382])).
fof(f15779,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))))) ) | $spl55),
  inference(subsumption_resolution,[],[f15767,f530])).
fof(f15767,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))))) ) | $spl55),
  inference(superposition,[],[f15746,f382])).
fof(f15744,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4)))))) ) | $spl52),
  inference(cnf_transformation,[],[f15744_D])).
fof(f15744_D,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4)))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f15746,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl55),
  inference(cnf_transformation,[],[f15746_D])).
fof(f15746_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f15712,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5)))))) ) | ($spl8 | $spl44)),
  inference(superposition,[],[f15664,f389])).
fof(f15664,plain,(
  ( ! [X50] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | ($spl8 | $spl44)),
  inference(subsumption_resolution,[],[f15623,f4706])).
fof(f15711,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4)))))) ) | ($spl8 | $spl44)),
  inference(superposition,[],[f15664,f387])).
fof(f10801,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | $spl51),
  inference(cnf_transformation,[],[f10801_D])).
fof(f10801_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f10795,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) )),
  inference(resolution,[],[f7228,f1700])).
fof(f7228,plain,(
  ( ! [X24,X23] : (sP0('const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK11)),X23,X24)) )),
  inference(resolution,[],[f7127,f370])).
fof(f10706,plain,(
  ( ! [X16] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X16)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl48),
  inference(cnf_transformation,[],[f10706_D])).
fof(f10706_D,plain,(
  ( ! [X16] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X16)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f10707,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X18)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X18),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X19))) ) | $spl8),
  inference(subsumption_resolution,[],[f10694,f530])).
fof(f10694,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,X19) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X18)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X18),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X19))) ) | $spl8),
  inference(superposition,[],[f9895,f382])).
fof(f9895,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X1))))) ) | $spl8),
  inference(superposition,[],[f2540,f398])).
fof(f10705,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X16)))) | ~p(s(bool,X17)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl8),
  inference(subsumption_resolution,[],[f10693,f530])).
fof(f10693,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),X16)))) | ~p(s(bool,X17)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl8),
  inference(superposition,[],[f9895,f382])).
fof(f832,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl11),
  inference(cnf_transformation,[],[f832_D])).
fof(f832_D,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f7146,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | $spl39),
  inference(cnf_transformation,[],[f7146_D])).
fof(f7146_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f10688,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl8),
  inference(superposition,[],[f9895,f568])).
fof(f7332,plain,(
  ( ! [X16] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl42),
  inference(cnf_transformation,[],[f7332_D])).
fof(f7332_D,plain,(
  ( ! [X16] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f10685,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X18),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X18),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X19))) ) | $spl8),
  inference(subsumption_resolution,[],[f10673,f530])).
fof(f10673,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,X19) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X18),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X18),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X19))) ) | $spl8),
  inference(superposition,[],[f9894,f382])).
fof(f9894,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl8),
  inference(superposition,[],[f2540,f398])).
fof(f10684,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X17)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl8),
  inference(subsumption_resolution,[],[f10672,f530])).
fof(f10672,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X17)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl8),
  inference(superposition,[],[f9894,f382])).
fof(f10667,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl8),
  inference(superposition,[],[f9894,f568])).
fof(f10655,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f10646])).
fof(f10646,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X4))) ) | $spl8),
  inference(resolution,[],[f9899,f372])).
fof(f9899,plain,(
  ( ! [X8,X9] : (~sP3(X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X8) | ~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X8)) ) | $spl8),
  inference(superposition,[],[f2540,f373])).
fof(f10637,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f10625])).
fof(f10625,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X4))) ) | $spl8),
  inference(resolution,[],[f9898,f372])).
fof(f9898,plain,(
  ( ! [X6,X7] : (~sP3(X7,X6,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | $spl8),
  inference(superposition,[],[f2540,f373])).
fof(f10619,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f10610])).
fof(f10610,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(resolution,[],[f9897,f374])).
fof(f9897,plain,(
  ( ! [X4,X5] : (sP3(X5,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X4) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl8),
  inference(superposition,[],[f2540,f371])).
fof(f10602,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f10590])).
fof(f10590,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(resolution,[],[f9896,f374])).
fof(f9896,plain,(
  ( ! [X2,X3] : (sP3(X3,X2,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl8),
  inference(superposition,[],[f2540,f371])).
fof(f10088,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X13))) ) | $spl7),
  inference(subsumption_resolution,[],[f10076,f507])).
fof(f10076,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X13))) ) | $spl7),
  inference(superposition,[],[f679,f383])).
fof(f10087,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(subsumption_resolution,[],[f10075,f507])).
fof(f10075,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(superposition,[],[f679,f383])).
fof(f9929,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f532])).
fof(f9928,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f530])).
fof(f9927,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f528])).
fof(f9926,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f526])).
fof(f9924,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f507])).
fof(f9923,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) )),
  inference(superposition,[],[f351,f383])).
fof(f9922,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f383])).
fof(f9921,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f351,f382])).
fof(f9920,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f382])).
fof(f9913,plain,(
  ( ! [X11] : (~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) ) | $spl8),
  inference(subsumption_resolution,[],[f9901,f530])).
fof(f9901,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) ) | $spl8),
  inference(superposition,[],[f2540,f382])).
fof(f9912,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl8),
  inference(subsumption_resolution,[],[f9900,f530])).
fof(f9900,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl8),
  inference(superposition,[],[f2540,f382])).
fof(f9893,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(subsumption_resolution,[],[f9891,f511])).
fof(f9891,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(resolution,[],[f8224,f484])).
fof(f9892,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(subsumption_resolution,[],[f9890,f511])).
fof(f9890,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(resolution,[],[f8224,f488])).
fof(f9889,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X1))) ) | $spl8),
  inference(subsumption_resolution,[],[f9887,f511])).
fof(f9887,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X1))) ) | $spl8),
  inference(resolution,[],[f8222,f484])).
fof(f9888,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f9886,f511])).
fof(f9886,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X0))) ) | $spl8),
  inference(resolution,[],[f8222,f488])).
fof(f9373,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f382])).
fof(f9374,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f351,f382])).
fof(f9375,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f383])).
fof(f9376,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) )),
  inference(superposition,[],[f351,f383])).
fof(f9377,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f507])).
fof(f9379,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f526])).
fof(f9380,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f528])).
fof(f9381,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f530])).
fof(f9382,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f532])).
fof(f9872,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f9859])).
fof(f9859,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl6),
  inference(resolution,[],[f9249,f374])).
fof(f9249,plain,(
  ( ! [X2,X3] : (sP3(X3,X2,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'F_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f1907,f371])).
fof(f9134,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f9122])).
fof(f9122,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl8),
  inference(resolution,[],[f8427,f374])).
fof(f8427,plain,(
  ( ! [X2,X3] : (sP3(X3,X2,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl8),
  inference(superposition,[],[f2540,f371])).
fof(f8611,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X13))) ) | $spl7),
  inference(subsumption_resolution,[],[f8599,f507])).
fof(f8599,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X13))) ) | $spl7),
  inference(superposition,[],[f679,f383])).
fof(f8610,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(subsumption_resolution,[],[f8598,f507])).
fof(f8598,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(superposition,[],[f679,f383])).
fof(f8460,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f532])).
fof(f8459,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f530])).
fof(f8458,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f528])).
fof(f8457,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f526])).
fof(f8455,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f507])).
fof(f8454,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) )),
  inference(superposition,[],[f351,f383])).
fof(f8453,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f383])).
fof(f8452,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f351,f382])).
fof(f8451,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f382])).
fof(f8444,plain,(
  ( ! [X11] : (~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) ) | $spl8),
  inference(subsumption_resolution,[],[f8432,f530])).
fof(f8432,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) ) | $spl8),
  inference(superposition,[],[f2540,f382])).
fof(f8443,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl8),
  inference(subsumption_resolution,[],[f8431,f530])).
fof(f8431,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl8),
  inference(superposition,[],[f2540,f382])).
fof(f8424,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(subsumption_resolution,[],[f8422,f511])).
fof(f8422,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(resolution,[],[f8224,f484])).
fof(f8423,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(subsumption_resolution,[],[f8421,f511])).
fof(f8421,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl8),
  inference(resolution,[],[f8224,f488])).
fof(f8420,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X1))) ) | $spl8),
  inference(subsumption_resolution,[],[f8418,f511])).
fof(f8418,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X1))) ) | $spl8),
  inference(resolution,[],[f8222,f484])).
fof(f8419,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f8417,f511])).
fof(f8417,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X0))) ) | $spl8),
  inference(resolution,[],[f8222,f488])).
fof(f8416,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f8325,f1778])).
fof(f8325,plain,(
  ( ! [X216,X215] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X215),s('type.num.num',X216))))) ) | $spl20),
  inference(resolution,[],[f1778,f498])).
fof(f8413,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f8321,f1778])).
fof(f8321,plain,(
  ( ! [X206,X207] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X206),s('type.num.num',X207))))) ) | $spl20),
  inference(resolution,[],[f1778,f463])).
fof(f8412,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f8320,f1778])).
fof(f8320,plain,(
  ( ! [X204,X205] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X204))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X205))))))) ) | $spl20),
  inference(resolution,[],[f1778,f498])).
fof(f8411,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f8317,f1778])).
fof(f8317,plain,(
  ( ! [X198,X199] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X198),s('type.num.num',X199))))) ) | $spl20),
  inference(resolution,[],[f1778,f463])).
fof(f8410,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f8409,f1778])).
fof(f8409,plain,(
  ( ! [X189,X190] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X189),s('type.pair.prod'('type.num.num','type.num.num'),X190))))) ) | $spl20),
  inference(subsumption_resolution,[],[f8313,f1778])).
fof(f8313,plain,(
  ( ! [X189,X190] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X189),s('type.pair.prod'('type.num.num','type.num.num'),X190)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X189),s('type.pair.prod'('type.num.num','type.num.num'),X190))))) ) | $spl20),
  inference(resolution,[],[f1778,f400])).
fof(f8408,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f8407,f1778])).
fof(f8407,plain,(
  ( ! [X187,X188] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X187),s('type.pair.prod'('type.num.num','type.num.num'),X188))))) ) | $spl20),
  inference(subsumption_resolution,[],[f8312,f1778])).
fof(f8312,plain,(
  ( ! [X187,X188] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X187),s('type.pair.prod'('type.num.num','type.num.num'),X188)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X188),s('type.pair.prod'('type.num.num','type.num.num'),X187))))) ) | $spl20),
  inference(resolution,[],[f1778,f400])).
fof(f8406,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f8405,f1778])).
fof(f8405,plain,(
  ( ! [X182,X181] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X181),s('type.pair.prod'('type.num.num','type.num.num'),X182))))) ) | $spl20),
  inference(subsumption_resolution,[],[f8309,f1778])).
fof(f8309,plain,(
  ( ! [X182,X181] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X181),s('type.pair.prod'('type.num.num','type.num.num'),X182)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X182),s('type.pair.prod'('type.num.num','type.num.num'),X181))))) ) | $spl20),
  inference(resolution,[],[f1778,f400])).
fof(f8400,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f8300,f1778])).
fof(f8300,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(resolution,[],[f1778,f540])).
fof(f540,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f347])).
fof(f347,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f208])).
fof(f208,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f24])).
fof(f24,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.NOT_CLAUSES|split|1')).
fof(f8399,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f8299,f1778])).
fof(f8299,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(resolution,[],[f1778,f542])).
fof(f542,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f209])).
fof(f209,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f26])).
fof(f26,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.integer.TINT_LT_WELLDEFL','thm.bool.NOT_CLAUSES|split|2')).
fof(f8297,plain,(
  $false | $spl20),
  inference(resolution,[],[f1778,f545])).
fof(f8298,plain,(
  $false | $spl20),
  inference(resolution,[],[f1778,f544])).
fof(f8311,plain,(
  $false | $spl20),
  inference(resolution,[],[f1778,f7127])).
fof(f8326,plain,(
  $false | $spl20),
  inference(resolution,[],[f1778,f538])).
fof(f7468,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f382])).
fof(f7469,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f351,f382])).
fof(f7470,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f383])).
fof(f7471,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) )),
  inference(superposition,[],[f351,f383])).
fof(f7472,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f507])).
fof(f7474,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f526])).
fof(f7475,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f528])).
fof(f7476,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f530])).
fof(f7477,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f532])).
fof(f8215,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f8144,f8046])).
fof(f8046,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f8045,f528])).
fof(f8045,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X12))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f677,f678])).
fof(f8144,plain,(
  ( ! [X222,X221] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X221),s('type.num.num',X222))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f8046,f497])).
fof(f8214,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f8139,f8046])).
fof(f8139,plain,(
  ( ! [X210,X209] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X209),s('type.num.num',X210))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f8046,f464])).
fof(f8213,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f8138,f8046])).
fof(f8138,plain,(
  ( ! [X208,X207] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X207))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X208))))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f8046,f497])).
fof(f8210,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f8135,f8046])).
fof(f8135,plain,(
  ( ! [X202,X201] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X201),s('type.num.num',X202))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f8046,f464])).
fof(f8205,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f8124,f8046])).
fof(f8124,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl6)),
  inference(resolution,[],[f8046,f543])).
fof(f543,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f348])).
fof(f8204,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f8123,f8046])).
fof(f8123,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl6)),
  inference(resolution,[],[f8046,f541])).
fof(f541,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f347])).
fof(f8122,plain,(
  $false | ($spl4 | $spl6)),
  inference(resolution,[],[f8046,f539])).
fof(f8130,plain,(
  $false | ($spl4 | $spl6)),
  inference(resolution,[],[f8046,f401])).
fof(f8133,plain,(
  $false | ($spl4 | $spl6)),
  inference(resolution,[],[f8046,f6814])).
fof(f7368,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) ) | $spl6),
  inference(subsumption_resolution,[],[f7356,f507])).
fof(f7356,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) ) | $spl6),
  inference(superposition,[],[f1907,f383])).
fof(f7367,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl6),
  inference(subsumption_resolution,[],[f7355,f507])).
fof(f7355,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl6),
  inference(superposition,[],[f1907,f383])).
fof(f4911,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f382])).
fof(f4912,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f351,f382])).
fof(f4913,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f383])).
fof(f4914,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) )),
  inference(superposition,[],[f351,f383])).
fof(f4915,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f507])).
fof(f4917,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f526])).
fof(f4918,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f528])).
fof(f4919,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f530])).
fof(f4920,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f532])).
fof(f7333,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X18),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X18),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X19))) ) | $spl8),
  inference(subsumption_resolution,[],[f7320,f530])).
fof(f7320,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,X19) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X18),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X18),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X19))) ) | $spl8),
  inference(superposition,[],[f4770,f382])).
fof(f4770,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),sK10))))) ) | $spl8),
  inference(superposition,[],[f2540,f398])).
fof(f7331,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X17)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl8),
  inference(subsumption_resolution,[],[f7319,f530])).
fof(f7319,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | ~p(s(bool,X17)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X16),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) ) | $spl8),
  inference(superposition,[],[f4770,f382])).
fof(f7314,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl8),
  inference(superposition,[],[f4770,f568])).
fof(f7139,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f7040,f1771])).
fof(f7040,plain,(
  ( ! [X10,X9] : (sP2(X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X10)) )),
  inference(resolution,[],[f6814,f361])).
fof(f7147,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl40),
  inference(cnf_transformation,[],[f7147_D])).
fof(f7147_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f7138,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) )),
  inference(resolution,[],[f7040,f1367])).
fof(f5480,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl28),
  inference(resolution,[],[f4940,f1700])).
fof(f4746,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X9))) ) | $spl1),
  inference(resolution,[],[f749,f1771])).
fof(f749,plain,(
  ( ! [X10,X9] : (sP2(X9,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)),X10)) ) | $spl1),
  inference(resolution,[],[f675,f361])).
fof(f4751,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl34),
  inference(cnf_transformation,[],[f4751_D])).
fof(f4751_D,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4745,plain,(
  ( ! [X8] : (sP3(X8,'F_0','const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) ) | $spl1),
  inference(resolution,[],[f749,f1367])).
fof(f4302,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f382])).
fof(f4303,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f351,f382])).
fof(f4304,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f383])).
fof(f4305,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) )),
  inference(superposition,[],[f351,f383])).
fof(f4306,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f507])).
fof(f4308,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f526])).
fof(f4309,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f528])).
fof(f4310,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f530])).
fof(f4311,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f532])).
fof(f4542,plain,(
  $false | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f4354,f539])).
fof(f4354,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl24)),
  inference(backward_demodulation,[],[f678,f1803])).
fof(f2309,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f528])).
fof(f2308,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f526])).
fof(f2306,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f507])).
fof(f4064,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3973,f3877])).
fof(f3877,plain,(
  ( ! [X10] : (~p(s(bool,X10))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3876,f532])).
fof(f3876,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10))) ) | ($spl2 | $spl8)),
  inference(forward_demodulation,[],[f676,f680])).
fof(f3973,plain,(
  ( ! [X216,X215] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X215),s('type.num.num',X216))))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f3877,f498])).
fof(f4061,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3969,f3877])).
fof(f3969,plain,(
  ( ! [X206,X207] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X206),s('type.num.num',X207))))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f3877,f463])).
fof(f4060,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3968,f3877])).
fof(f3968,plain,(
  ( ! [X204,X205] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X204))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X205))))))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f3877,f498])).
fof(f4059,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3965,f3877])).
fof(f3965,plain,(
  ( ! [X198,X199] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X198),s('type.num.num',X199))))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f3877,f463])).
fof(f4058,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4057,f3877])).
fof(f4057,plain,(
  ( ! [X189,X190] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X189),s('type.pair.prod'('type.num.num','type.num.num'),X190))))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3961,f3877])).
fof(f3961,plain,(
  ( ! [X189,X190] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X189),s('type.pair.prod'('type.num.num','type.num.num'),X190)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X189),s('type.pair.prod'('type.num.num','type.num.num'),X190))))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f3877,f400])).
fof(f4056,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4055,f3877])).
fof(f4055,plain,(
  ( ! [X187,X188] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X187),s('type.pair.prod'('type.num.num','type.num.num'),X188))))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3960,f3877])).
fof(f3960,plain,(
  ( ! [X187,X188] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X187),s('type.pair.prod'('type.num.num','type.num.num'),X188)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X188),s('type.pair.prod'('type.num.num','type.num.num'),X187))))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f3877,f400])).
fof(f4054,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4053,f3877])).
fof(f4053,plain,(
  ( ! [X182,X181] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X181),s('type.pair.prod'('type.num.num','type.num.num'),X182))))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3958,f3877])).
fof(f3958,plain,(
  ( ! [X182,X181] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X181),s('type.pair.prod'('type.num.num','type.num.num'),X182)))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),X182),s('type.pair.prod'('type.num.num','type.num.num'),X181))))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f3877,f400])).
fof(f4048,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3949,f3877])).
fof(f3949,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl8)),
  inference(resolution,[],[f3877,f540])).
fof(f4047,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f3948,f3877])).
fof(f3948,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl8)),
  inference(resolution,[],[f3877,f542])).
fof(f3946,plain,(
  $false | ($spl2 | $spl8)),
  inference(resolution,[],[f3877,f545])).
fof(f3947,plain,(
  $false | ($spl2 | $spl8)),
  inference(resolution,[],[f3877,f544])).
fof(f3974,plain,(
  $false | ($spl2 | $spl8)),
  inference(resolution,[],[f3877,f538])).
fof(f3014,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) | $spl33),
  inference(cnf_transformation,[],[f3014_D])).
fof(f3014_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f3007,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) ) | $spl12),
  inference(resolution,[],[f1700,f906])).
fof(f2510,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X8))) ) | $spl1),
  inference(resolution,[],[f749,f1771])).
fof(f2302,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f382])).
fof(f2303,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f351,f382])).
fof(f2304,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f383])).
fof(f2305,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) )),
  inference(superposition,[],[f351,f383])).
fof(f2310,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f530])).
fof(f2311,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f532])).
fof(f2408,plain,(
  $false | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f2401,f1797])).
fof(f1797,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl7),
  inference(trivial_inequality_removal,[],[f1794])).
fof(f1794,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl7),
  inference(superposition,[],[f679,f528])).
fof(f2401,plain,(
  ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl9),
  inference(trivial_inequality_removal,[],[f2400])).
fof(f2400,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl9),
  inference(superposition,[],[f681,f532])).
fof(f2407,plain,(
  $false | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f2402,f1797])).
fof(f2402,plain,(
  ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl9),
  inference(trivial_inequality_removal,[],[f2399])).
fof(f2399,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl9),
  inference(superposition,[],[f681,f530])).
fof(f2111,plain,(
  p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl30),
  inference(cnf_transformation,[],[f2111_D])).
fof(f2111_D,plain,(
  p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2109,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl6),
  inference(subsumption_resolution,[],[f1927,f511])).
fof(f1927,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl6),
  inference(superposition,[],[f400,f678])).
fof(f2106,plain,(
  p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl26),
  inference(cnf_transformation,[],[f2106_D])).
fof(f2106_D,plain,(
  p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2104,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl6),
  inference(subsumption_resolution,[],[f1926,f511])).
fof(f1926,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK11),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | $spl6),
  inference(superposition,[],[f400,f678])).
fof(f1804,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X13))) ) | $spl7),
  inference(subsumption_resolution,[],[f1790,f507])).
fof(f1790,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11)))) | p(s(bool,X13))) ) | $spl7),
  inference(superposition,[],[f679,f383])).
fof(f1801,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(subsumption_resolution,[],[f1789,f507])).
fof(f1789,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))))) ) | $spl7),
  inference(superposition,[],[f679,f383])).
fof(f1776,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))) | ~p(s(bool,X8))) ) | $spl1),
  inference(resolution,[],[f1771,f749])).
fof(f1286,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1220,f1119])).
fof(f1220,plain,(
  ( ! [X222,X221] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X221),s('type.num.num',X222))))) ) | $spl16),
  inference(resolution,[],[f1119,f497])).
fof(f1285,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1215,f1119])).
fof(f1215,plain,(
  ( ! [X210,X209] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X209),s('type.num.num',X210))))) ) | $spl16),
  inference(resolution,[],[f1119,f464])).
fof(f1284,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1214,f1119])).
fof(f1214,plain,(
  ( ! [X208,X207] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X207))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X208))))))) ) | $spl16),
  inference(resolution,[],[f1119,f497])).
fof(f1281,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1211,f1119])).
fof(f1211,plain,(
  ( ! [X202,X201] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X201),s('type.num.num',X202))))) ) | $spl16),
  inference(resolution,[],[f1119,f464])).
fof(f1276,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1200,f1119])).
fof(f1200,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(resolution,[],[f1119,f543])).
fof(f1275,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1199,f1119])).
fof(f1199,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(resolution,[],[f1119,f541])).
fof(f1198,plain,(
  $false | $spl16),
  inference(resolution,[],[f1119,f539])).
fof(f1206,plain,(
  $false | $spl16),
  inference(resolution,[],[f1119,f401])).
fof(f1209,plain,(
  $false | ($spl1 | $spl16)),
  inference(resolution,[],[f1119,f675])).
fof(f1120,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) | $spl18),
  inference(cnf_transformation,[],[f1120_D])).
fof(f1120_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1118,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f1113])).
fof(f1113,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f1094,f374])).
fof(f1094,plain,(
  ( ! [X4,X5] : (sP3(X5,X4,'F_0') | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) )),
  inference(superposition,[],[f1088,f371])).
fof(f835,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl12),
  inference(cnf_transformation,[],[f835_D])).
fof(f835_D,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f743,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | $spl1),
  inference(resolution,[],[f675,f400])).
fof(f742,plain,(
  p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK12),s('type.pair.prod'('type.num.num','type.num.num'),sK10)))) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | $spl1),
  inference(resolution,[],[f675,f400])).
fof(f673,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f532])).
fof(f672,plain,(
  s(bool,'T_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f530])).
fof(f671,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f528])).
fof(f670,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f526])).
fof(f668,plain,(
  s(bool,'F_0') != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))),
  inference(superposition,[],[f351,f507])).
fof(f667,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | p(s(bool,X13))) )),
  inference(superposition,[],[f351,f383])).
fof(f666,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | p(s(bool,X12)) | p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f383])).
fof(f665,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f351,f382])).
fof(f664,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK11))) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.tint_lt_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK10),s('type.pair.prod'('type.num.num','type.num.num'),sK12))))) )),
  inference(superposition,[],[f351,f382])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1692
% Passive clauses: 16362
% Generated clauses: 79270
% Final active clauses: 1004
% Final passive clauses: 9093
% Input formulas: 133
% Initial clauses: 214
% 
% Selected by SInE selection: 74
% SInE iterations: 1
% 
% Duplicate literals: 3205
% Trivial inequalities: 536
% Fw subsumption resolutions: 9709
% Fw demodulations: 15152
% Bw demodulations: 60
% 
% Simple tautologies: 5637
% Equational tautologies: 47
% Forward subsumptions: 29786
% 
% Binary resolution: 27812
% Forward superposition: 13878
% Backward superposition: 7311
% Equality resolution: 12
% 
% Split clauses: 499
% Split components: 1014
% SAT solver clauses: 591
% SAT solver unit clauses: 32
% SAT solver binary clauses: 157
% 
% Sat splits: 499
% Sat splitting refutations: 82
% 
% Memory used [KB]: 18166
% Time elapsed: 1.094 s
% ------------------------------
% ------------------------------
% Success in time 12.828 s
