Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.divides
% (54685)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 790
% Passive clauses: 9039
% Generated clauses: 25001
% Final active clauses: 648
% Final passive clauses: 6867
% Input formulas: 49
% Initial clauses: 185
% 
% Duplicate literals: 291
% Trivial inequalities: 113
% Fw subsumption resolutions: 1761
% Fw demodulations: 832
% Bw demodulations: 80
% 
% Simple tautologies: 242
% Equational tautologies: 55
% Forward subsumptions: 12458
% 
% Binary resolution: 9875
% Factoring: 1000
% Forward superposition: 4091
% Backward superposition: 6424
% Self superposition: 58
% Equality factoring: 26
% Equality resolution: 26
% 
% Split clauses: 313
% Split components: 627
% SAT solver clauses: 674
% SAT solver unit clauses: 21
% SAT solver binary clauses: 135
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 4
% 
% Sat splits: 313
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 670
% TWLsolver calls for satisfiability: 437
% 
% Memory used [KB]: 11385
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.divides
% (54695)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 464
% Passive clauses: 2856
% Generated clauses: 28081
% Final active clauses: 281
% Final passive clauses: 1980
% Input formulas: 49
% Initial clauses: 185
% 
% Duplicate literals: 82
% Trivial inequalities: 304
% Fw subsumption resolutions: 1068
% Fw demodulations: 5861
% Bw demodulations: 409
% Condensations: 18
% Global subsumptions: 13
% 
% Simple tautologies: 375
% Equational tautologies: 6
% Forward subsumptions: 13721
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 45
% 
% Binary resolution: 1243
% Factoring: 10
% Forward superposition: 7755
% Backward superposition: 6800
% Self superposition: 78
% Equality factoring: 2
% Equality resolution: 20
% 
% Split clauses: 213
% Split components: 426
% SAT solver clauses: 3507
% SAT solver unit clauses: 201
% SAT solver binary clauses: 93
% SAT solver learnt clauses: 32
% SAT solver learnt literals: 21
% 
% Sat splits: 213
% Sat splitting refutations: 74
% 
% TWLsolver clauses: 3449
% TWLsolver calls for satisfiability: 1978
% 
% Memory used [KB]: 22259
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.divides
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.divides
% SZS output start Proof for thm.divides
fof(f3635,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f256,f452,f450,f255,f451,f449,f655,f665,f663,f1866,f1870,f1873,f1876,f1879,f1882,f1885,f1888,f1891,f1894,f1897,f1900,f1903,f1906,f1909,f1912,f1915,f1918,f1949,f1867,f1951,f1871,f1953,f1874,f1955,f1877,f1957,f1880,f1959,f1883,f1961,f1886,f1963,f1889,f1965,f1892,f1967,f1895,f1969,f1898,f1971,f1901,f1973,f1904,f1975,f1907,f1977,f1910,f1979,f1913,f1981,f1916,f1983,f1919,f1998,f1999,f2001,f2444,f2446,f2450,f2454,f2460,f2464,f2490,f2504,f2506,f2511,f2513,f2518,f2520,f2523,f2525,f2533,f2538,f2544,f2549,f2551,f2557,f2562,f2568,f2573,f2577,f2579,f2587,f2589,f2609,f2614,f2620,f2625,f2645,f2649,f2654,f2658,f2670,f2688,f2718,f2720,f2724,f2730,f2732,f2737,f2751,f2757,f2782,f2783,f2784,f2785,f2786,f2787,f2788,f2789,f2790,f2791,f2792,f2793,f2794,f2795,f2796,f2797,f2798,f2799,f2800,f2801,f2802,f2803,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f2813,f2814,f2815,f2816,f2817,f2818,f2819,f2820,f2821,f2822,f2823,f2828,f2829,f3262,f3265,f3263,f3307,f3310,f1868,f3308,f3384,f3393,f3391,f3609])).
fof(f3609,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3608,f3402])).
fof(f3402,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK13)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f452,f328])).
fof(f328,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) & (s('type.num.num',X0) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) & ((s('type.num.num',X0) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X9] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X9))) <=> (s('type.num.num',X6) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & s('type.num.num',X9) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.divides.DIVIDES_ONE','thm.arithmetic.MULT_EQ_|49|')).
fof(f3608,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK18('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),sK13)),s('type.num.num',sK13))) | $spl0),
  inference(subsumption_resolution,[],[f3559,f257])).
fof(f257,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f16])).
fof(f16,axiom,(
  p(s(bool,'T_0'))),
  file('thm.divides.DIVIDES_ONE','thm.bool.TRUTH')).
fof(f3559,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK18('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),sK13)),s('type.num.num',sK13))) | $spl0),
  inference(superposition,[],[f325,f2836])).
fof(f2836,plain,(
  s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f257,f449,f320])).
fof(f320,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.divides.DIVIDES_ONE','thm.bool.IMP_ANTISYM_AX')).
fof(f325,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK18(X1,X0)),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK18(X1,X0)),s('type.num.num',X0)))) & (! [X2] : s('type.num.num',X1) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f200])).
fof(f200,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ? [X3] : s('type.num.num',X1) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X0)))) & (! [X2] : s('type.num.num',X1) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f199])).
fof(f199,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ? [X2] : s('type.num.num',X1) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0)))) & (! [X2] : s('type.num.num',X1) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> ? [X2] : s('type.num.num',X1) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X18,X19] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X18),s('type.num.num',X19)))) <=> ? [X20] : s('type.num.num',X19) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',X18))))),
  file('thm.divides.DIVIDES_ONE','thm.divides.divides_def')).
fof(f3391,plain,(
  ( ! [X1] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl50),
  inference(cnf_transformation,[],[f3391_D])).
fof(f3391_D,plain,(
  ( ! [X1] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f3393,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl53),
  inference(cnf_transformation,[],[f3393_D])).
fof(f3393_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f3384,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f327,f265])).
fof(f265,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.divides.DIVIDES_ONE','thm.arithmetic.MULT_|48|')).
fof(f327,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(cnf_transformation,[],[f203])).
fof(f3308,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK15('T_0',X2,X1)))))) ) | $spl48),
  inference(cnf_transformation,[],[f3308_D])).
fof(f3308_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK15('T_0',X2,X1)))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1868,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f1868_D])).
fof(f1868_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f3310,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK15('T_0',X2,X1)))))) )),
  inference(condensation,[],[f3309])).
fof(f3309,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,happ(s(fun(X30,bool),X31),s(X30,X32)))) | ~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,sK15('T_0',X31,X30))))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f3302,f881])).
fof(f881,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f875,f453])).
fof(f453,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f260,f257])).
fof(f260,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f25])).
fof(f25,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.divides.DIVIDES_ONE','thm.bool.NOT_CLAUSES|split|2')).
fof(f875,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f317,f268])).
fof(f268,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.divides.DIVIDES_ONE',aHL_BOOL_CASES)).
fof(f317,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f135])).
fof(f3302,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,happ(s(fun(X30,bool),X31),s(X30,X32)))) | ~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,sK15('T_0',X31,X30))))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f309,f782])).
fof(f782,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f268,f268])).
fof(f309,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X2,X1,X0))))) | ~p(s(bool,X2)) | ! [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X2,X1,X0))))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : ((? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2)) | ! [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2)) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : (((? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : ((! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : ! [X1,X2] : ((! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X10] : ! [X11,X15] : ((! [X6] : p(s(bool,happ(s(fun(X10,bool),X11),s(X10,X6)))) & p(s(bool,X15))) <=> ! [X6] : (p(s(bool,happ(s(fun(X10,bool),X11),s(X10,X6)))) & p(s(bool,X15))))),
  file('thm.divides.DIVIDES_ONE','thm.bool.LEFT_AND_FORALL_THM')).
fof(f3307,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK15('T_0',X2,X1)))))) )),
  inference(condensation,[],[f3306])).
fof(f3306,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK15('T_0',X26,X25))))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f3301,f881])).
fof(f3301,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK15('T_0',X26,X25))))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f309,f782])).
fof(f3263,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK14('T_0',X2,X1)))))) ) | $spl46),
  inference(cnf_transformation,[],[f3263_D])).
fof(f3263_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK14('T_0',X2,X1)))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f3265,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK14('T_0',X2,X1)))))) )),
  inference(condensation,[],[f3264])).
fof(f3264,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,happ(s(fun(X30,bool),X31),s(X30,X32)))) | ~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,sK14('T_0',X31,X30))))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f3257,f881])).
fof(f3257,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,happ(s(fun(X30,bool),X31),s(X30,X32)))) | ~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,sK14('T_0',X31,X30))))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f307,f782])).
fof(f307,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f194])).
fof(f3262,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK14('T_0',X2,X1)))))) )),
  inference(condensation,[],[f3261])).
fof(f3261,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK14('T_0',X26,X25))))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f3256,f881])).
fof(f3256,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK14('T_0',X26,X25))))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f307,f782])).
fof(f2829,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2449,f257])).
fof(f2449,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f523,f2404,f320])).
fof(f2404,plain,(
  ( ! [X0] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X0))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f456,f326])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X1) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f201])).
fof(f456,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK13),s('type.num.num',X0)))) ) | $spl2),
  inference(backward_demodulation,[],[f451,f266])).
fof(f266,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X7)))),
  file('thm.divides.DIVIDES_ONE','thm.arithmetic.MULT_LEFT_|49|')).
fof(f523,plain,(
  s(bool,'T_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f465,f275])).
fof(f275,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.divides.DIVIDES_ONE','thm.bool.EQ_CLAUSES|split|1')).
fof(f465,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f450,f451])).
fof(f2828,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2824,f522])).
fof(f522,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f453,f275])).
fof(f2824,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f2451,f530])).
fof(f530,plain,(
  s(bool,'F_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f465,f280])).
fof(f280,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.divides.DIVIDES_ONE','thm.bool.EQ_CLAUSES|split|3')).
fof(f2451,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f257,f2404,f320])).
fof(f2823,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2453,f257])).
fof(f2453,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f523,f2404,f320])).
fof(f2822,plain,(
  $false | ($spl1 | $spl2 | $spl12)),
  inference(subsumption_resolution,[],[f2459,f257])).
fof(f2459,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl12)),
  inference(unit_resulting_resolution,[],[f2353,f2404,f346])).
fof(f346,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f211])).
fof(f211,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f2353,plain,(
  sP1('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0') | ($spl1 | $spl2 | $spl12)),
  inference(subsumption_resolution,[],[f2348,f453])).
fof(f2348,plain,(
  p(s(bool,'F_0')) | sP1('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0') | ($spl1 | $spl2 | $spl12)),
  inference(superposition,[],[f1871,f530])).
fof(f2821,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2463,f257])).
fof(f2463,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f2317,f2404,f346])).
fof(f2317,plain,(
  sP1('T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13))) | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2312,f453])).
fof(f2312,plain,(
  p(s(bool,'F_0')) | sP1('T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13))) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f1867,f530])).
fof(f2820,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2489,f257])).
fof(f2489,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f927,f2404,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP3(X2,X1,X0) & sP4(X1,X2,X0) & sP5(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP3(X2,X1,X0) | ~sP4(X1,X2,X0) | ~sP5(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f221])).
fof(f221,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP6(X2,X1,X0)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP6(X2,X1,X0)))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X2,X1,X0] : (sP6(X2,X1,X0) <=> (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f927,plain,(
  sP6('T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f530,f386])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP6(X2,X1,X0)) & (~sP6(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP6(X2,X1,X0))),
  inference(definition_folding,[],[f115,f159,f158,f157,f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : (sP3(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f157,plain,(
  ! [X1,X0,X2] : (sP4(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f158,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f115,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X21,X20,X22] : ((p(s(bool,X22)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X22)) | p(s(bool,X20)) | p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X20))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.divides.DIVIDES_ONE','thm.sat.sat_2|49|_prove_ltb')).
fof(f2819,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2492,f484])).
fof(f484,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f257,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(nnf_transformation,[],[f156])).
fof(f2492,plain,(
  ~sP3('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0','T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f481,f476,f2404,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~sP3(X2,X1,X0) | ~p(s(bool,X2)) | ~sP5(X0,X2,X1) | ~sP4(X1,X2,X0) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f222])).
fof(f476,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),X1)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f465,f375])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f158])).
fof(f481,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),X1)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f465,f379])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f227])).
fof(f227,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(nnf_transformation,[],[f157])).
fof(f939,plain,(
  ~sP6('T_0','T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f523,f387])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f2818,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2494,f484])).
fof(f2494,plain,(
  ~sP3('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0','T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f479,f476,f2404,f370])).
fof(f479,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f257,f378])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f228])).
fof(f2817,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2496,f484])).
fof(f2496,plain,(
  ~sP3('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0','T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f481,f474,f2404,f370])).
fof(f474,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f257,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f2816,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2497,f484])).
fof(f2497,plain,(
  ~sP3('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0','T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f479,f474,f2404,f370])).
fof(f2815,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2500,f474])).
fof(f2500,plain,(
  ~sP5('T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f481,f484,f2404,f370])).
fof(f2814,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2501,f474])).
fof(f2501,plain,(
  ~sP5('T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f479,f484,f2404,f370])).
fof(f2813,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2502,f479])).
fof(f2502,plain,(
  ~sP4('T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f476,f484,f2404,f370])).
fof(f2812,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2509,f479])).
fof(f2509,plain,(
  ~sP4('T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f474,f484,f2404,f370])).
fof(f2811,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2514,f474])).
fof(f2514,plain,(
  ~sP5('T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f481,f485,f2404,f370])).
fof(f485,plain,(
  ( ! [X0,X1] : (sP3(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f257,f383])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f230])).
fof(f2810,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2515,f474])).
fof(f2515,plain,(
  ~sP5('T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f479,f485,f2404,f370])).
fof(f2809,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2516,f479])).
fof(f2516,plain,(
  ~sP4('T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f476,f485,f2404,f370])).
fof(f2808,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2521,f479])).
fof(f2521,plain,(
  ~sP4('T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f474,f485,f2404,f370])).
fof(f2807,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2527,f2582])).
fof(f2582,plain,(
  ( ! [X0] : (~sP6('const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X0)),'T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f523,f2404,f387])).
fof(f2527,plain,(
  sP6('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f481,f478,f485,f2404,f370])).
fof(f478,plain,(
  ( ! [X0,X1] : (sP5('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),X0,X1)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f465,f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f2806,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2532,f2750])).
fof(f2750,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f632,f2404,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f149,f168,f167])).
fof(f167,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f168,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f149,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X21,X20,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X20)) => p(s(bool,X21)))) <=> ((p(s(bool,X22)) | p(s(bool,X20))) & (p(s(bool,X22)) | ~p(s(bool,X21))) & (~p(s(bool,X20)) | p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.divides.DIVIDES_ONE','thm.sat.sat_27_prove_ltb')).
fof(f632,plain,(
  ( ! [X0] : (sP12('T_0',X0,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f497,f257,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP11(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f248])).
fof(f248,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1)) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f168])).
fof(f497,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f465,f431])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f251])).
fof(f251,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f167])).
fof(f2532,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f478,f2404,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f225])).
fof(f2805,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2536,f2750])).
fof(f2536,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f476,f2404,f373])).
fof(f2804,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2543,f2750])).
fof(f2543,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f476,f2404,f373])).
fof(f2803,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2547,f2750])).
fof(f2547,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f478,f2404,f373])).
fof(f2802,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2556,f2750])).
fof(f2556,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f483,f2404,f377])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f228])).
fof(f483,plain,(
  ( ! [X0,X1] : (sP4('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),X0,X1)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f465,f380])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f2801,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2560,f2750])).
fof(f2560,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f481,f2404,f377])).
fof(f2800,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2567,f2750])).
fof(f2567,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f481,f2404,f377])).
fof(f2799,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2571,f2750])).
fof(f2571,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f483,f2404,f377])).
fof(f2798,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2586,f522])).
fof(f2586,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1299,f2404,f387])).
fof(f1299,plain,(
  sP6('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'F_0','T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f482,f478,f486,f257,f370])).
fof(f486,plain,(
  ( ! [X0,X1] : (sP3('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f257,f384])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f230])).
fof(f482,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f453,f380])).
fof(f2797,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2588,f523])).
fof(f2588,plain,(
  s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1300,f2404,f387])).
fof(f1300,plain,(
  sP6('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f483,f478,f486,f257,f370])).
fof(f2796,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2608,f2750])).
fof(f2608,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f491,f2404,f396])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f491,plain,(
  ( ! [X0,X1] : (sP7('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),X0,X1)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f465,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f236])).
fof(f2795,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2612,f2750])).
fof(f2612,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f489,f2404,f396])).
fof(f489,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),X1)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f465,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f236])).
fof(f2794,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2619,f2750])).
fof(f2619,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f489,f2404,f396])).
fof(f2793,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2623,f2750])).
fof(f2623,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f491,f2404,f396])).
fof(f2792,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2644,f2750])).
fof(f2644,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f650,f2404,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP8(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f117,f162,f161])).
fof(f162,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f117,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X21,X20,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X20)) & p(s(bool,X21)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X20)) | ~p(s(bool,X22))) & (p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.divides.DIVIDES_ONE','thm.sat.sat_23_prove_ltb')).
fof(f650,plain,(
  ( ! [X0] : (sP8('F_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),X0)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f491,f453,f442])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f395])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (sP7(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f233])).
fof(f233,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f162])).
fof(f2791,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2647,f2750])).
fof(f2647,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f648,f2404,f405])).
fof(f648,plain,(
  ( ! [X0] : (sP8('F_0',X0,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f489,f453,f442])).
fof(f2790,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2653,f2750])).
fof(f2653,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f648,f2404,f405])).
fof(f2789,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2656,f2750])).
fof(f2656,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f650,f2404,f405])).
fof(f2788,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2716,f2750])).
fof(f2716,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f500,f2404,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f500,plain,(
  ( ! [X0,X1] : (sP11('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),X0,X1)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f465,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f252])).
fof(f2787,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2722,f2750])).
fof(f2722,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f497,f2404,f430])).
fof(f2786,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2728,f2750])).
fof(f2728,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f497,f2404,f430])).
fof(f2785,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2735,f2750])).
fof(f2735,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f500,f2404,f430])).
fof(f2784,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2749,f257])).
fof(f2749,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f632,f2404,f437])).
fof(f2783,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2755,f453])).
fof(f2755,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f1110,f2404,f437])).
fof(f1110,plain,(
  sP12('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'F_0','T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f500,f257,f453,f429])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f249])).
fof(f2782,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2756,f257])).
fof(f2756,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f1110,f2404,f437])).
fof(f2757,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f453,f1110,f2404,f437])).
fof(f2751,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f257,f632,f2404,f437])).
fof(f2737,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f2404,f500,f2404,f430])).
fof(f2732,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f2404,f497,f2404,f430])).
fof(f2730,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f257,f497,f2404,f430])).
fof(f2724,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f2404,f497,f2404,f430])).
fof(f2720,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f2404,f500,f2404,f430])).
fof(f2718,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f257,f500,f2404,f430])).
fof(f2688,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f453,f1065,f2404,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f119,f165,f164])).
fof(f164,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f165,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f119,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X21,X20,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X20)) | p(s(bool,X21)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X20))) & (p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X20)) | p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.divides.DIVIDES_ONE','thm.sat.sat_25_prove_ltb')).
fof(f1065,plain,(
  sP10('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'F_0','F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f495,f453,f453,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP9(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f240])).
fof(f240,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1)) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f165])).
fof(f495,plain,(
  ( ! [X0,X1] : (sP9('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),X0,X1)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f465,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f243])).
fof(f243,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f164])).
fof(f2670,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f453,f495,f2404,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f244])).
fof(f2658,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f2404,f650,f2404,f405])).
fof(f2654,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f453,f648,f2404,f405])).
fof(f2649,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f2404,f648,f2404,f405])).
fof(f2645,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f453,f650,f2404,f405])).
fof(f2625,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f2404,f491,f2404,f396])).
fof(f2620,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f453,f489,f2404,f396])).
fof(f2614,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f2404,f489,f2404,f396])).
fof(f2609,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f453,f491,f2404,f396])).
fof(f2589,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f523,f1300,f2404,f387])).
fof(f2587,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f522,f1299,f2404,f387])).
fof(f2579,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1348,f2404,f384])).
fof(f1348,plain,(
  ~sP3('const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'F_0','F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f952,f481,f477,f453,f371])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (~sP3(X2,X1,X0) | p(s(bool,X1)) | ~sP5(X0,X2,X1) | ~sP4(X1,X2,X0) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f222])).
fof(f477,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f453,f376])).
fof(f952,plain,(
  ~sP6('F_0','F_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f530,f388])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f231])).
fof(f2577,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1390,f2404,f383])).
fof(f1390,plain,(
  ~sP3('F_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f951,f483,f475,f453,f372])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (~sP3(X2,X1,X0) | p(s(bool,X0)) | ~sP5(X0,X2,X1) | ~sP4(X1,X2,X0) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f222])).
fof(f475,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f453,f375])).
fof(f951,plain,(
  ~sP6('F_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f530,f388])).
fof(f2573,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f2404,f483,f2404,f377])).
fof(f2568,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f453,f481,f2404,f377])).
fof(f2562,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f2404,f481,f2404,f377])).
fof(f2557,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f453,f483,f2404,f377])).
fof(f2551,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1304,f2404,f374])).
fof(f1304,plain,(
  ~sP5('T_0','T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f935,f483,f484,f257,f370])).
fof(f935,plain,(
  ~sP6('T_0','const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num',sK13)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f523,f387])).
fof(f2549,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f2404,f478,f2404,f373])).
fof(f2544,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f453,f476,f2404,f373])).
fof(f2538,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f2404,f476,f2404,f373])).
fof(f2533,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f453,f478,f2404,f373])).
fof(f2525,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f479,f474,f485,f2404,f370])).
fof(f2523,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f481,f474,f485,f2404,f370])).
fof(f2520,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f479,f476,f485,f2404,f370])).
fof(f2518,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f481,f476,f485,f2404,f370])).
fof(f2513,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f479,f474,f484,f2404,f370])).
fof(f2511,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f481,f474,f484,f2404,f370])).
fof(f2506,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f479,f476,f484,f2404,f370])).
fof(f2504,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f939,f481,f476,f484,f2404,f370])).
fof(f2490,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f453,f927,f2404,f369])).
fof(f2464,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f257,f2317,f2404,f346])).
fof(f2460,plain,(
  $false | ($spl1 | $spl2 | $spl12)),
  inference(unit_resulting_resolution,[],[f257,f2353,f2404,f346])).
fof(f2454,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f523,f2404,f320])).
fof(f2450,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f257,f523,f2404,f320])).
fof(f2446,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f530,f2404,f279])).
fof(f279,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f176])).
fof(f2444,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f523,f2404,f276])).
fof(f276,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f174])).
fof(f2001,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f256,f255,f257,f453,f264,f265,f266,f267,f268,f276,f275,f280,f279,f303,f454,f309,f307,f314,f312,f315,f316,f320,f317,f326,f325,f329,f328,f327,f332,f331,f455,f350,f347,f346,f345,f355,f361,f360,f359,f443,f357,f356,f372,f371,f370,f369,f368,f367,f366,f376,f375,f374,f373,f380,f379,f378,f377,f384,f383,f382,f381,f388,f387,f386,f385,f442,f392,f391,f390,f389,f399,f398,f397,f396,f405,f402,f401,f400,f412,f411,f410,f441,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f429,f428,f427,f440,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f472,f473,f474,f475,f477,f479,f480,f482,f484,f485,f486,f487,f488,f490,f492,f493,f494,f496,f498,f499,f501,f522,f542,f543,f546,f552,f556,f562,f565,f575,f579,f580,f583,f589,f590,f594,f597,f601,f609,f610,f612,f620,f782,f881,f920,f918,f942,f938,f979,f996,f1005,f1009,f1014,f1018,f1029,f1038,f1064,f1073,f1081,f1109,f1116,f1125,f1307,f1351,f1463,f1573,f1866,f1870,f1873,f1876,f1879,f1882,f1885,f1888,f1891,f1894,f1897,f1900,f1903,f1906,f1909,f1912,f1915,f1918,f1920,f1921,f1922,f1923,f1924,f1925,f1926,f1927,f1928,f1929,f1930,f1931,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1949,f1951,f1953,f1955,f1957,f1959,f1961,f1963,f1965,f1967,f1969,f1971,f1973,f1975,f1977,f1979,f1981,f1983,f1868,f1997,f2000])).
fof(f2000,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK18(X1,X0)),s('type.num.num',X0)))) ) | $spl10),
  inference(subsumption_resolution,[],[f1988,f257])).
fof(f1988,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK18(X1,X0)),s('type.num.num',X0)))) ) | $spl10),
  inference(backward_demodulation,[],[f1868,f325])).
fof(f1997,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl10),
  inference(subsumption_resolution,[],[f1984,f1868])).
fof(f1984,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl10),
  inference(backward_demodulation,[],[f1868,f279])).
fof(f1943,plain,(
  ( ! [X121,X122,X120,X119] : (p(s(bool,'|dollar|exists_1'(s(fun(X119,bool),X120)))) | s(bool,X122) = s(bool,happ(s(fun(X119,bool),X120),s(X119,X121))) | s(bool,'T_0') = s(bool,X122)) )),
  inference(subsumption_resolution,[],[f1646,f257])).
fof(f1646,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X119,bool),X120)))) | s(bool,X122) = s(bool,happ(s(fun(X119,bool),X120),s(X119,X121))) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f454,f782])).
fof(f1942,plain,(
  ( ! [X101,X99,X102,X100] : (~sP9(X99,X100,X101) | sP10(X99,X100,X101) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f1635,f257])).
fof(f1635,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,'T_0')) | ~sP9(X99,X100,X101) | sP10(X99,X100,X101) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f441,f782])).
fof(f1941,plain,(
  ( ! [X97,X95,X98,X96] : (~sP11(X95,X96,X97) | sP12(X95,X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f1634,f257])).
fof(f1634,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,'T_0')) | ~sP11(X95,X96,X97) | sP12(X95,X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f440,f782])).
fof(f1940,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X92)) | p(s(bool,X93)) | ~sP12(X91,X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f1633,f257])).
fof(f1633,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | ~p(s(bool,X92)) | p(s(bool,X93)) | ~sP12(X91,X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f437,f782])).
fof(f1939,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X88)) | sP12(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1632,f257])).
fof(f1632,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,X88)) | sP12(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f436,f782])).
fof(f1938,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X84)) | sP12(X83,X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1631,f257])).
fof(f1631,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X84)) | sP12(X83,X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f435,f782])).
fof(f1937,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | ~p(s(bool,X81)) | ~sP11(X81,X80,X79) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1630,f257])).
fof(f1630,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | p(s(bool,X80)) | ~p(s(bool,X81)) | ~sP11(X81,X80,X79) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f430,f782])).
fof(f1936,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | p(s(bool,X77)) | ~sP10(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1629,f257])).
fof(f1629,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X76)) | p(s(bool,X77)) | ~sP10(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f420,f782])).
fof(f1935,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | sP10(X71,X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1628,f257])).
fof(f1628,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X72)) | sP10(X71,X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f419,f782])).
fof(f1934,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X68)) | sP10(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1627,f257])).
fof(f1627,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X68)) | sP10(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f418,f782])).
fof(f1933,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1626,f257])).
fof(f1626,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f413,f782])).
fof(f1932,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | p(s(bool,X61)) | ~sP8(X61,X60,X59) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1625,f257])).
fof(f1625,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | p(s(bool,X61)) | ~sP8(X61,X60,X59) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f405,f782])).
fof(f1931,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP7(X56,X55,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1624,f257])).
fof(f1624,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP7(X56,X55,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f396,f782])).
fof(f1930,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP4(X52,X51,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1623,f257])).
fof(f1623,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP4(X52,X51,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f377,f782])).
fof(f1929,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP5(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1622,f257])).
fof(f1622,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP5(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f373,f782])).
fof(f1928,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1621,f257])).
fof(f1621,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f369,f782])).
fof(f1927,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | sP2(X40,X41,X39) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1620,f257])).
fof(f1620,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | sP2(X40,X41,X39) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f361,f782])).
fof(f1926,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | sP2(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1619,f257])).
fof(f1619,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | sP2(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f360,f782])).
fof(f1925,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f1617,f257])).
fof(f1617,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f350,f782])).
fof(f1924,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f1616,f257])).
fof(f1616,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f347,f782])).
fof(f1923,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1615,f257])).
fof(f1615,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f332,f782])).
fof(f1922,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1614,f257])).
fof(f1614,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f331,f782])).
fof(f1921,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1613,f257])).
fof(f1613,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f320,f782])).
fof(f1920,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1609,f257])).
fof(f1609,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f279,f782])).
fof(f1573,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f275,f782])).
fof(f1463,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f782,f782])).
fof(f1351,plain,(
  sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f482,f475,f484,f453,f371])).
fof(f1307,plain,(
  sP6('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f479,f474,f485,f257,f370])).
fof(f1125,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1116,f425])).
fof(f1116,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f257,f453,f257,f430])).
fof(f1109,plain,(
  sP12('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f499,f257,f453,f429])).
fof(f1081,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1073,f408])).
fof(f1073,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f453,f453,f257,f413])).
fof(f1064,plain,(
  sP10('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f494,f453,f453,f412])).
fof(f1038,plain,(
  ~sP8('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1029,f389])).
fof(f1029,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f453,f257,f257,f396])).
fof(f1018,plain,(
  sP8('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f487,f257,f257,f392])).
fof(f1014,plain,(
  ~sP6('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1009,f366])).
fof(f1009,plain,(
  ~sP3('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f453,f453,f453,f381])).
fof(f1005,plain,(
  ~sP6('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f996,f367])).
fof(f996,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f453,f257,f257,f377])).
fof(f979,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f453,f257,f257,f373])).
fof(f938,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f257,f522,f387])).
fof(f942,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f257,f522,f387])).
fof(f918,plain,(
  sP6('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f453,f522,f385])).
fof(f920,plain,(
  sP6('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f453,f522,f385])).
fof(f620,plain,(
  ( ! [X0] : (sP12('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f257,f257,f436])).
fof(f612,plain,(
  ( ! [X0] : (sP12('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f453,f257,f435])).
fof(f610,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f257,f453,f424])).
fof(f609,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f453,f453,f423])).
fof(f601,plain,(
  ( ! [X0] : (sP10('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f257,f257,f419])).
fof(f597,plain,(
  ( ! [X0] : (sP10('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f257,f257,f418])).
fof(f594,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f257,f453,f407])).
fof(f590,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f257,f453,f406])).
fof(f589,plain,(
  ( ! [X0] : (sP8('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f453,f453,f401])).
fof(f583,plain,(
  ( ! [X0] : (sP8('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f453,f453,f400])).
fof(f580,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f257,f453,f391])).
fof(f579,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f257,f453,f390])).
fof(f575,plain,(
  ( ! [X0] : (sP2('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f257,f257,f361])).
fof(f565,plain,(
  ( ! [X0] : (~sP2('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f453,f453,f357])).
fof(f562,plain,(
  ( ! [X0] : (~sP2(X0,'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f453,f453,f356])).
fof(f556,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f453,f453,f455])).
fof(f552,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f453,f257,f350])).
fof(f546,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f453,f257,f347])).
fof(f543,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f257,f257,f346])).
fof(f542,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f453,f453,f345])).
fof(f501,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f257,f443])).
fof(f499,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f453,f433])).
fof(f498,plain,(
  ( ! [X0,X1] : (sP11(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f257,f432])).
fof(f496,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f453,f431])).
fof(f494,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f453,f416])).
fof(f493,plain,(
  ( ! [X0,X1] : (sP9(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f257,f415])).
fof(f492,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f257,f414])).
fof(f490,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f453,f399])).
fof(f488,plain,(
  ( ! [X0,X1] : (sP7(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f453,f398])).
fof(f487,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f257,f397])).
fof(f480,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f453,f379])).
fof(f473,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f257,f332])).
fof(f472,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f257,f331])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f254])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f254])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f254])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f252])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f249])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f249])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f249])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f249])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f249])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f246])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f244])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f244])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f241])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f409])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f241])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f241])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f238])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f238])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f238])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f234])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f234])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f234])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f234])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f231])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f230])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f222])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f222])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f222])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f216])).
fof(f216,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f217])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f217])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f217])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f217])).
fof(f355,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK19(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK19(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK19(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK19(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.divides.DIVIDES_ONE',aHL_EXT)).
fof(f345,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f212])).
fof(f347,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f212])).
fof(f350,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f212])).
fof(f455,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f330,f453])).
fof(f330,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f204])).
fof(f204,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f331,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f205])).
fof(f332,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f205])).
fof(f329,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f203])).
fof(f316,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7,X8] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X8)))))),
  file('thm.divides.DIVIDES_ONE','thm.arithmetic.MULT_SUC')).
fof(f315,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X7,X8] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X8))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X8)))),
  file('thm.divides.DIVIDES_ONE','thm.arithmetic.MULT|split|1')).
fof(f312,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK16(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK16(X2,X1,X0)))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK17(X2,X1,X0)))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X10] : ! [X11,X15] : (! [X6] : (p(s(bool,X11)) | p(s(bool,happ(s(fun(X10,bool),X15),s(X10,X6))))) <=> (p(s(bool,X11)) | ! [X6] : p(s(bool,happ(s(fun(X10,bool),X15),s(X10,X6))))))),
  file('thm.divides.DIVIDES_ONE','thm.bool.RIGHT_FORALL_OR_THM')).
fof(f314,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK17(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f198])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f304,f303])).
fof(f304,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X10] : ! [X11,X6] : (p(s(bool,happ(s(fun(X10,bool),X11),s(X10,X6)))) => p(s(bool,happ(s(fun(X10,bool),X11),s(X10,'const.min.@_1'(s(fun(X10,bool),X11)))))))),
  file('thm.divides.DIVIDES_ONE','thm.bool.SELECT_AX')).
fof(f303,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X10,bool),X5))) = s(bool,happ(s(fun(X10,bool),X5),s(X10,'const.min.@_1'(s(fun(X10,bool),X5)))))),
  file('thm.divides.DIVIDES_ONE','thm.bool.EXISTS_DEF')).
fof(f267,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.divides.DIVIDES_ONE','thm.arithmetic.MULT_RIGHT_|49|')).
fof(f264,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.divides.DIVIDES_ONE','thm.arithmetic.MULT_CLAUSES|split|0')).
fof(f1999,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1986,f257])).
fof(f1986,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(backward_demodulation,[],[f1868,f453])).
fof(f1998,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f256,f255,f257,f453,f264,f265,f266,f267,f268,f276,f275,f280,f279,f303,f454,f309,f307,f314,f312,f315,f316,f320,f317,f326,f325,f329,f328,f327,f332,f331,f455,f350,f347,f346,f345,f355,f361,f360,f359,f443,f357,f356,f372,f371,f370,f369,f368,f367,f366,f376,f375,f374,f373,f380,f379,f378,f377,f384,f383,f382,f381,f388,f387,f386,f385,f442,f392,f391,f390,f389,f399,f398,f397,f396,f405,f402,f401,f400,f412,f411,f410,f441,f408,f407,f406,f416,f415,f414,f413,f420,f419,f418,f417,f429,f428,f427,f440,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f472,f473,f474,f475,f477,f479,f480,f482,f484,f485,f486,f487,f488,f490,f492,f493,f494,f496,f498,f499,f501,f522,f542,f543,f546,f552,f556,f562,f565,f575,f579,f580,f583,f589,f590,f594,f597,f601,f609,f610,f612,f620,f782,f881,f920,f918,f942,f938,f979,f996,f1005,f1009,f1014,f1018,f1029,f1038,f1064,f1073,f1081,f1109,f1116,f1125,f1307,f1351,f1463,f1573,f1866,f1870,f1873,f1876,f1879,f1882,f1885,f1888,f1891,f1894,f1897,f1900,f1903,f1906,f1909,f1912,f1915,f1918,f1920,f1921,f1922,f1923,f1924,f1925,f1926,f1927,f1928,f1929,f1930,f1931,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1949,f1951,f1953,f1955,f1957,f1959,f1961,f1963,f1965,f1967,f1969,f1971,f1973,f1975,f1977,f1979,f1981,f1983,f1868,f1997])).
fof(f1919,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) ) | $spl44),
  inference(cnf_transformation,[],[f1919_D])).
fof(f1919_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1983,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f1982])).
fof(f1982,plain,(
  ( ! [X107,X105,X106,X104] : (~sP9('T_0',X105,X106) | sP10('T_0',X105,X106) | s(bool,X104) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(subsumption_resolution,[],[f1678,f881])).
fof(f1678,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,X104)) | ~sP9('T_0',X105,X106) | sP10('T_0',X105,X106) | s(bool,X104) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f441,f782])).
fof(f1916,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f1916_D])).
fof(f1916_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1981,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f1980])).
fof(f1980,plain,(
  ( ! [X103,X101,X102,X100] : (~sP11('T_0',X101,X102) | sP12('T_0',X101,X102) | s(bool,X100) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(subsumption_resolution,[],[f1677,f881])).
fof(f1677,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,X100)) | ~sP11('T_0',X101,X102) | sP12('T_0',X101,X102) | s(bool,X100) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f440,f782])).
fof(f1913,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) ) | $spl40),
  inference(cnf_transformation,[],[f1913_D])).
fof(f1913_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1979,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f1978])).
fof(f1978,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12('T_0',X98,X97) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f1676,f881])).
fof(f1676,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X96)) | ~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12('T_0',X98,X97) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f437,f782])).
fof(f1910,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | sP12('T_0',X1,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f1910_D])).
fof(f1910_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | sP12('T_0',X1,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1977,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f1976])).
fof(f1976,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X93)) | sP12('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f1675,f881])).
fof(f1675,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~p(s(bool,X93)) | sP12('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f436,f782])).
fof(f1907,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | sP12('T_0',X2,X1)) ) | $spl36),
  inference(cnf_transformation,[],[f1907_D])).
fof(f1907_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | sP12('T_0',X2,X1)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1975,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f1974])).
fof(f1974,plain,(
  ( ! [X90,X88,X91,X89] : (p(s(bool,X89)) | sP12('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f1674,f881])).
fof(f1674,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | p(s(bool,X89)) | sP12('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f435,f782])).
fof(f1904,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) ) | $spl34),
  inference(cnf_transformation,[],[f1904_D])).
fof(f1904_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1973,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f1972])).
fof(f1972,plain,(
  ( ! [X87,X85,X86,X84] : (p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,'T_0') | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f1673,f881])).
fof(f1673,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,'T_0') | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f430,f782])).
fof(f1901,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) ) | $spl32),
  inference(cnf_transformation,[],[f1901_D])).
fof(f1901_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1971,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f1970])).
fof(f1970,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X81)) | p(s(bool,X82)) | ~sP10('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f1672,f881])).
fof(f1672,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | p(s(bool,X82)) | ~sP10('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f420,f782])).
fof(f1898,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f1898_D])).
fof(f1898_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1969,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f1968])).
fof(f1968,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X77)) | sP10('T_0',X77,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f1671,f881])).
fof(f1671,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | ~p(s(bool,X77)) | sP10('T_0',X77,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f419,f782])).
fof(f1895,plain,(
  ( ! [X2,X1] : (sP10('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f1895_D])).
fof(f1895_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1967,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f1966])).
fof(f1966,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X73)) | sP10('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f1670,f881])).
fof(f1670,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | sP10('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f418,f782])).
fof(f1892,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | ~sP9('T_0',X1,X2) | p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f1892_D])).
fof(f1892_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | ~sP9('T_0',X1,X2) | p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1965,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1964])).
fof(f1964,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f1669,f881])).
fof(f1669,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f413,f782])).
fof(f1889,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) ) | $spl24),
  inference(cnf_transformation,[],[f1889_D])).
fof(f1889_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1963,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f1962])).
fof(f1962,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,'T_0') | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f1668,f881])).
fof(f1668,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,'T_0') | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f405,f782])).
fof(f1886,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP7(X1,'T_0',X2) | p(s(bool,X2))) ) | $spl22),
  inference(cnf_transformation,[],[f1886_D])).
fof(f1886_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP7(X1,'T_0',X2) | p(s(bool,X2))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1961,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f1960])).
fof(f1960,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,'T_0',X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f1667,f881])).
fof(f1667,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,'T_0',X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f396,f782])).
fof(f1883,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP4(X1,'T_0',X2) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f1883_D])).
fof(f1883_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP4(X1,'T_0',X2) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1959,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1958])).
fof(f1958,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,'T_0',X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f1666,f881])).
fof(f1666,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,'T_0',X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f377,f782])).
fof(f1880,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP5(X1,'T_0',X2) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f1880_D])).
fof(f1880_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP5(X1,'T_0',X2) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1957,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f1956])).
fof(f1956,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f1665,f881])).
fof(f1665,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f373,f782])).
fof(f1877,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP6(X1,'T_0',X2) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f1877_D])).
fof(f1877_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP6(X1,'T_0',X2) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1955,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1954])).
fof(f1954,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f1664,f881])).
fof(f1664,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f369,f782])).
fof(f1874,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0') | ~p(s(bool,X1))) ) | $spl14),
  inference(cnf_transformation,[],[f1874_D])).
fof(f1874_D,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0') | ~p(s(bool,X1))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1953,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2(X1,X2,'T_0')) )),
  inference(condensation,[],[f1952])).
fof(f1952,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | sP2(X45,X46,'T_0') | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f1663,f881])).
fof(f1663,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | sP2(X45,X46,'T_0') | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f361,f782])).
fof(f1871,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl12),
  inference(cnf_transformation,[],[f1871_D])).
fof(f1871_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1951,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1950])).
fof(f1950,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f1660,f881])).
fof(f1660,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f350,f782])).
fof(f1867,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl8),
  inference(cnf_transformation,[],[f1867_D])).
fof(f1867_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1949,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1948])).
fof(f1948,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f1659,f881])).
fof(f1659,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f347,f782])).
fof(f1918,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f1917])).
fof(f1917,plain,(
  ( ! [X107,X105,X106,X104] : (~sP9('T_0',X105,X106) | sP10('T_0',X105,X106) | s(bool,X104) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(subsumption_resolution,[],[f1603,f881])).
fof(f1603,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,X104)) | ~sP9('T_0',X105,X106) | sP10('T_0',X105,X106) | s(bool,X104) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f441,f782])).
fof(f1915,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f1914])).
fof(f1914,plain,(
  ( ! [X103,X101,X102,X100] : (~sP11('T_0',X101,X102) | sP12('T_0',X101,X102) | s(bool,X100) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(subsumption_resolution,[],[f1602,f881])).
fof(f1602,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,X100)) | ~sP11('T_0',X101,X102) | sP12('T_0',X101,X102) | s(bool,X100) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f440,f782])).
fof(f1912,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f1911])).
fof(f1911,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12('T_0',X98,X97) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f1601,f881])).
fof(f1601,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X96)) | ~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12('T_0',X98,X97) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f437,f782])).
fof(f1909,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f1908])).
fof(f1908,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X93)) | sP12('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f1600,f881])).
fof(f1600,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~p(s(bool,X93)) | sP12('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f436,f782])).
fof(f1906,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f1905])).
fof(f1905,plain,(
  ( ! [X90,X88,X91,X89] : (p(s(bool,X89)) | sP12('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f1599,f881])).
fof(f1599,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | p(s(bool,X89)) | sP12('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f435,f782])).
fof(f1903,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f1902])).
fof(f1902,plain,(
  ( ! [X87,X85,X86,X84] : (p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,'T_0') | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f1598,f881])).
fof(f1598,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,'T_0') | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f430,f782])).
fof(f1900,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f1899])).
fof(f1899,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X81)) | p(s(bool,X82)) | ~sP10('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f1597,f881])).
fof(f1597,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | p(s(bool,X82)) | ~sP10('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f420,f782])).
fof(f1897,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f1896])).
fof(f1896,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X77)) | sP10('T_0',X77,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f1596,f881])).
fof(f1596,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | ~p(s(bool,X77)) | sP10('T_0',X77,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f419,f782])).
fof(f1894,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f1893])).
fof(f1893,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X73)) | sP10('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f1595,f881])).
fof(f1595,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | sP10('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f418,f782])).
fof(f1891,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1890])).
fof(f1890,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f1594,f881])).
fof(f1594,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f413,f782])).
fof(f1888,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f1887])).
fof(f1887,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,'T_0') | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f1593,f881])).
fof(f1593,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,'T_0') | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f405,f782])).
fof(f1885,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f1884])).
fof(f1884,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,'T_0',X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f1592,f881])).
fof(f1592,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,'T_0',X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f396,f782])).
fof(f1882,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1881])).
fof(f1881,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,'T_0',X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f1591,f881])).
fof(f1591,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,'T_0',X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f377,f782])).
fof(f1879,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f1878])).
fof(f1878,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f1590,f881])).
fof(f1590,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f373,f782])).
fof(f1876,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1875])).
fof(f1875,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f1589,f881])).
fof(f1589,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f369,f782])).
fof(f1873,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2(X1,X2,'T_0')) )),
  inference(condensation,[],[f1872])).
fof(f1872,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | sP2(X45,X46,'T_0') | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f1588,f881])).
fof(f1588,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | sP2(X45,X46,'T_0') | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f361,f782])).
fof(f1870,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1869])).
fof(f1869,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f1585,f881])).
fof(f1585,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f350,f782])).
fof(f1866,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1865])).
fof(f1865,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f1584,f881])).
fof(f1584,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f347,f782])).
fof(f663,plain,(
  ( ! [X0] : (s('type.num.num',sK13) = s('type.num.num',X0)) ) | $spl4),
  inference(cnf_transformation,[],[f663_D])).
fof(f663_D,plain,(
  ( ! [X0] : (s('type.num.num',sK13) = s('type.num.num',X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f665,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK13) | $spl7),
  inference(cnf_transformation,[],[f665_D])).
fof(f665_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK13) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f655,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK13) | s('type.num.num',sK13) = s('type.num.num',X0)) ) | $spl2),
  inference(superposition,[],[f462,f264])).
fof(f462,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK13) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num',sK13) = s('type.num.num',X1)) ) | $spl2),
  inference(forward_demodulation,[],[f459,f451])).
fof(f459,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK13) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl2),
  inference(backward_demodulation,[],[f451,f328])).
fof(f449,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl0),
  inference(cnf_transformation,[],[f449_D])).
fof(f449_D,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f451,plain,(
  s('type.num.num',sK13) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl2),
  inference(cnf_transformation,[],[f451_D])).
fof(f451_D,plain,(
  s('type.num.num',sK13) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f255,plain,(
  s('type.num.num',sK13) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  (p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.num.num',sK13) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.num.num',sK13) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f170])).
fof(f170,plain,(
  ? [X0] : ((p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ? [X0] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <~> s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ~! [X0] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> s('type.num.num',X6) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> s('type.num.num',X6) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  file('thm.divides.DIVIDES_ONE','thm.divides.DIVIDES_ONE')).
fof(f450,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl1),
  inference(cnf_transformation,[],[f450_D])).
fof(f450_D,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f452,plain,(
  s('type.num.num',sK13) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl3),
  inference(cnf_transformation,[],[f452_D])).
fof(f452_D,plain,(
  s('type.num.num',sK13) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f256,plain,(
  s('type.num.num',sK13) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(cnf_transformation,[],[f171])).
% SZS output end Proof for thm.divides
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 272
% Passive clauses: 563
% Generated clauses: 3377
% Final active clauses: 220
% Final passive clauses: 187
% Input formulas: 49
% Initial clauses: 185
% 
% Duplicate literals: 61
% Trivial inequalities: 14
% Fw subsumption resolutions: 180
% Fw demodulations: 281
% Bw demodulations: 24
% Condensations: 112
% Global subsumptions: 4
% 
% Simple tautologies: 77
% Equational tautologies: 53
% Forward subsumptions: 1894
% 
% Unit resulting resolution: 1459
% Factoring: 1
% Forward superposition: 434
% Backward superposition: 506
% Self superposition: 34
% Equality factoring: 28
% Equality resolution: 6
% 
% Split clauses: 44
% Split components: 88
% SAT solver clauses: 751
% SAT solver unit clauses: 96
% SAT solver binary clauses: 360
% 
% Sat splits: 44
% Sat splitting refutations: 94
% 
% Memory used [KB]: 2558
% Time elapsed: 0.094 s
% ------------------------------
% ------------------------------
% Success in time 1.413 s
